% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET103-7
% SZS output start Proof for SET103-7
fof(f296,plain,(
  $false),
  inference(subsumption_resolution,[],[f295,f108])).
fof(f108,axiom,(
  ( ! [X0,X1] : (unordered_pair(X0,X1) = unordered_pair(X1,X0)) )),
  file('SET103-7.tptp',unknown)).
fof(f295,plain,(
  unordered_pair(unordered_pair(null_class,x),unordered_pair(x,x)) != unordered_pair(unordered_pair(x,x),unordered_pair(null_class,x))),
  inference(forward_demodulation,[],[f294,f108])).
fof(f294,plain,(
  unordered_pair(unordered_pair(x,x),unordered_pair(x,null_class)) != unordered_pair(unordered_pair(null_class,x),unordered_pair(x,x))),
  inference(backward_demodulation,[],[f292,f240])).
fof(f240,plain,(
  unordered_pair(unordered_pair(x,x),unordered_pair(x,unordered_pair(y,y))) != unordered_pair(unordered_pair(null_class,x),unordered_pair(x,x))),
  inference(forward_demodulation,[],[f239,f108])).
fof(f239,plain,(
  unordered_pair(unordered_pair(x,x),unordered_pair(x,unordered_pair(y,y))) != unordered_pair(unordered_pair(x,null_class),unordered_pair(x,x))),
  inference(forward_demodulation,[],[f235,f108])).
fof(f235,plain,(
  unordered_pair(unordered_pair(x,x),unordered_pair(x,null_class)) != unordered_pair(unordered_pair(x,x),unordered_pair(x,unordered_pair(y,y)))),
  inference(definition_unfolding,[],[f148,f12,f154])).
fof(f154,plain,(
  ( ! [X0,X1] : (ordered_pair(X0,X1) = unordered_pair(unordered_pair(X0,X0),unordered_pair(X0,unordered_pair(X1,X1)))) )),
  inference(definition_unfolding,[],[f13,f12,f12])).
fof(f13,axiom,(
  ( ! [X0,X1] : (unordered_pair(singleton(X0),unordered_pair(X0,singleton(X1))) = ordered_pair(X0,X1)) )),
  file('SET103-7.tptp',unknown)).
fof(f12,axiom,(
  ( ! [X0] : (unordered_pair(X0,X0) = singleton(X0)) )),
  file('SET103-7.tptp',unknown)).
fof(f148,axiom,(
  unordered_pair(singleton(x),unordered_pair(x,null_class)) != ordered_pair(x,y)),
  file('SET103-7.tptp',unknown)).
fof(f292,plain,(
  null_class = unordered_pair(y,y)),
  inference(resolution,[],[f216,f149])).
fof(f149,axiom,(
  ~member(y,universal_class)),
  file('SET103-7.tptp',unknown)).
fof(f216,plain,(
  ( ! [X0] : (member(X0,universal_class) | unordered_pair(X0,X0) = null_class) )),
  inference(definition_unfolding,[],[f126,f12])).
fof(f126,axiom,(
  ( ! [X0] : (singleton(X0) = null_class | member(X0,universal_class)) )),
  file('SET103-7.tptp',unknown)).
% SZS output end Proof for SET103-7
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
