% Refutation found. Thanks to Tanya!
% SZS status Theorem for AGT025+1
% SZS output start Proof for AGT025+1
fof(f12055748,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f30607,f1423556,f3044823,f4765784,f4810947,f7810543,f12055745,f12055747])).
fof(f12055747,plain,(
  spl0_644 | ~spl0_24 | ~spl0_732),
  inference(avatar_split_clause,[],[f12055746,f202540,f2547,f150689])).
fof(f150689,plain,(
  spl0_644 <=> ! [X10] : (accept_city(X10,towna) | ~accept_city(X10,citya))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_644])])).
fof(f2547,plain,(
  spl0_24 <=> less(n1,n25)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_24])])).
fof(f202540,plain,(
  spl0_732 <=> less(n24,n25)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_732])])).
fof(f12055746,plain,(
  ( ! [X11] : (accept_city(X11,towna) | ~accept_city(X11,citya)) ) | (~spl0_24 | ~spl0_732)),
  inference(subsumption_resolution,[],[f12055728,f202541])).
fof(f202541,plain,(
  less(n24,n25) | ~spl0_732),
  inference(avatar_component_clause,[],[f202540])).
fof(f12055728,plain,(
  ( ! [X11] : (accept_city(X11,towna) | ~accept_city(X11,citya) | ~less(n24,n25)) ) | ~spl0_24),
  inference(resolution,[],[f12055715,f1700])).
fof(f1700,plain,(
  ( ! [X2,X0,X1] : (accept_population(X1,X2,X0) | ~less(X0,n25)) )),
  inference(resolution,[],[f639,f1257])).
fof(f1257,plain,(
  ( ! [X0,X1] : (accept_population(X0,X1,n25)) )),
  inference(cnf_transformation,[],[f603])).
fof(f603,plain,(
  ! [X0,X1] : accept_population(X0,X1,n25)),
  inference(ennf_transformation,[],[f576])).
fof(f576,plain,(
  ~? [X0,X1] : ~accept_population(X0,X1,n25)),
  inference(rectify,[],[f557])).
fof(f557,negated_conjecture,(
  ~? [X6,X7] : ~accept_population(X6,X7,n25)),
  inference(negated_conjecture,[],[f556])).
fof(f556,conjecture,(
  ? [X6,X7] : ~accept_population(X6,X7,n25)),
  file('AGT025+1.tptp',unknown)).
fof(f639,plain,(
  ( ! [X2,X0,X3,X1] : (~accept_population(X0,X3,X1) | ~less(X2,X1) | accept_population(X0,X3,X2)) )),
  inference(cnf_transformation,[],[f584])).
fof(f584,plain,(
  ! [X0,X1,X2,X3] : (accept_population(X0,X3,X2) | ~less(X2,X1) | ~accept_population(X0,X3,X1))),
  inference(flattening,[],[f583])).
fof(f583,plain,(
  ! [X0,X1,X2,X3] : (accept_population(X0,X3,X2) | (~less(X2,X1) | ~accept_population(X0,X3,X1)))),
  inference(ennf_transformation,[],[f559])).
fof(f559,plain,(
  ! [X0,X1,X2,X3] : ((less(X2,X1) & accept_population(X0,X3,X1)) => accept_population(X0,X3,X2))),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X0,X2,X4,X5] : ((less(X4,X2) & accept_population(X0,X5,X2)) => accept_population(X0,X5,X4))),
  file('AGT025+1.tptp',unknown)).
fof(f12055715,plain,(
  ( ! [X4] : (~accept_population(X4,christian,n24) | accept_city(X4,towna) | ~accept_city(X4,citya)) ) | ~spl0_24),
  inference(subsumption_resolution,[],[f30280,f2548])).
fof(f2548,plain,(
  less(n1,n25) | ~spl0_24),
  inference(avatar_component_clause,[],[f2547])).
fof(f30280,plain,(
  ( ! [X4] : (~less(n1,n25) | accept_city(X4,towna) | ~accept_population(X4,christian,n24) | ~accept_city(X4,citya)) )),
  inference(resolution,[],[f1700,f2726])).
fof(f2726,plain,(
  ( ! [X0] : (~accept_population(X0,muslim,n1) | accept_city(X0,towna) | ~accept_population(X0,christian,n24) | ~accept_city(X0,citya)) )),
  inference(subsumption_resolution,[],[f2725,f699])).
fof(f699,plain,(
  ( ! [X0] : (accept_population(X0,atheist,n75) | ~accept_city(X0,citya)) )),
  inference(cnf_transformation,[],[f625])).
fof(f625,plain,(
  ! [X0] : (((accept_population(X0,other,n0) & accept_population(X0,native,n0) & accept_population(X0,muslim,n0) & accept_population(X0,christian,n25) & accept_population(X0,atheist,n75)) | ~accept_city(X0,citya)) & (accept_city(X0,citya) | ~accept_population(X0,other,n0) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n0) | ~accept_population(X0,christian,n25) | ~accept_population(X0,atheist,n75)))),
  inference(flattening,[],[f624])).
fof(f624,plain,(
  ! [X0] : (((accept_population(X0,other,n0) & accept_population(X0,native,n0) & accept_population(X0,muslim,n0) & accept_population(X0,christian,n25) & accept_population(X0,atheist,n75)) | ~accept_city(X0,citya)) & (accept_city(X0,citya) | (~accept_population(X0,other,n0) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n0) | ~accept_population(X0,christian,n25) | ~accept_population(X0,atheist,n75))))),
  inference(nnf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0] : ((accept_population(X0,other,n0) & accept_population(X0,native,n0) & accept_population(X0,muslim,n0) & accept_population(X0,christian,n25) & accept_population(X0,atheist,n75)) <=> accept_city(X0,citya))),
  file('AGT025+1.tptp',unknown)).
fof(f2725,plain,(
  ( ! [X0] : (accept_city(X0,towna) | ~accept_population(X0,muslim,n1) | ~accept_population(X0,christian,n24) | ~accept_population(X0,atheist,n75) | ~accept_city(X0,citya)) )),
  inference(subsumption_resolution,[],[f2716,f702])).
fof(f702,plain,(
  ( ! [X0] : (accept_population(X0,native,n0) | ~accept_city(X0,citya)) )),
  inference(cnf_transformation,[],[f625])).
fof(f2716,plain,(
  ( ! [X0] : (accept_city(X0,towna) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n1) | ~accept_population(X0,christian,n24) | ~accept_population(X0,atheist,n75) | ~accept_city(X0,citya)) )),
  inference(resolution,[],[f692,f703])).
fof(f703,plain,(
  ( ! [X0] : (accept_population(X0,other,n0) | ~accept_city(X0,citya)) )),
  inference(cnf_transformation,[],[f625])).
fof(f692,plain,(
  ( ! [X0] : (~accept_population(X0,other,n0) | accept_city(X0,towna) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n1) | ~accept_population(X0,christian,n24) | ~accept_population(X0,atheist,n75)) )),
  inference(cnf_transformation,[],[f623])).
fof(f623,plain,(
  ! [X0] : (((accept_population(X0,other,n0) & accept_population(X0,native,n0) & accept_population(X0,muslim,n1) & accept_population(X0,christian,n24) & accept_population(X0,atheist,n75)) | ~accept_city(X0,towna)) & (accept_city(X0,towna) | ~accept_population(X0,other,n0) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n1) | ~accept_population(X0,christian,n24) | ~accept_population(X0,atheist,n75)))),
  inference(flattening,[],[f622])).
fof(f622,plain,(
  ! [X0] : (((accept_population(X0,other,n0) & accept_population(X0,native,n0) & accept_population(X0,muslim,n1) & accept_population(X0,christian,n24) & accept_population(X0,atheist,n75)) | ~accept_city(X0,towna)) & (accept_city(X0,towna) | (~accept_population(X0,other,n0) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n1) | ~accept_population(X0,christian,n24) | ~accept_population(X0,atheist,n75))))),
  inference(nnf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0] : ((accept_population(X0,other,n0) & accept_population(X0,native,n0) & accept_population(X0,muslim,n1) & accept_population(X0,christian,n24) & accept_population(X0,atheist,n75)) <=> accept_city(X0,towna))),
  file('AGT025+1.tptp',unknown)).
fof(f12055745,plain,(
  ~spl0_407 | ~spl0_644),
  inference(avatar_contradiction_clause,[],[f12055744])).
fof(f12055744,plain,(
  $false | (~spl0_407 | ~spl0_644)),
  inference(subsumption_resolution,[],[f12055738,f1698])).
fof(f1698,plain,(
  ~accept_city(muslimcountrybhumanitarianorganization,towna)),
  inference(subsumption_resolution,[],[f1695,f1337])).
fof(f1337,plain,(
  accept_leader(muslimcountrybhumanitarianorganization,countryacivilorganization)),
  inference(resolution,[],[f642,f733])).
fof(f733,plain,(
  any_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,countryacivilorganization,towna)),
  inference(cnf_transformation,[],[f32])).
fof(f32,axiom,(
  any_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,countryacivilorganization,towna)),
  file('AGT025+1.tptp',unknown)).
fof(f642,plain,(
  ( ! [X2,X0,X1] : (~any_agent_in_all_proposed_teams(X0,X1,X2) | accept_leader(X0,X1)) )),
  inference(cnf_transformation,[],[f586])).
fof(f586,plain,(
  ! [X0,X1,X2] : (accept_leader(X0,X1) | ~any_agent_in_all_proposed_teams(X0,X1,X2))),
  inference(ennf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0,X1,X2] : (any_agent_in_all_proposed_teams(X0,X1,X2) => accept_leader(X0,X1))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X0,X3,X1] : (any_agent_in_all_proposed_teams(X0,X3,X1) => accept_leader(X0,X3))),
  file('AGT025+1.tptp',unknown)).
fof(f1695,plain,(
  ~accept_city(muslimcountrybhumanitarianorganization,towna) | ~accept_leader(muslimcountrybhumanitarianorganization,countryacivilorganization)),
  inference(resolution,[],[f643,f732])).
fof(f732,plain,(
  the_agent_not_in_any_proposed_teams(muslimcountrybhumanitarianorganization,countryacivilorganization,towna)),
  inference(cnf_transformation,[],[f31])).
fof(f31,axiom,(
  the_agent_not_in_any_proposed_teams(muslimcountrybhumanitarianorganization,countryacivilorganization,towna)),
  file('AGT025+1.tptp',unknown)).
fof(f643,plain,(
  ( ! [X2,X0,X1] : (~the_agent_not_in_any_proposed_teams(X0,X1,X2) | ~accept_city(X0,X2) | ~accept_leader(X0,X1)) )),
  inference(cnf_transformation,[],[f588])).
fof(f588,plain,(
  ! [X0,X1,X2] : (~accept_leader(X0,X1) | ~accept_city(X0,X2) | ~the_agent_not_in_any_proposed_teams(X0,X1,X2))),
  inference(flattening,[],[f587])).
fof(f587,plain,(
  ! [X0,X1,X2] : ((~accept_leader(X0,X1) | ~accept_city(X0,X2)) | ~the_agent_not_in_any_proposed_teams(X0,X1,X2))),
  inference(ennf_transformation,[],[f562])).
fof(f562,plain,(
  ! [X0,X1,X2] : (the_agent_not_in_any_proposed_teams(X0,X1,X2) => ~(accept_leader(X0,X1) & accept_city(X0,X2)))),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X0,X3,X1] : (the_agent_not_in_any_proposed_teams(X0,X3,X1) => ~(accept_leader(X0,X3) & accept_city(X0,X1)))),
  file('AGT025+1.tptp',unknown)).
fof(f12055738,plain,(
  accept_city(muslimcountrybhumanitarianorganization,towna) | (~spl0_407 | ~spl0_644)),
  inference(resolution,[],[f150690,f11953962])).
fof(f11953962,plain,(
  accept_city(muslimcountrybhumanitarianorganization,citya) | ~spl0_407),
  inference(resolution,[],[f76840,f1408])).
fof(f1408,plain,(
  accept_city(muslimcountrybhumanitarianorganization,cityb)),
  inference(resolution,[],[f634,f729])).
fof(f729,plain,(
  accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,cityb,n4)),
  inference(cnf_transformation,[],[f28])).
fof(f28,axiom,(
  accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,cityb,n4)),
  file('AGT025+1.tptp',unknown)).
fof(f634,plain,(
  ( ! [X2,X0,X3,X1] : (~accept_team(X0,X3,X1,X2) | accept_city(X0,X1)) )),
  inference(cnf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2,X3] : ((accept_team(X0,X3,X1,X2) | ~accept_number(X0,X2) | ~accept_leader(X0,X3) | ~accept_city(X0,X1)) & ((accept_number(X0,X2) & accept_leader(X0,X3) & accept_city(X0,X1)) | ~accept_team(X0,X3,X1,X2)))),
  inference(flattening,[],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2,X3] : ((accept_team(X0,X3,X1,X2) | (~accept_number(X0,X2) | ~accept_leader(X0,X3) | ~accept_city(X0,X1))) & ((accept_number(X0,X2) & accept_leader(X0,X3) & accept_city(X0,X1)) | ~accept_team(X0,X3,X1,X2)))),
  inference(nnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2,X3] : (accept_team(X0,X3,X1,X2) <=> (accept_number(X0,X2) & accept_leader(X0,X3) & accept_city(X0,X1)))),
  file('AGT025+1.tptp',unknown)).
fof(f76840,plain,(
  ( ! [X14] : (~accept_city(X14,cityb) | accept_city(X14,citya)) ) | ~spl0_407),
  inference(avatar_component_clause,[],[f76839])).
fof(f76839,plain,(
  spl0_407 <=> ! [X14] : (accept_city(X14,citya) | ~accept_city(X14,cityb))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_407])])).
fof(f150690,plain,(
  ( ! [X10] : (~accept_city(X10,citya) | accept_city(X10,towna)) ) | ~spl0_644),
  inference(avatar_component_clause,[],[f150689])).
fof(f7810543,plain,(
  spl0_44),
  inference(avatar_split_clause,[],[f170784,f2683])).
fof(f2683,plain,(
  spl0_44 <=> less(n0,n25)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_44])])).
fof(f170784,plain,(
  less(n0,n25)),
  inference(resolution,[],[f28665,f971])).
fof(f971,plain,(
  rdn_translate(n0,rdn_pos(rdnn(n0)))),
  inference(cnf_transformation,[],[f270])).
fof(f270,axiom,(
  rdn_translate(n0,rdn_pos(rdnn(n0)))),
  file('AGT025+1.tptp',unknown)).
fof(f28665,plain,(
  ( ! [X0,X1] : (~rdn_translate(X0,rdn_pos(rdnn(X1))) | less(X0,n25)) )),
  inference(subsumption_resolution,[],[f28662,f1265])).
fof(f1265,plain,(
  rdn_non_zero(rdnn(n2))),
  inference(resolution,[],[f1249,f1228])).
fof(f1228,plain,(
  rdn_non_zero_digit(rdnn(n2))),
  inference(cnf_transformation,[],[f527])).
fof(f527,axiom,(
  rdn_non_zero_digit(rdnn(n2))),
  file('AGT025+1.tptp',unknown)).
fof(f1249,plain,(
  ( ! [X0] : (~rdn_non_zero_digit(rdnn(X0)) | rdn_non_zero(rdnn(X0))) )),
  inference(cnf_transformation,[],[f595])).
fof(f595,plain,(
  ! [X0] : (rdn_non_zero(rdnn(X0)) | ~rdn_non_zero_digit(rdnn(X0)))),
  inference(ennf_transformation,[],[f568])).
fof(f568,plain,(
  ! [X0] : (rdn_non_zero_digit(rdnn(X0)) => rdn_non_zero(rdnn(X0)))),
  inference(rectify,[],[f548])).
fof(f548,axiom,(
  ! [X6] : (rdn_non_zero_digit(rdnn(X6)) => rdn_non_zero(rdnn(X6)))),
  file('AGT025+1.tptp',unknown)).
fof(f28662,plain,(
  ( ! [X0,X1] : (less(X0,n25) | ~rdn_translate(X0,rdn_pos(rdnn(X1))) | ~rdn_non_zero(rdnn(n2))) )),
  inference(resolution,[],[f2246,f1248])).
fof(f1248,plain,(
  ( ! [X2,X0,X1] : (rdn_positive_less(rdnn(X0),rdn(rdnn(X1),X2)) | ~rdn_non_zero(X2)) )),
  inference(cnf_transformation,[],[f594])).
fof(f594,plain,(
  ! [X0,X1,X2] : (rdn_positive_less(rdnn(X0),rdn(rdnn(X1),X2)) | ~rdn_non_zero(X2))),
  inference(ennf_transformation,[],[f567])).
fof(f567,plain,(
  ! [X0,X1,X2] : (rdn_non_zero(X2) => rdn_positive_less(rdnn(X0),rdn(rdnn(X1),X2)))),
  inference(rectify,[],[f547])).
fof(f547,axiom,(
  ! [X14,X11,X12] : (rdn_non_zero(X12) => rdn_positive_less(rdnn(X14),rdn(rdnn(X11),X12)))),
  file('AGT025+1.tptp',unknown)).
fof(f2246,plain,(
  ( ! [X146,X147] : (~rdn_positive_less(X146,rdn(rdnn(n5),rdnn(n2))) | less(X147,n25) | ~rdn_translate(X147,rdn_pos(X146))) )),
  inference(resolution,[],[f1251,f996])).
fof(f996,plain,(
  rdn_translate(n25,rdn_pos(rdn(rdnn(n5),rdnn(n2))))),
  inference(cnf_transformation,[],[f295])).
fof(f295,axiom,(
  rdn_translate(n25,rdn_pos(rdn(rdnn(n5),rdnn(n2))))),
  file('AGT025+1.tptp',unknown)).
fof(f1251,plain,(
  ( ! [X2,X0,X3,X1] : (~rdn_translate(X1,rdn_pos(X3)) | ~rdn_positive_less(X2,X3) | less(X0,X1) | ~rdn_translate(X0,rdn_pos(X2))) )),
  inference(cnf_transformation,[],[f598])).
fof(f598,plain,(
  ! [X0,X1,X2,X3] : (less(X0,X1) | ~rdn_positive_less(X2,X3) | ~rdn_translate(X1,rdn_pos(X3)) | ~rdn_translate(X0,rdn_pos(X2)))),
  inference(flattening,[],[f597])).
fof(f597,plain,(
  ! [X0,X1,X2,X3] : (less(X0,X1) | (~rdn_positive_less(X2,X3) | ~rdn_translate(X1,rdn_pos(X3)) | ~rdn_translate(X0,rdn_pos(X2))))),
  inference(ennf_transformation,[],[f570])).
fof(f570,plain,(
  ! [X0,X1,X2,X3] : ((rdn_positive_less(X2,X3) & rdn_translate(X1,rdn_pos(X3)) & rdn_translate(X0,rdn_pos(X2))) => less(X0,X1))),
  inference(rectify,[],[f550])).
fof(f550,axiom,(
  ! [X6,X7,X15,X16] : ((rdn_positive_less(X15,X16) & rdn_translate(X7,rdn_pos(X16)) & rdn_translate(X6,rdn_pos(X15))) => less(X6,X7))),
  file('AGT025+1.tptp',unknown)).
fof(f4810947,plain,(
  spl0_732),
  inference(avatar_split_clause,[],[f4810946,f202540])).
fof(f4810946,plain,(
  less(n24,n25)),
  inference(subsumption_resolution,[],[f4810930,f1240])).
fof(f1240,plain,(
  rdn_positive_less(rdnn(n4),rdnn(n5))),
  inference(cnf_transformation,[],[f539])).
fof(f539,axiom,(
  rdn_positive_less(rdnn(n4),rdnn(n5))),
  file('AGT025+1.tptp',unknown)).
fof(f4810930,plain,(
  less(n24,n25) | ~rdn_positive_less(rdnn(n4),rdnn(n5))),
  inference(resolution,[],[f28666,f995])).
fof(f995,plain,(
  rdn_translate(n24,rdn_pos(rdn(rdnn(n4),rdnn(n2))))),
  inference(cnf_transformation,[],[f294])).
fof(f294,axiom,(
  rdn_translate(n24,rdn_pos(rdn(rdnn(n4),rdnn(n2))))),
  file('AGT025+1.tptp',unknown)).
fof(f28666,plain,(
  ( ! [X6,X5] : (~rdn_translate(X5,rdn_pos(rdn(rdnn(X6),rdnn(n2)))) | less(X5,n25) | ~rdn_positive_less(rdnn(X6),rdnn(n5))) )),
  inference(subsumption_resolution,[],[f28664,f1265])).
fof(f28664,plain,(
  ( ! [X6,X5] : (less(X5,n25) | ~rdn_translate(X5,rdn_pos(rdn(rdnn(X6),rdnn(n2)))) | ~rdn_non_zero(rdnn(n2)) | ~rdn_positive_less(rdnn(X6),rdnn(n5))) )),
  inference(resolution,[],[f2246,f1247])).
fof(f1247,plain,(
  ( ! [X2,X0,X1] : (rdn_positive_less(rdn(rdnn(X0),X1),rdn(rdnn(X2),X1)) | ~rdn_non_zero(X1) | ~rdn_positive_less(rdnn(X0),rdnn(X2))) )),
  inference(cnf_transformation,[],[f593])).
fof(f593,plain,(
  ! [X0,X1,X2] : (rdn_positive_less(rdn(rdnn(X0),X1),rdn(rdnn(X2),X1)) | ~rdn_non_zero(X1) | ~rdn_positive_less(rdnn(X0),rdnn(X2)))),
  inference(flattening,[],[f592])).
fof(f592,plain,(
  ! [X0,X1,X2] : (rdn_positive_less(rdn(rdnn(X0),X1),rdn(rdnn(X2),X1)) | (~rdn_non_zero(X1) | ~rdn_positive_less(rdnn(X0),rdnn(X2))))),
  inference(ennf_transformation,[],[f566])).
fof(f566,plain,(
  ! [X0,X1,X2] : ((rdn_non_zero(X1) & rdn_positive_less(rdnn(X0),rdnn(X2))) => rdn_positive_less(rdn(rdnn(X0),X1),rdn(rdnn(X2),X1)))),
  inference(rectify,[],[f546])).
fof(f546,axiom,(
  ! [X9,X13,X11] : ((rdn_non_zero(X13) & rdn_positive_less(rdnn(X9),rdnn(X11))) => rdn_positive_less(rdn(rdnn(X9),X13),rdn(rdnn(X11),X13)))),
  file('AGT025+1.tptp',unknown)).
fof(f4765784,plain,(
  spl0_406),
  inference(avatar_split_clause,[],[f4765783,f76835])).
fof(f76835,plain,(
  spl0_406 <=> less(n75,n78)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_406])])).
fof(f4765783,plain,(
  less(n75,n78)),
  inference(subsumption_resolution,[],[f4765766,f73646])).
fof(f73646,plain,(
  rdn_positive_less(rdnn(n5),rdnn(n8))),
  inference(resolution,[],[f28032,f2429])).
fof(f2429,plain,(
  ( ! [X0] : (~rdn_positive_less(rdnn(X0),rdnn(n7)) | rdn_positive_less(rdnn(X0),rdnn(n8))) )),
  inference(resolution,[],[f1245,f1243])).
fof(f1243,plain,(
  rdn_positive_less(rdnn(n7),rdnn(n8))),
  inference(cnf_transformation,[],[f542])).
fof(f542,axiom,(
  rdn_positive_less(rdnn(n7),rdnn(n8))),
  file('AGT025+1.tptp',unknown)).
fof(f1245,plain,(
  ( ! [X2,X0,X1] : (~rdn_positive_less(rdnn(X1),rdnn(X2)) | rdn_positive_less(rdnn(X0),rdnn(X2)) | ~rdn_positive_less(rdnn(X0),rdnn(X1))) )),
  inference(cnf_transformation,[],[f590])).
fof(f590,plain,(
  ! [X0,X1,X2] : (rdn_positive_less(rdnn(X0),rdnn(X2)) | ~rdn_positive_less(rdnn(X1),rdnn(X2)) | ~rdn_positive_less(rdnn(X0),rdnn(X1)))),
  inference(flattening,[],[f589])).
fof(f589,plain,(
  ! [X0,X1,X2] : (rdn_positive_less(rdnn(X0),rdnn(X2)) | (~rdn_positive_less(rdnn(X1),rdnn(X2)) | ~rdn_positive_less(rdnn(X0),rdnn(X1))))),
  inference(ennf_transformation,[],[f564])).
fof(f564,plain,(
  ! [X0,X1,X2] : ((rdn_positive_less(rdnn(X1),rdnn(X2)) & rdn_positive_less(rdnn(X0),rdnn(X1))) => rdn_positive_less(rdnn(X0),rdnn(X2)))),
  inference(rectify,[],[f544])).
fof(f544,axiom,(
  ! [X6,X7,X8] : ((rdn_positive_less(rdnn(X7),rdnn(X8)) & rdn_positive_less(rdnn(X6),rdnn(X7))) => rdn_positive_less(rdnn(X6),rdnn(X8)))),
  file('AGT025+1.tptp',unknown)).
fof(f28032,plain,(
  rdn_positive_less(rdnn(n5),rdnn(n7))),
  inference(resolution,[],[f2437,f1241])).
fof(f1241,plain,(
  rdn_positive_less(rdnn(n5),rdnn(n6))),
  inference(cnf_transformation,[],[f540])).
fof(f540,axiom,(
  rdn_positive_less(rdnn(n5),rdnn(n6))),
  file('AGT025+1.tptp',unknown)).
fof(f2437,plain,(
  ( ! [X8] : (~rdn_positive_less(rdnn(X8),rdnn(n6)) | rdn_positive_less(rdnn(X8),rdnn(n7))) )),
  inference(resolution,[],[f1245,f1242])).
fof(f1242,plain,(
  rdn_positive_less(rdnn(n6),rdnn(n7))),
  inference(cnf_transformation,[],[f541])).
fof(f541,axiom,(
  rdn_positive_less(rdnn(n6),rdnn(n7))),
  file('AGT025+1.tptp',unknown)).
fof(f4765766,plain,(
  less(n75,n78) | ~rdn_positive_less(rdnn(n5),rdnn(n8))),
  inference(resolution,[],[f28281,f1046])).
fof(f1046,plain,(
  rdn_translate(n75,rdn_pos(rdn(rdnn(n5),rdnn(n7))))),
  inference(cnf_transformation,[],[f345])).
fof(f345,axiom,(
  rdn_translate(n75,rdn_pos(rdn(rdnn(n5),rdnn(n7))))),
  file('AGT025+1.tptp',unknown)).
fof(f28281,plain,(
  ( ! [X6,X5] : (~rdn_translate(X5,rdn_pos(rdn(rdnn(X6),rdnn(n7)))) | less(X5,n78) | ~rdn_positive_less(rdnn(X6),rdnn(n8))) )),
  inference(subsumption_resolution,[],[f28279,f1259])).
fof(f1259,plain,(
  rdn_non_zero(rdnn(n7))),
  inference(resolution,[],[f1249,f1233])).
fof(f1233,plain,(
  rdn_non_zero_digit(rdnn(n7))),
  inference(cnf_transformation,[],[f532])).
fof(f532,axiom,(
  rdn_non_zero_digit(rdnn(n7))),
  file('AGT025+1.tptp',unknown)).
fof(f28279,plain,(
  ( ! [X6,X5] : (less(X5,n78) | ~rdn_translate(X5,rdn_pos(rdn(rdnn(X6),rdnn(n7)))) | ~rdn_non_zero(rdnn(n7)) | ~rdn_positive_less(rdnn(X6),rdnn(n8))) )),
  inference(resolution,[],[f2185,f1247])).
fof(f2185,plain,(
  ( ! [X24,X25] : (~rdn_positive_less(X24,rdn(rdnn(n8),rdnn(n7))) | less(X25,n78) | ~rdn_translate(X25,rdn_pos(X24))) )),
  inference(resolution,[],[f1251,f1049])).
fof(f1049,plain,(
  rdn_translate(n78,rdn_pos(rdn(rdnn(n8),rdnn(n7))))),
  inference(cnf_transformation,[],[f348])).
fof(f348,axiom,(
  rdn_translate(n78,rdn_pos(rdn(rdnn(n8),rdnn(n7))))),
  file('AGT025+1.tptp',unknown)).
fof(f3044823,plain,(
  ~spl0_406 | spl0_407 | ~spl0_178),
  inference(avatar_split_clause,[],[f1398151,f30501,f76839,f76835])).
fof(f30501,plain,(
  spl0_178 <=> ! [X15] : (accept_city(X15,citya) | ~accept_population(X15,atheist,n75))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_178])])).
fof(f1398151,plain,(
  ( ! [X13] : (accept_city(X13,citya) | ~less(n75,n78) | ~accept_city(X13,cityb)) ) | ~spl0_178),
  inference(resolution,[],[f30502,f1713])).
fof(f1713,plain,(
  ( ! [X28,X27] : (accept_population(X28,atheist,X27) | ~less(X27,n78) | ~accept_city(X28,cityb)) )),
  inference(resolution,[],[f639,f711])).
fof(f711,plain,(
  ( ! [X0] : (accept_population(X0,atheist,n78) | ~accept_city(X0,cityb)) )),
  inference(cnf_transformation,[],[f629])).
fof(f629,plain,(
  ! [X0] : (((accept_population(X0,other,n1) & accept_population(X0,native,n0) & accept_population(X0,muslim,n1) & accept_population(X0,christian,n20) & accept_population(X0,atheist,n78)) | ~accept_city(X0,cityb)) & (accept_city(X0,cityb) | ~accept_population(X0,other,n1) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n1) | ~accept_population(X0,christian,n20) | ~accept_population(X0,atheist,n78)))),
  inference(flattening,[],[f628])).
fof(f628,plain,(
  ! [X0] : (((accept_population(X0,other,n1) & accept_population(X0,native,n0) & accept_population(X0,muslim,n1) & accept_population(X0,christian,n20) & accept_population(X0,atheist,n78)) | ~accept_city(X0,cityb)) & (accept_city(X0,cityb) | (~accept_population(X0,other,n1) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n1) | ~accept_population(X0,christian,n20) | ~accept_population(X0,atheist,n78))))),
  inference(nnf_transformation,[],[f19])).
fof(f19,axiom,(
  ! [X0] : ((accept_population(X0,other,n1) & accept_population(X0,native,n0) & accept_population(X0,muslim,n1) & accept_population(X0,christian,n20) & accept_population(X0,atheist,n78)) <=> accept_city(X0,cityb))),
  file('AGT025+1.tptp',unknown)).
fof(f30502,plain,(
  ( ! [X15] : (~accept_population(X15,atheist,n75) | accept_city(X15,citya)) ) | ~spl0_178),
  inference(avatar_component_clause,[],[f30501])).
fof(f1423556,plain,(
  spl0_24),
  inference(avatar_contradiction_clause,[],[f1423555])).
fof(f1423555,plain,(
  $false | spl0_24),
  inference(subsumption_resolution,[],[f170782,f2549])).
fof(f2549,plain,(
  ~less(n1,n25) | spl0_24),
  inference(avatar_component_clause,[],[f2547])).
fof(f170782,plain,(
  less(n1,n25)),
  inference(resolution,[],[f28665,f972])).
fof(f972,plain,(
  rdn_translate(n1,rdn_pos(rdnn(n1)))),
  inference(cnf_transformation,[],[f271])).
fof(f271,axiom,(
  rdn_translate(n1,rdn_pos(rdnn(n1)))),
  file('AGT025+1.tptp',unknown)).
fof(f30607,plain,(
  spl0_178 | ~spl0_44),
  inference(avatar_split_clause,[],[f20207,f2683,f30501])).
fof(f20207,plain,(
  ( ! [X10] : (~less(n0,n25) | accept_city(X10,citya) | ~accept_population(X10,atheist,n75)) )),
  inference(subsumption_resolution,[],[f20206,f1700])).
fof(f20206,plain,(
  ( ! [X10] : (~less(n0,n25) | accept_city(X10,citya) | ~accept_population(X10,muslim,n0) | ~accept_population(X10,atheist,n75)) )),
  inference(subsumption_resolution,[],[f2862,f1700])).
fof(f2862,plain,(
  ( ! [X10] : (~less(n0,n25) | accept_city(X10,citya) | ~accept_population(X10,native,n0) | ~accept_population(X10,muslim,n0) | ~accept_population(X10,atheist,n75)) )),
  inference(resolution,[],[f1700,f2782])).
fof(f2782,plain,(
  ( ! [X0] : (~accept_population(X0,other,n0) | accept_city(X0,citya) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n0) | ~accept_population(X0,atheist,n75)) )),
  inference(subsumption_resolution,[],[f698,f1257])).
fof(f698,plain,(
  ( ! [X0] : (accept_city(X0,citya) | ~accept_population(X0,other,n0) | ~accept_population(X0,native,n0) | ~accept_population(X0,muslim,n0) | ~accept_population(X0,christian,n25) | ~accept_population(X0,atheist,n75)) )),
  inference(cnf_transformation,[],[f625])).
% SZS output end Proof for AGT025+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 2464053
% Time elapsed: 45.397 s
% ------------------------------
% ------------------------------
