% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for CAT006-4
% SZS output start Proof for CAT006-4
fof(f56,plain,(
  $false),
  inference(subsumption_resolution,[],[f55,f15])).
fof(f15,axiom,(
  d != codomain(a)),
  file('CAT006-4.tptp',unknown)).
fof(f55,plain,(
  d = codomain(a)),
  inference(forward_demodulation,[],[f54,f30])).
fof(f30,plain,(
  d = domain(d)),
  inference(forward_demodulation,[],[f29,f10])).
fof(f10,axiom,(
  ( ! [X0] : (compose(X0,domain(X0)) = X0) )),
  file('CAT006-4.tptp',unknown)).
fof(f29,plain,(
  domain(d) = compose(d,domain(d))),
  inference(subsumption_resolution,[],[f27,f22])).
fof(f22,plain,(
  there_exists(d)),
  inference(resolution,[],[f19,f4])).
fof(f4,axiom,(
  ( ! [X0] : (~there_exists(domain(X0)) | there_exists(X0)) )),
  file('CAT006-4.tptp',unknown)).
fof(f19,plain,(
  there_exists(domain(d))),
  inference(resolution,[],[f6,f12])).
fof(f12,axiom,(
  there_exists(compose(d,a))),
  file('CAT006-4.tptp',unknown)).
fof(f6,axiom,(
  ( ! [X0,X1] : (~there_exists(compose(X0,X1)) | there_exists(domain(X0))) )),
  file('CAT006-4.tptp',unknown)).
fof(f27,plain,(
  ~there_exists(d) | domain(d) = compose(d,domain(d))),
  inference(superposition,[],[f14,f10])).
fof(f14,axiom,(
  ( ! [X0] : (~there_exists(compose(d,X0)) | compose(d,X0) = X0) )),
  file('CAT006-4.tptp',unknown)).
fof(f54,plain,(
  codomain(a) = domain(d)),
  inference(subsumption_resolution,[],[f52,f28])).
fof(f28,plain,(
  there_exists(a)),
  inference(backward_demodulation,[],[f26,f12])).
fof(f26,plain,(
  a = compose(d,a)),
  inference(resolution,[],[f14,f12])).
fof(f52,plain,(
  ~there_exists(a) | codomain(a) = domain(d)),
  inference(superposition,[],[f7,f26])).
fof(f7,axiom,(
  ( ! [X0,X1] : (~there_exists(compose(X0,X1)) | domain(X0) = codomain(X1)) )),
  file('CAT006-4.tptp',unknown)).
% SZS output end Proof for CAT006-4
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
