% Refutation found. Thanks to Tanya!
% SZS status Theorem for CSR050+3
% SZS output start Proof for CSR050+3
fof(f68803,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f27695,f44623,f45778,f47488,f68802])).
fof(f68802,plain,(
  ~spl0_263),
  inference(avatar_contradiction_clause,[],[f68801])).
fof(f68801,plain,(
  $false | ~spl0_263),
  inference(subsumption_resolution,[],[f68798,f24406])).
fof(f24406,plain,(
  ~tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma))),
  inference(cnf_transformation,[],[f17118])).
fof(f17118,plain,(
  ~tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma)) & mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32))),
  inference(ennf_transformation,[],[f8007])).
fof(f8007,negated_conjecture,(
  ~(mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32)) => tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma)))),
  inference(negated_conjecture,[],[f8006])).
fof(f8006,conjecture,(
  mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32)) => tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma))),
  file('CSR050+3.tptp',unknown)).
fof(f68798,plain,(
  tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma)) | ~spl0_263),
  inference(resolution,[],[f47912,f18124])).
fof(f18124,plain,(
  ( ! [X0,X1] : (~tptptypes_8_692(X0,X1) | tptptypes_7_691(X1,X0)) )),
  inference(cnf_transformation,[],[f12862])).
fof(f12862,plain,(
  ! [X0,X1] : (tptptypes_7_691(X1,X0) | ~tptptypes_8_692(X0,X1))),
  inference(ennf_transformation,[],[f8060])).
fof(f8060,plain,(
  ! [X0,X1] : (tptptypes_8_692(X0,X1) => tptptypes_7_691(X1,X0))),
  inference(rectify,[],[f1024])).
fof(f1024,axiom,(
  ! [X2,X3] : (tptptypes_8_692(X2,X3) => tptptypes_7_691(X3,X2))),
  file('CSR050+3.tptp',unknown)).
fof(f47912,plain,(
  tptptypes_8_692(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939) | ~spl0_263),
  inference(resolution,[],[f25619,f18939])).
fof(f18939,plain,(
  ( ! [X0,X1] : (~tptptypes_9_693(X0,X1) | tptptypes_8_692(X0,X1)) )),
  inference(cnf_transformation,[],[f13196])).
fof(f13196,plain,(
  ! [X0,X1] : (tptptypes_8_692(X0,X1) | ~tptptypes_9_693(X0,X1))),
  inference(ennf_transformation,[],[f8102])).
fof(f8102,plain,(
  ! [X0,X1] : (tptptypes_9_693(X0,X1) => tptptypes_8_692(X0,X1))),
  inference(rectify,[],[f1858])).
fof(f1858,axiom,(
  ! [X2,X3] : (tptptypes_9_693(X2,X3) => tptptypes_8_692(X2,X3))),
  file('CSR050+3.tptp',unknown)).
fof(f25619,plain,(
  tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939) | ~spl0_263),
  inference(avatar_component_clause,[],[f25617])).
fof(f25617,plain,(
  spl0_263 <=> tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_263])])).
fof(f47488,plain,(
  ~spl0_2021 | spl0_39),
  inference(avatar_split_clause,[],[f44093,f24577,f45775])).
fof(f45775,plain,(
  spl0_2021 <=> mtvisible(c_saicterrorismeventsdatamt)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2021])])).
fof(f24577,plain,(
  spl0_39 <=> mtvisible(c_currentworlddatacollectormt_nonhomocentric)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_39])])).
fof(f44093,plain,(
  mtvisible(c_currentworlddatacollectormt_nonhomocentric) | ~mtvisible(c_saicterrorismeventsdatamt)),
  inference(resolution,[],[f24396,f19874])).
fof(f19874,plain,(
  genlmt(c_saicterrorismeventsdatamt,c_currentworlddatacollectormt_nonhomocentric)),
  inference(cnf_transformation,[],[f2803])).
fof(f2803,axiom,(
  genlmt(c_saicterrorismeventsdatamt,c_currentworlddatacollectormt_nonhomocentric)),
  file('CSR050+3.tptp',unknown)).
fof(f24396,plain,(
  ( ! [X0,X1] : (~genlmt(X0,X1) | mtvisible(X1) | ~mtvisible(X0)) )),
  inference(cnf_transformation,[],[f17109])).
fof(f17109,plain,(
  ! [X0,X1] : (mtvisible(X1) | ~genlmt(X0,X1) | ~mtvisible(X0))),
  inference(flattening,[],[f17108])).
fof(f17108,plain,(
  ! [X0,X1] : (mtvisible(X1) | (~genlmt(X0,X1) | ~mtvisible(X0)))),
  inference(ennf_transformation,[],[f11718])).
fof(f11718,plain,(
  ! [X0,X1] : ((genlmt(X0,X1) & mtvisible(X0)) => mtvisible(X1))),
  inference(rectify,[],[f7997])).
fof(f7997,axiom,(
  ! [X20,X21] : ((genlmt(X20,X21) & mtvisible(X20)) => mtvisible(X21))),
  file('CSR050+3.tptp',unknown)).
fof(f45778,plain,(
  ~spl0_1726 | spl0_2021),
  inference(avatar_split_clause,[],[f43406,f45775,f44273])).
fof(f44273,plain,(
  spl0_1726 <=> mtvisible(c_machinelearningspindleheadmt)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1726])])).
fof(f43406,plain,(
  mtvisible(c_saicterrorismeventsdatamt) | ~mtvisible(c_machinelearningspindleheadmt)),
  inference(resolution,[],[f24396,f20242])).
fof(f20242,plain,(
  genlmt(c_machinelearningspindleheadmt,c_saicterrorismeventsdatamt)),
  inference(cnf_transformation,[],[f3178])).
fof(f3178,axiom,(
  genlmt(c_machinelearningspindleheadmt,c_saicterrorismeventsdatamt)),
  file('CSR050+3.tptp',unknown)).
fof(f44623,plain,(
  spl0_1726),
  inference(avatar_split_clause,[],[f44622,f44273])).
fof(f44622,plain,(
  mtvisible(c_machinelearningspindleheadmt)),
  inference(subsumption_resolution,[],[f43115,f24405])).
fof(f24405,plain,(
  mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32))),
  inference(cnf_transformation,[],[f17118])).
fof(f43115,plain,(
  mtvisible(c_machinelearningspindleheadmt) | ~mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32))),
  inference(resolution,[],[f24396,f18140])).
fof(f18140,plain,(
  genlmt(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32),c_machinelearningspindleheadmt)),
  inference(cnf_transformation,[],[f1040])).
fof(f1040,axiom,(
  genlmt(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32),c_machinelearningspindleheadmt)),
  file('CSR050+3.tptp',unknown)).
fof(f27695,plain,(
  ~spl0_39 | spl0_263),
  inference(avatar_split_clause,[],[f21000,f25617,f24577])).
fof(f21000,plain,(
  tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939) | ~mtvisible(c_currentworlddatacollectormt_nonhomocentric)),
  inference(cnf_transformation,[],[f14027])).
fof(f14027,plain,(
  tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939) | ~mtvisible(c_currentworlddatacollectormt_nonhomocentric)),
  inference(ennf_transformation,[],[f3946])).
fof(f3946,axiom,(
  mtvisible(c_currentworlddatacollectormt_nonhomocentric) => tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939)),
  file('CSR050+3.tptp',unknown)).
% SZS output end Proof for CSR050+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 35692
% Time elapsed: 0.381 s
% ------------------------------
% ------------------------------
