% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for FLD025-4
% SZS output start Proof for FLD025-4
fof(f1463,plain,(
  $false),
  inference(subsumption_resolution,[],[f1462,f28])).
fof(f28,axiom,(
  defined(b)),
  file('FLD025-4.tptp',unknown)).
fof(f1462,plain,(
  ~defined(b)),
  inference(subsumption_resolution,[],[f1455,f35])).
fof(f35,axiom,(
  ~product(d,b,u)),
  file('FLD025-4.tptp',unknown)).
fof(f1455,plain,(
  product(d,b,u) | ~defined(b)),
  inference(resolution,[],[f458,f1189])).
fof(f1189,plain,(
  product(c,b,u)),
  inference(resolution,[],[f1175,f10])).
fof(f10,axiom,(
  ( ! [X0,X5,X3] : (~product(X0,X3,X5) | product(X3,X0,X5)) )),
  file('FLD025-4.tptp',unknown)).
fof(f1175,plain,(
  product(b,c,u)),
  inference(resolution,[],[f299,f32])).
fof(f32,axiom,(
  product(multiplicative_identity,a,b)),
  file('FLD025-4.tptp',unknown)).
fof(f299,plain,(
  ( ! [X4] : (~product(multiplicative_identity,a,X4) | product(X4,c,u)) )),
  inference(subsumption_resolution,[],[f295,f31])).
fof(f31,axiom,(
  defined(u)),
  file('FLD025-4.tptp',unknown)).
fof(f295,plain,(
  ( ! [X4] : (~product(multiplicative_identity,a,X4) | product(X4,c,u) | ~defined(u)) )),
  inference(resolution,[],[f158,f8])).
fof(f8,axiom,(
  ( ! [X0] : (product(multiplicative_identity,X0,X0) | ~defined(X0)) )),
  file('FLD025-4.tptp',unknown)).
fof(f158,plain,(
  ( ! [X24,X23,X25] : (~product(X23,u,X24) | ~product(X23,a,X25) | product(X25,c,X24)) )),
  inference(resolution,[],[f7,f34])).
fof(f34,axiom,(
  product(a,c,u)),
  file('FLD025-4.tptp',unknown)).
fof(f7,axiom,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~product(X3,X5,X1) | ~product(X0,X1,X2) | ~product(X0,X3,X4) | product(X4,X5,X2)) )),
  file('FLD025-4.tptp',unknown)).
fof(f458,plain,(
  ( ! [X17,X18] : (~product(c,X17,X18) | product(d,X17,X18) | ~defined(X17)) )),
  inference(resolution,[],[f156,f39])).
fof(f39,plain,(
  product(c,multiplicative_identity,d)),
  inference(resolution,[],[f10,f33])).
fof(f33,axiom,(
  product(multiplicative_identity,c,d)),
  file('FLD025-4.tptp',unknown)).
fof(f156,plain,(
  ( ! [X17,X15,X18,X16] : (~product(X15,multiplicative_identity,X18) | ~product(X15,X16,X17) | product(X18,X16,X17) | ~defined(X16)) )),
  inference(resolution,[],[f7,f8])).
% SZS output end Proof for FLD025-4
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5628
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
