% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for FLD015-3
% SZS output start Proof for FLD015-3
fof(f3010,plain,(
  $false),
  inference(subsumption_resolution,[],[f3009,f28])).
fof(f28,axiom,(
  defined(b)),
  file('FLD015-3.tptp',unknown)).
fof(f3009,plain,(
  ~defined(b)),
  inference(subsumption_resolution,[],[f2997,f30])).
fof(f30,axiom,(
  ~sum(additive_identity,a,b)),
  file('FLD015-3.tptp',unknown)).
fof(f2997,plain,(
  sum(additive_identity,a,b) | ~defined(b)),
  inference(resolution,[],[f2961,f410])).
fof(f410,plain,(
  ( ! [X4,X3] : (~sum(additive_identity,X3,X4) | sum(additive_identity,X4,X3) | ~defined(X3)) )),
  inference(subsumption_resolution,[],[f403,f14])).
fof(f14,axiom,(
  defined(additive_identity)),
  file('FLD015-3.tptp',unknown)).
fof(f403,plain,(
  ( ! [X4,X3] : (~sum(additive_identity,X3,X4) | sum(additive_identity,X4,X3) | ~defined(X3) | ~defined(additive_identity)) )),
  inference(resolution,[],[f108,f3])).
fof(f3,axiom,(
  ( ! [X0] : (sum(additive_identity,X0,X0) | ~defined(X0)) )),
  file('FLD015-3.tptp',unknown)).
fof(f108,plain,(
  ( ! [X6,X4,X5,X3] : (~sum(X6,X3,additive_identity) | ~sum(X3,X4,X5) | sum(X6,X5,X4) | ~defined(X4)) )),
  inference(resolution,[],[f1,f3])).
fof(f1,axiom,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sum(X4,X5,X2) | ~sum(X3,X5,X1) | ~sum(X0,X3,X4) | sum(X0,X1,X2)) )),
  file('FLD015-3.tptp',unknown)).
fof(f2961,plain,(
  sum(additive_identity,b,a)),
  inference(resolution,[],[f2926,f5])).
fof(f5,axiom,(
  ( ! [X0,X5,X3] : (~sum(X0,X3,X5) | sum(X3,X0,X5)) )),
  file('FLD015-3.tptp',unknown)).
fof(f2926,plain,(
  sum(b,additive_identity,a)),
  inference(subsumption_resolution,[],[f2925,f28])).
fof(f2925,plain,(
  sum(b,additive_identity,a) | ~defined(b)),
  inference(subsumption_resolution,[],[f2910,f27])).
fof(f27,axiom,(
  defined(a)),
  file('FLD015-3.tptp',unknown)).
fof(f2910,plain,(
  sum(b,additive_identity,a) | ~defined(a) | ~defined(b)),
  inference(resolution,[],[f1575,f409])).
fof(f409,plain,(
  ( ! [X21,X19,X20] : (~sum(additive_inverse(X19),X20,X21) | sum(X19,X21,X20) | ~defined(X20) | ~defined(X19)) )),
  inference(resolution,[],[f108,f33])).
fof(f33,plain,(
  ( ! [X1] : (sum(X1,additive_inverse(X1),additive_identity) | ~defined(X1)) )),
  inference(resolution,[],[f5,f4])).
fof(f4,axiom,(
  ( ! [X0] : (sum(additive_inverse(X0),X0,additive_identity) | ~defined(X0)) )),
  file('FLD015-3.tptp',unknown)).
fof(f1575,plain,(
  sum(additive_inverse(b),a,additive_identity)),
  inference(subsumption_resolution,[],[f1574,f27])).
fof(f1574,plain,(
  sum(additive_inverse(b),a,additive_identity) | ~defined(a)),
  inference(duplicate_literal_removal,[],[f1556])).
fof(f1556,plain,(
  sum(additive_inverse(b),a,additive_identity) | ~defined(a) | ~defined(a)),
  inference(resolution,[],[f414,f4])).
fof(f414,plain,(
  ( ! [X8,X7] : (~sum(additive_inverse(a),X7,X8) | sum(additive_inverse(b),X7,X8) | ~defined(X7)) )),
  inference(resolution,[],[f120,f31])).
fof(f31,plain,(
  sum(additive_inverse(a),additive_identity,additive_inverse(b))),
  inference(resolution,[],[f5,f29])).
fof(f29,axiom,(
  sum(additive_identity,additive_inverse(a),additive_inverse(b))),
  file('FLD015-3.tptp',unknown)).
fof(f120,plain,(
  ( ! [X6,X4,X5,X3] : (~sum(X3,additive_identity,X6) | ~sum(X3,X4,X5) | sum(X6,X4,X5) | ~defined(X4)) )),
  inference(resolution,[],[f2,f3])).
fof(f2,axiom,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sum(X3,X5,X1) | ~sum(X0,X1,X2) | ~sum(X0,X3,X4) | sum(X4,X5,X2)) )),
  file('FLD015-3.tptp',unknown)).
% SZS output end Proof for FLD015-3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6524
% Time elapsed: 0.060 s
% ------------------------------
% ------------------------------
