% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP473-1
% SZS output start Proof for GRP473-1
fof(f16658,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f16645])).
fof(f16645,plain,(
  a2 != a2),
  inference(superposition,[],[f4,f10276])).
fof(f10276,plain,(
  ( ! [X88,X85] : (divide(divide(X85,X85),inverse(X88)) = X88) )),
  inference(backward_demodulation,[],[f10273,f8059])).
fof(f8059,plain,(
  ( ! [X88,X85] : (divide(divide(X85,inverse(inverse(X85))),inverse(X88)) = X88) )),
  inference(backward_demodulation,[],[f8057,f8028])).
fof(f8028,plain,(
  ( ! [X88,X85] : (divide(divide(inverse(inverse(X85)),X85),inverse(X88)) = X88) )),
  inference(backward_demodulation,[],[f8027,f5497])).
fof(f5497,plain,(
  ( ! [X88,X85,X84] : (divide(divide(divide(X85,X84),divide(X85,X84)),inverse(X88)) = X88) )),
  inference(backward_demodulation,[],[f4950,f4813])).
fof(f4813,plain,(
  ( ! [X88,X85,X84] : (divide(divide(divide(X85,X84),inverse(divide(X84,X85))),inverse(X88)) = X88) )),
  inference(backward_demodulation,[],[f4778,f989])).
fof(f989,plain,(
  ( ! [X88,X87,X85,X86,X84] : (divide(divide(divide(X85,X84),inverse(divide(X84,X85))),divide(inverse(X87),divide(divide(X88,X86),inverse(divide(X86,X87))))) = X88) )),
  inference(superposition,[],[f166,f714])).
fof(f714,plain,(
  ( ! [X28,X26,X29,X25] : (divide(X26,X25) = divide(divide(inverse(divide(X29,X28)),divide(X28,X29)),divide(X25,X26))) )),
  inference(superposition,[],[f379,f5])).
fof(f5,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (divide(X3,X2) = divide(divide(inverse(X1),divide(divide(X4,X5),divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0)))),divide(X5,X4))) )),
  inference(superposition,[],[f1,f1])).
fof(f1,axiom,(
  ( ! [X2,X0,X3,X1] : (divide(divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0)),divide(X3,X2)) = X1) )),
  file('GRP473-1.tptp',unknown)).
fof(f379,plain,(
  ( ! [X294,X296,X298,X295,X297] : (divide(divide(inverse(divide(divide(inverse(X295),divide(X294,divide(inverse(divide(X296,X295)),divide(X297,X296)))),X298)),X297),X294) = X298) )),
  inference(forward_demodulation,[],[f371,f1])).
fof(f371,plain,(
  ( ! [X292,X294,X296,X298,X293,X295,X291,X297] : (divide(divide(inverse(divide(divide(inverse(X295),divide(divide(divide(inverse(divide(X293,X294)),divide(divide(X292,X291),X293)),divide(X291,X292)),divide(inverse(divide(X296,X295)),divide(X297,X296)))),X298)),X297),X294) = X298) )),
  inference(superposition,[],[f8,f166])).
fof(f8,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (divide(divide(inverse(divide(X4,X5)),divide(divide(divide(X3,X2),divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0))),X4)),X1) = X5) )),
  inference(superposition,[],[f1,f1])).
fof(f166,plain,(
  ( ! [X6,X4,X7,X5,X1] : (divide(divide(X4,X5),divide(inverse(X6),divide(divide(X5,X4),divide(inverse(divide(X7,X6)),divide(X1,X7))))) = X1) )),
  inference(superposition,[],[f49,f1])).
fof(f49,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (divide(X21,X22) = divide(divide(X19,X18),divide(inverse(X17),divide(divide(X18,X19),divide(inverse(divide(X20,X17)),divide(divide(X21,X22),X20)))))) )),
  inference(forward_demodulation,[],[f31,f5])).
fof(f31,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (divide(divide(X19,X18),divide(inverse(X17),divide(divide(X18,X19),divide(inverse(divide(X20,X17)),divide(divide(X21,X22),X20))))) = divide(divide(inverse(X23),divide(divide(X24,X25),divide(inverse(divide(X26,X23)),divide(divide(X22,X21),X26)))),divide(X25,X24))) )),
  inference(superposition,[],[f5,f5])).
fof(f4778,plain,(
  ( ! [X12,X10,X11] : (inverse(X12) = divide(inverse(X11),divide(divide(X12,X10),inverse(divide(X10,X11))))) )),
  inference(backward_demodulation,[],[f4708,f3456])).
fof(f3456,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (divide(divide(divide(divide(X13,X14),X15),inverse(divide(X15,X16))),divide(inverse(divide(divide(X14,X13),X12)),X16)) = divide(inverse(X11),divide(divide(X12,X10),inverse(divide(X10,X11))))) )),
  inference(superposition,[],[f1633,f20])).
fof(f20,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (divide(divide(inverse(divide(X12,X13)),divide(X9,X12)),divide(divide(divide(divide(X8,X7),X11),inverse(divide(X11,X10))),divide(inverse(divide(divide(X7,X8),X9)),X10))) = X13) )),
  inference(superposition,[],[f1,f7])).
fof(f7,plain,(
  ( ! [X4,X2,X0,X3,X1] : (divide(divide(inverse(divide(divide(X3,X2),X4)),X1),divide(divide(divide(X2,X3),X0),inverse(divide(X0,X1)))) = X4) )),
  inference(superposition,[],[f1,f1])).
fof(f1633,plain,(
  ( ! [X14,X15,X13] : (divide(inverse(divide(divide(X14,X15),X13)),divide(X15,X14)) = X13) )),
  inference(forward_demodulation,[],[f1378,f1058])).
fof(f1058,plain,(
  ( ! [X14,X17,X15,X16] : (divide(X16,X17) = divide(divide(X16,X17),divide(inverse(divide(X14,X15)),divide(X15,X14)))) )),
  inference(forward_demodulation,[],[f954,f714])).
fof(f954,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (divide(divide(X16,X17),divide(inverse(divide(X14,X15)),divide(X15,X14))) = divide(divide(inverse(divide(X18,X19)),divide(X19,X18)),divide(X17,X16))) )),
  inference(superposition,[],[f714,f714])).
fof(f1378,plain,(
  ( ! [X14,X12,X15,X13,X11] : (divide(inverse(divide(divide(X14,X15),X13)),divide(divide(X15,X14),divide(inverse(divide(X11,X12)),divide(X12,X11)))) = X13) )),
  inference(superposition,[],[f968,f714])).
fof(f968,plain,(
  ( ! [X21,X19,X20,X18] : (divide(inverse(divide(X21,divide(X20,divide(X18,X19)))),divide(divide(X19,X18),X21)) = X20) )),
  inference(superposition,[],[f714,f7])).
fof(f4708,plain,(
  ( ! [X26,X24,X23,X25,X22] : (inverse(X22) = divide(divide(divide(divide(X23,X24),X25),inverse(divide(X25,X26))),divide(inverse(divide(divide(X24,X23),X22)),X26))) )),
  inference(backward_demodulation,[],[f4515,f2042])).
fof(f2042,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X22,X20] : (divide(divide(divide(divide(X23,X24),X25),inverse(divide(X25,X26))),divide(inverse(divide(divide(X24,X23),X22)),X26)) = divide(divide(inverse(X21),divide(X27,divide(inverse(divide(X20,X21)),divide(X22,X20)))),divide(divide(divide(divide(X28,X29),X30),inverse(divide(X30,X31))),divide(inverse(divide(divide(X29,X28),X27)),X31)))) )),
  inference(superposition,[],[f20,f20])).
fof(f4515,plain,(
  ( ! [X10,X19,X17,X15,X11,X9,X18,X16] : (inverse(X11) = divide(divide(inverse(X10),divide(X15,divide(inverse(divide(X9,X10)),divide(X11,X9)))),divide(divide(divide(divide(X16,X17),X18),inverse(divide(X18,X19))),divide(inverse(divide(divide(X17,X16),X15)),X19)))) )),
  inference(backward_demodulation,[],[f4271,f2041])).
fof(f2041,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X9,X18,X16] : (divide(divide(X12,X13),divide(inverse(divide(X14,X11)),divide(divide(X13,X12),X14))) = divide(divide(inverse(X10),divide(X15,divide(inverse(divide(X9,X10)),divide(X11,X9)))),divide(divide(divide(divide(X16,X17),X18),inverse(divide(X18,X19))),divide(inverse(divide(divide(X17,X16),X15)),X19)))) )),
  inference(superposition,[],[f20,f6])).
fof(f6,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (divide(divide(inverse(divide(X4,X5)),divide(X1,X4)),divide(divide(X3,X2),divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0)))) = X5) )),
  inference(superposition,[],[f1,f1])).
fof(f4271,plain,(
  ( ! [X2,X0,X3,X1] : (inverse(X1) = divide(divide(X3,X2),divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0)))) )),
  inference(backward_demodulation,[],[f4254,f952])).
fof(f952,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (divide(divide(X3,X2),divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0))) = divide(divide(inverse(divide(X4,X5)),divide(X5,X4)),X1)) )),
  inference(superposition,[],[f714,f1])).
fof(f4254,plain,(
  ( ! [X522,X518,X521] : (divide(divide(inverse(divide(X521,X522)),divide(X522,X521)),X518) = inverse(X518)) )),
  inference(backward_demodulation,[],[f4185,f1484])).
fof(f1484,plain,(
  ( ! [X520,X522,X518,X521,X517,X519] : (divide(divide(divide(X520,X519),X517),inverse(divide(X517,divide(X518,divide(X519,X520))))) = divide(divide(inverse(divide(X521,X522)),divide(X522,X521)),X518)) )),
  inference(superposition,[],[f714,f968])).
fof(f4185,plain,(
  ( ! [X80,X78,X81,X79] : (inverse(X80) = divide(divide(divide(X79,X78),X81),inverse(divide(X81,divide(X80,divide(X78,X79)))))) )),
  inference(backward_demodulation,[],[f4174,f2691])).
fof(f2691,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (divide(divide(divide(X79,X78),X81),inverse(divide(X81,divide(X80,divide(X78,X79))))) = divide(divide(inverse(X80),divide(X76,X77)),divide(X77,X76))) )),
  inference(superposition,[],[f21,f2013])).
fof(f2013,plain,(
  ( ! [X94,X90,X89] : (divide(X94,divide(inverse(divide(X89,X90)),divide(X90,X89))) = X94) )),
  inference(backward_demodulation,[],[f2012,f990])).
fof(f990,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (divide(divide(inverse(divide(divide(inverse(X92),divide(divide(X93,X91),inverse(divide(X91,X92)))),X94)),X93),divide(inverse(divide(X89,X90)),divide(X90,X89))) = X94) )),
  inference(superposition,[],[f379,f714])).
fof(f2012,plain,(
  ( ! [X45,X43,X44,X40] : (divide(inverse(divide(divide(inverse(X40),divide(divide(X44,X43),inverse(divide(X43,X40)))),X45)),X44) = X45) )),
  inference(forward_demodulation,[],[f1883,f2007])).
fof(f2007,plain,(
  ( ! [X39,X35,X36,X40] : (divide(X39,X40) = divide(divide(divide(X36,X35),inverse(divide(X35,X36))),divide(X40,X39))) )),
  inference(backward_demodulation,[],[f2005,f981])).
fof(f981,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (divide(X39,X40) = divide(divide(divide(X36,X35),inverse(divide(X35,X36))),divide(inverse(X38),divide(divide(divide(X39,X40),X37),inverse(divide(X37,X38)))))) )),
  inference(superposition,[],[f49,f714])).
fof(f2005,plain,(
  ( ! [X33,X29,X34,X32] : (divide(X34,X33) = divide(inverse(X29),divide(divide(divide(X33,X34),X32),inverse(divide(X32,X29))))) )),
  inference(forward_demodulation,[],[f1881,f714])).
fof(f1881,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (divide(X34,X33) = divide(inverse(X29),divide(divide(inverse(divide(X30,X31)),divide(X31,X30)),divide(inverse(divide(X32,X29)),divide(divide(X33,X34),X32))))) )),
  inference(superposition,[],[f1640,f5])).
fof(f1640,plain,(
  ( ! [X177,X178,X181] : (divide(X181,divide(divide(X178,X177),inverse(divide(X177,X178)))) = X181) )),
  inference(backward_demodulation,[],[f1633,f1116])).
fof(f1116,plain,(
  ( ! [X177,X180,X178,X181,X179] : (divide(divide(inverse(divide(divide(X179,X180),X181)),divide(X180,X179)),divide(divide(X178,X177),inverse(divide(X177,X178)))) = X181) )),
  inference(forward_demodulation,[],[f1002,f972])).
fof(f972,plain,(
  ( ! [X43,X41,X44,X42,X40] : (divide(X44,X43) = divide(divide(X40,X41),divide(inverse(divide(X42,divide(X43,X44))),divide(divide(X41,X40),X42)))) )),
  inference(superposition,[],[f714,f8])).
fof(f1002,plain,(
  ( ! [X177,X182,X180,X178,X184,X183,X181,X179] : (divide(divide(inverse(divide(divide(X179,X180),X181)),divide(X180,X179)),divide(divide(X182,X183),divide(inverse(divide(X184,divide(inverse(divide(X177,X178)),divide(X178,X177)))),divide(divide(X183,X182),X184)))) = X181) )),
  inference(superposition,[],[f6,f714])).
fof(f1883,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (divide(inverse(divide(divide(inverse(X40),divide(divide(divide(X41,X42),inverse(divide(X42,X41))),divide(inverse(divide(X43,X40)),divide(X44,X43)))),X45)),X44) = X45) )),
  inference(superposition,[],[f1640,f379])).
fof(f21,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (divide(divide(divide(X15,X14),X18),inverse(divide(X18,X17))) = divide(divide(inverse(X16),divide(divide(X19,X20),divide(inverse(divide(divide(X14,X15),X16)),X17))),divide(X20,X19))) )),
  inference(superposition,[],[f1,f7])).
fof(f4174,plain,(
  ( ! [X573,X569,X570] : (divide(divide(X573,divide(X569,X570)),divide(X570,X569)) = X573) )),
  inference(backward_demodulation,[],[f3494,f2423])).
fof(f2423,plain,(
  ( ! [X573,X569,X571,X572,X570] : (divide(divide(inverse(divide(divide(inverse(divide(X571,X572)),divide(X572,X571)),X573)),divide(X569,X570)),divide(X570,X569)) = X573) )),
  inference(superposition,[],[f1,f2013])).
fof(f3494,plain,(
  ( ! [X10,X8,X9] : (inverse(divide(divide(inverse(divide(X8,X9)),divide(X9,X8)),X10)) = X10) )),
  inference(superposition,[],[f1633,f1640])).
fof(f4950,plain,(
  ( ! [X553,X555,X554] : (divide(X553,divide(X554,X555)) = divide(X553,inverse(divide(X555,X554)))) )),
  inference(backward_demodulation,[],[f4836,f1489])).
fof(f1489,plain,(
  ( ! [X553,X555,X549,X551,X552,X554,X550] : (divide(X553,divide(X554,X555)) = divide(X553,divide(X552,inverse(divide(divide(X550,X549),divide(divide(X555,X554),divide(divide(divide(X549,X550),X551),inverse(divide(X551,X552))))))))) )),
  inference(superposition,[],[f15,f968])).
fof(f15,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (divide(divide(inverse(divide(divide(divide(divide(X2,X3),X0),inverse(divide(X0,X1))),X4)),X5),divide(X1,inverse(divide(divide(X3,X2),X5)))) = X4) )),
  inference(superposition,[],[f7,f1])).
fof(f4836,plain,(
  ( ! [X52,X50,X48,X51,X49] : (inverse(X48) = divide(X52,inverse(divide(divide(X49,X50),divide(X48,divide(divide(divide(X50,X49),X51),inverse(divide(X51,X52)))))))) )),
  inference(backward_demodulation,[],[f4782,f3698])).
fof(f3698,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (divide(X52,inverse(divide(divide(X49,X50),divide(X48,divide(divide(divide(X50,X49),X51),inverse(divide(X51,X52))))))) = divide(divide(inverse(X53),X54),divide(inverse(divide(X55,X48)),divide(X54,inverse(divide(X53,X55)))))) )),
  inference(backward_demodulation,[],[f3637,f1545])).
fof(f1545,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (divide(divide(inverse(X53),X54),divide(divide(divide(divide(X48,X55),inverse(divide(X55,X53))),X56),inverse(divide(X56,X54)))) = divide(X52,inverse(divide(divide(X49,X50),divide(X48,divide(divide(divide(X50,X49),X51),inverse(divide(X51,X52)))))))) )),
  inference(backward_demodulation,[],[f1375,f232])).
fof(f232,plain,(
  ( ! [X47,X54,X52,X50,X48,X56,X55,X53,X51,X49] : (divide(X52,divide(inverse(divide(X47,X48)),divide(divide(divide(X49,X50),divide(inverse(divide(X51,X52)),divide(divide(X50,X49),X51))),X47))) = divide(divide(inverse(X53),X54),divide(divide(divide(divide(X48,X55),inverse(divide(X55,X53))),X56),inverse(divide(X56,X54))))) )),
  inference(superposition,[],[f11,f8])).
fof(f11,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (divide(X3,X2) = divide(divide(inverse(X1),X4),divide(divide(divide(divide(divide(X2,X3),X0),inverse(divide(X0,X1))),X5),inverse(divide(X5,X4))))) )),
  inference(superposition,[],[f7,f1])).
fof(f1375,plain,(
  ( ! [X103,X105,X102,X106,X104] : (inverse(divide(X102,divide(X103,divide(X104,X105)))) = divide(inverse(divide(X106,X103)),divide(divide(X102,divide(X105,X104)),X106))) )),
  inference(superposition,[],[f968,f968])).
fof(f3637,plain,(
  ( ! [X94,X90,X95,X93,X91,X96] : (divide(divide(divide(divide(X90,X91),inverse(divide(X93,X94))),X95),inverse(divide(X95,X96))) = divide(inverse(divide(X91,X90)),divide(X96,inverse(divide(X94,X93))))) )),
  inference(backward_demodulation,[],[f3587,f3328])).
fof(f3328,plain,(
  ( ! [X94,X90,X97,X95,X93,X91,X98,X96] : (divide(divide(inverse(divide(X91,X90)),divide(divide(X97,X98),divide(inverse(divide(X94,X93)),X96))),divide(X98,X97)) = divide(divide(divide(divide(X90,X91),inverse(divide(X93,X94))),X95),inverse(divide(X95,X96)))) )),
  inference(forward_demodulation,[],[f2668,f2921])).
fof(f2921,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (divide(divide(divide(X16,X17),divide(inverse(divide(X18,X15)),divide(divide(X19,X20),X18))),inverse(X15)) = divide(divide(X16,X17),inverse(divide(X19,X20)))) )),
  inference(backward_demodulation,[],[f2920,f34])).
fof(f34,plain,(
  ( ! [X23,X21,X19,X17,X15,X22,X20,X18,X16] : (divide(divide(divide(X16,X17),divide(inverse(divide(X18,X15)),divide(divide(X19,X20),X18))),inverse(X15)) = divide(divide(inverse(X21),divide(divide(X22,X23),divide(inverse(divide(divide(X17,X16),X21)),divide(X20,X19)))),divide(X23,X22))) )),
  inference(superposition,[],[f5,f5])).
fof(f2920,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (divide(divide(inverse(X38),divide(divide(X39,X40),divide(inverse(divide(divide(X37,X36),X38)),divide(X34,X35)))),divide(X40,X39)) = divide(divide(X36,X37),inverse(divide(X35,X34)))) )),
  inference(forward_demodulation,[],[f2598,f1058])).
fof(f2598,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32,X40] : (divide(divide(inverse(X38),divide(divide(X39,X40),divide(inverse(divide(divide(X37,X36),X38)),divide(X34,X35)))),divide(X40,X39)) = divide(divide(divide(X36,X37),divide(inverse(divide(X32,X33)),divide(X33,X32))),inverse(divide(X35,X34)))) )),
  inference(superposition,[],[f21,f714])).
fof(f2668,plain,(
  ( ! [X94,X92,X90,X97,X95,X93,X91,X89,X98,X96] : (divide(divide(divide(divide(divide(X90,X91),divide(inverse(divide(X92,X89)),divide(divide(X93,X94),X92))),inverse(X89)),X95),inverse(divide(X95,X96))) = divide(divide(inverse(divide(X91,X90)),divide(divide(X97,X98),divide(inverse(divide(X94,X93)),X96))),divide(X98,X97))) )),
  inference(superposition,[],[f21,f5])).
fof(f3587,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (divide(divide(inverse(divide(X3,X2)),divide(divide(X6,X7),divide(inverse(X0),X1))),divide(X7,X6)) = divide(inverse(divide(X3,X2)),divide(X1,inverse(X0)))) )),
  inference(backward_demodulation,[],[f3464,f256])).
fof(f256,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (divide(divide(divide(divide(divide(X2,X3),X4),inverse(divide(X4,X0))),X5),inverse(divide(X5,X1))) = divide(divide(inverse(divide(X3,X2)),divide(divide(X6,X7),divide(inverse(X0),X1))),divide(X7,X6))) )),
  inference(superposition,[],[f1,f11])).
fof(f3464,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (divide(divide(divide(divide(divide(X61,X62),X63),inverse(divide(X63,X59))),X64),inverse(divide(X64,X60))) = divide(inverse(divide(X62,X61)),divide(X60,inverse(X59)))) )),
  inference(superposition,[],[f1633,f11])).
fof(f4782,plain,(
  ( ! [X24,X23,X25,X20] : (inverse(X20) = divide(divide(inverse(X23),X24),divide(inverse(divide(X25,X20)),divide(X24,inverse(divide(X23,X25)))))) )),
  inference(backward_demodulation,[],[f4708,f3652])).
fof(f3652,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (divide(divide(divide(divide(X19,X18),X22),inverse(divide(X22,X21))),divide(inverse(divide(divide(X18,X19),X20)),X21)) = divide(divide(inverse(X23),X24),divide(inverse(divide(X25,X20)),divide(X24,inverse(divide(X23,X25)))))) )),
  inference(backward_demodulation,[],[f3637,f229])).
fof(f229,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (divide(divide(divide(divide(X19,X18),X22),inverse(divide(X22,X21))),divide(inverse(divide(divide(X18,X19),X20)),X21)) = divide(divide(inverse(X23),X24),divide(divide(divide(divide(X20,X25),inverse(divide(X25,X23))),X26),inverse(divide(X26,X24))))) )),
  inference(superposition,[],[f11,f7])).
fof(f8027,plain,(
  ( ! [X109,X107,X110] : (divide(divide(X107,X110),divide(X109,X110)) = divide(inverse(inverse(X107)),X109)) )),
  inference(forward_demodulation,[],[f8008,f6834])).
fof(f6834,plain,(
  ( ! [X8,X9] : (divide(divide(X9,X8),inverse(X8)) = X9) )),
  inference(backward_demodulation,[],[f6721,f4505])).
fof(f4505,plain,(
  ( ! [X8,X9] : (divide(inverse(divide(X8,X9)),inverse(X8)) = X9) )),
  inference(backward_demodulation,[],[f4271,f1877])).
fof(f1877,plain,(
  ( ! [X12,X10,X8,X11,X9] : (divide(inverse(divide(X8,X9)),divide(divide(X10,X11),divide(inverse(divide(X12,X8)),divide(divide(X11,X10),X12)))) = X9) )),
  inference(superposition,[],[f1640,f9])).
fof(f9,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (divide(divide(inverse(divide(X1,X4)),X5),divide(divide(divide(divide(X3,X2),divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0))),X6),inverse(divide(X6,X5)))) = X4) )),
  inference(superposition,[],[f7,f1])).
fof(f6721,plain,(
  ( ! [X21,X20] : (divide(X20,X21) = inverse(divide(X21,X20))) )),
  inference(forward_demodulation,[],[f5559,f4169])).
fof(f4169,plain,(
  ( ! [X463,X465,X460,X464] : (divide(divide(X463,X464),divide(inverse(X465),divide(divide(X464,X463),divide(X465,X460)))) = X460) )),
  inference(backward_demodulation,[],[f3494,f2404])).
fof(f2404,plain,(
  ( ! [X461,X463,X465,X460,X462,X464] : (divide(divide(X463,X464),divide(inverse(X465),divide(divide(X464,X463),divide(inverse(divide(divide(inverse(divide(X461,X462)),divide(X462,X461)),X465)),X460)))) = X460) )),
  inference(superposition,[],[f166,f2013])).
fof(f5559,plain,(
  ( ! [X24,X23,X21,X22,X20] : (inverse(divide(X21,X20)) = divide(divide(X23,X24),divide(inverse(X22),divide(divide(X24,X23),divide(X22,divide(X20,X21)))))) )),
  inference(backward_demodulation,[],[f4950,f2337])).
fof(f2337,plain,(
  ( ! [X24,X23,X21,X22,X20] : (inverse(divide(X21,X20)) = divide(divide(X23,X24),divide(inverse(X22),divide(divide(X24,X23),inverse(divide(divide(X20,X21),X22)))))) )),
  inference(superposition,[],[f166,f2013])).
fof(f8008,plain,(
  ( ! [X111,X109,X107,X110] : (divide(divide(X107,X110),divide(X109,X110)) = divide(divide(divide(inverse(inverse(X107)),X109),X111),inverse(X111))) )),
  inference(backward_demodulation,[],[f8006,f7633])).
fof(f7633,plain,(
  ( ! [X111,X109,X107,X110,X108] : (divide(divide(X107,X110),divide(X109,X110)) = divide(divide(inverse(X108),divide(X111,divide(divide(X108,inverse(X107)),X109))),inverse(X111))) )),
  inference(forward_demodulation,[],[f6843,f6721])).
fof(f6843,plain,(
  ( ! [X111,X109,X107,X110,X108] : (divide(divide(inverse(X108),divide(X111,divide(inverse(divide(inverse(X107),X108)),X109))),inverse(X111)) = divide(divide(X107,X110),divide(X109,X110))) )),
  inference(backward_demodulation,[],[f6721,f4672])).
fof(f4672,plain,(
  ( ! [X111,X109,X107,X110,X108] : (divide(divide(X107,X110),inverse(divide(X110,X109))) = divide(divide(inverse(X108),divide(X111,divide(inverse(divide(inverse(X107),X108)),X109))),inverse(X111))) )),
  inference(forward_demodulation,[],[f4412,f4271])).
fof(f4412,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106,X104] : (divide(divide(X107,X110),inverse(divide(X110,X109))) = divide(divide(inverse(X108),divide(X111,divide(inverse(divide(divide(divide(X104,X105),divide(inverse(divide(X106,X107)),divide(divide(X105,X104),X106))),X108)),X109))),inverse(X111))) )),
  inference(backward_demodulation,[],[f4271,f866])).
fof(f866,plain,(
  ( ! [X111,X109,X107,X105,X114,X112,X110,X108,X106,X104,X113] : (divide(divide(X107,X110),inverse(divide(X110,X109))) = divide(divide(inverse(X108),divide(X111,divide(inverse(divide(divide(divide(X104,X105),divide(inverse(divide(X106,X107)),divide(divide(X105,X104),X106))),X108)),X109))),divide(divide(X112,X113),divide(inverse(divide(X114,X111)),divide(divide(X113,X112),X114))))) )),
  inference(superposition,[],[f6,f13])).
fof(f13,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (divide(divide(inverse(divide(divide(divide(X3,X2),divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0))),X4)),X5),divide(divide(X1,X6),inverse(divide(X6,X5)))) = X4) )),
  inference(superposition,[],[f7,f1])).
fof(f8006,plain,(
  ( ! [X47,X52,X50,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(X47,divide(divide(X50,X52),X49)))) )),
  inference(forward_demodulation,[],[f7952,f7816])).
fof(f7816,plain,(
  ( ! [X24,X19,X17,X18] : (divide(divide(X18,X17),X19) = divide(divide(X18,X24),divide(X19,divide(X24,X17)))) )),
  inference(forward_demodulation,[],[f7815,f4174])).
fof(f7815,plain,(
  ( ! [X24,X21,X19,X17,X20,X18] : (divide(divide(X18,X17),X19) = divide(divide(X18,X24),divide(X19,divide(X24,divide(divide(X17,divide(X21,X20)),divide(X20,X21)))))) )),
  inference(forward_demodulation,[],[f7814,f7723])).
fof(f7723,plain,(
  ( ! [X35,X31,X36,X34,X32] : (divide(divide(X35,divide(divide(X31,X32),X34)),X36) = divide(X35,divide(divide(X36,divide(X32,X31)),X34))) )),
  inference(backward_demodulation,[],[f7709,f7680])).
fof(f7680,plain,(
  ( ! [X37,X35,X31,X36,X34,X32] : (divide(divide(X35,X37),divide(divide(divide(X36,divide(X32,X31)),X34),X37)) = divide(divide(X35,divide(divide(X31,X32),X34)),X36)) )),
  inference(forward_demodulation,[],[f7679,f7440])).
fof(f7440,plain,(
  ( ! [X14,X17,X15,X18] : (divide(divide(X15,X14),X17) = divide(divide(divide(X15,X14),X18),divide(X17,X18))) )),
  inference(forward_demodulation,[],[f7439,f5548])).
fof(f5548,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (divide(divide(X99,X98),X96) = divide(divide(inverse(X97),divide(divide(X100,X101),divide(divide(X97,divide(X98,X99)),X96))),divide(X101,X100))) )),
  inference(backward_demodulation,[],[f4950,f1431])).
fof(f1431,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (divide(divide(X99,X98),X96) = divide(divide(inverse(X97),divide(divide(X100,X101),inverse(divide(X96,divide(X97,divide(X98,X99)))))),divide(X101,X100))) )),
  inference(superposition,[],[f1,f968])).
fof(f7439,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (divide(divide(divide(X15,X14),X18),divide(X17,X18)) = divide(divide(inverse(X16),divide(divide(X19,X20),divide(divide(X16,divide(X14,X15)),X17))),divide(X20,X19))) )),
  inference(forward_demodulation,[],[f6723,f6721])).
fof(f6723,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (divide(divide(inverse(X16),divide(divide(X19,X20),divide(inverse(divide(divide(X14,X15),X16)),X17))),divide(X20,X19)) = divide(divide(divide(X15,X14),X18),divide(X17,X18))) )),
  inference(backward_demodulation,[],[f6721,f21])).
fof(f7679,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (divide(divide(X35,X37),divide(divide(divide(X36,divide(X32,X31)),X34),X37)) = divide(divide(X35,divide(divide(divide(X31,X32),X33),divide(X34,X33))),X36)) )),
  inference(forward_demodulation,[],[f7678,f6721])).
fof(f7678,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (divide(inverse(divide(divide(divide(divide(X31,X32),X33),divide(X34,X33)),X35)),X36) = divide(divide(X35,X37),divide(divide(divide(X36,divide(X32,X31)),X34),X37))) )),
  inference(forward_demodulation,[],[f6875,f6721])).
fof(f6875,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (divide(inverse(divide(divide(divide(divide(X31,X32),X33),divide(X34,X33)),X35)),X36) = divide(divide(X35,X37),divide(divide(inverse(divide(divide(X32,X31),X36)),X34),X37))) )),
  inference(backward_demodulation,[],[f6721,f5053])).
fof(f5053,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (divide(inverse(divide(X37,X35)),divide(divide(inverse(divide(divide(X32,X31),X36)),X34),X37)) = divide(inverse(divide(divide(divide(divide(X31,X32),X33),divide(X34,X33)),X35)),X36)) )),
  inference(backward_demodulation,[],[f4950,f1365])).
fof(f1365,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (divide(inverse(divide(divide(divide(divide(X31,X32),X33),inverse(divide(X33,X34))),X35)),X36) = divide(inverse(divide(X37,X35)),divide(divide(inverse(divide(divide(X32,X31),X36)),X34),X37))) )),
  inference(superposition,[],[f968,f15])).
fof(f7709,plain,(
  ( ! [X2,X0,X3,X1] : (divide(X1,divide(X2,X3)) = divide(divide(X1,X0),divide(divide(X2,X3),X0))) )),
  inference(backward_demodulation,[],[f7703,f7477])).
fof(f7477,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (divide(divide(X1,X0),divide(divide(X2,X3),X0)) = divide(divide(X4,X5),divide(inverse(X6),divide(divide(X5,X4),divide(divide(X6,divide(X3,X2)),X1))))) )),
  inference(forward_demodulation,[],[f6724,f6721])).
fof(f6724,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (divide(divide(X4,X5),divide(inverse(X6),divide(divide(X5,X4),divide(inverse(divide(divide(X3,X2),X6)),X1)))) = divide(divide(X1,X0),divide(divide(X2,X3),X0))) )),
  inference(backward_demodulation,[],[f6721,f172])).
fof(f172,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0)) = divide(divide(X4,X5),divide(inverse(X6),divide(divide(X5,X4),divide(inverse(divide(divide(X3,X2),X6)),X1))))) )),
  inference(superposition,[],[f49,f1])).
fof(f7703,plain,(
  ( ! [X693,X689,X691,X692,X694,X690] : (divide(X694,divide(X693,X692)) = divide(divide(X691,X690),divide(inverse(X689),divide(divide(X690,X691),divide(divide(X689,divide(X692,X693)),X694))))) )),
  inference(forward_demodulation,[],[f7702,f6721])).
fof(f7702,plain,(
  ( ! [X693,X689,X691,X692,X694,X690] : (divide(divide(X691,X690),divide(inverse(X689),divide(divide(X690,X691),divide(inverse(divide(divide(X692,X693),X689)),X694)))) = divide(X694,divide(X693,X692))) )),
  inference(forward_demodulation,[],[f7701,f4172])).
fof(f4172,plain,(
  ( ! [X505,X504,X501] : (divide(X504,X505) = divide(divide(X501,X501),divide(X505,X504))) )),
  inference(backward_demodulation,[],[f3494,f2411])).
fof(f2411,plain,(
  ( ! [X505,X502,X504,X501,X503] : (divide(X504,X505) = divide(divide(inverse(divide(divide(inverse(divide(X502,X503)),divide(X503,X502)),X501)),X501),divide(X505,X504))) )),
  inference(superposition,[],[f714,f2013])).
fof(f7701,plain,(
  ( ! [X697,X693,X689,X691,X696,X692,X694,X690] : (divide(divide(X691,X690),divide(inverse(X689),divide(divide(X690,X691),divide(inverse(divide(divide(X692,X693),X689)),X694)))) = divide(divide(divide(X697,X696),divide(X697,X696)),divide(divide(X693,X692),X694))) )),
  inference(forward_demodulation,[],[f6884,f7440])).
fof(f6884,plain,(
  ( ! [X697,X693,X695,X689,X691,X696,X692,X694,X690] : (divide(divide(X691,X690),divide(inverse(X689),divide(divide(X690,X691),divide(inverse(divide(divide(X692,X693),X689)),X694)))) = divide(divide(divide(X697,X696),divide(X697,X696)),divide(divide(divide(X693,X692),X695),divide(X694,X695)))) )),
  inference(backward_demodulation,[],[f6721,f5103])).
fof(f5103,plain,(
  ( ! [X697,X693,X695,X689,X691,X696,X692,X694,X690] : (divide(divide(X691,X690),divide(inverse(X689),divide(divide(X690,X691),divide(inverse(divide(divide(X692,X693),X689)),X694)))) = divide(divide(inverse(divide(X696,X697)),divide(X697,X696)),divide(divide(divide(X693,X692),X695),divide(X694,X695)))) )),
  inference(backward_demodulation,[],[f4950,f2884])).
fof(f2884,plain,(
  ( ! [X697,X693,X695,X689,X691,X696,X692,X694,X690] : (divide(divide(X691,X690),divide(inverse(X689),divide(divide(X690,X691),divide(inverse(divide(divide(X692,X693),X689)),X694)))) = divide(divide(inverse(divide(X696,X697)),divide(X697,X696)),divide(divide(divide(X693,X692),X695),inverse(divide(X695,X694))))) )),
  inference(superposition,[],[f714,f21])).
fof(f7814,plain,(
  ( ! [X24,X21,X19,X17,X20,X18] : (divide(divide(X18,X17),X19) = divide(divide(X18,X24),divide(X19,divide(divide(X24,divide(divide(X20,X21),divide(X20,X21))),X17)))) )),
  inference(forward_demodulation,[],[f7765,f7723])).
fof(f7765,plain,(
  ( ! [X24,X21,X19,X17,X20,X18] : (divide(divide(X18,X17),X19) = divide(divide(X18,X24),divide(divide(X19,divide(divide(divide(X20,X21),divide(X20,X21)),X17)),X24))) )),
  inference(backward_demodulation,[],[f7723,f7684])).
fof(f7684,plain,(
  ( ! [X24,X21,X19,X17,X20,X18] : (divide(divide(X18,X17),X19) = divide(divide(X18,X24),divide(divide(divide(X19,divide(divide(X21,X20),X17)),divide(X20,X21)),X24))) )),
  inference(forward_demodulation,[],[f7683,f6721])).
fof(f7683,plain,(
  ( ! [X24,X21,X19,X17,X20,X18] : (divide(divide(X18,X17),X19) = divide(inverse(divide(X24,X18)),divide(divide(divide(X19,divide(divide(X21,X20),X17)),divide(X20,X21)),X24))) )),
  inference(forward_demodulation,[],[f7682,f7440])).
fof(f7682,plain,(
  ( ! [X24,X21,X19,X17,X22,X20,X18] : (divide(divide(X18,X17),X19) = divide(inverse(divide(X24,X18)),divide(divide(divide(X19,divide(divide(divide(X21,X20),X22),divide(X17,X22))),divide(X20,X21)),X24))) )),
  inference(forward_demodulation,[],[f6877,f6721])).
fof(f6877,plain,(
  ( ! [X24,X21,X19,X17,X22,X20,X18] : (divide(divide(X18,X17),X19) = divide(inverse(divide(X24,X18)),divide(inverse(divide(divide(X20,X21),divide(X19,divide(divide(divide(X21,X20),X22),divide(X17,X22))))),X24))) )),
  inference(backward_demodulation,[],[f6721,f5064])).
fof(f5064,plain,(
  ( ! [X24,X21,X19,X17,X22,X20,X18] : (divide(inverse(divide(X17,X18)),X19) = divide(inverse(divide(X24,X18)),divide(inverse(divide(divide(X20,X21),divide(X19,divide(divide(divide(X21,X20),X22),divide(X17,X22))))),X24))) )),
  inference(backward_demodulation,[],[f4950,f1570])).
fof(f1570,plain,(
  ( ! [X24,X21,X19,X17,X22,X20,X18] : (divide(inverse(divide(X17,X18)),X19) = divide(inverse(divide(X24,X18)),divide(inverse(divide(divide(X20,X21),divide(X19,divide(divide(divide(X21,X20),X22),inverse(divide(X22,X17)))))),X24))) )),
  inference(backward_demodulation,[],[f1375,f1363])).
fof(f1363,plain,(
  ( ! [X24,X23,X21,X19,X17,X22,X20,X18] : (divide(inverse(divide(X17,X18)),X19) = divide(inverse(divide(X24,X18)),divide(divide(inverse(divide(X23,X19)),divide(divide(divide(X20,X21),divide(inverse(divide(X22,X17)),divide(divide(X21,X20),X22))),X23)),X24))) )),
  inference(superposition,[],[f968,f9])).
fof(f7952,plain,(
  ( ! [X47,X52,X50,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(X47,divide(divide(X50,X51),divide(X49,divide(X51,X52)))))) )),
  inference(backward_demodulation,[],[f7946,f7864])).
fof(f7864,plain,(
  ( ! [X47,X52,X50,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(X47,divide(divide(divide(X50,X51),divide(X52,X51)),X49)))) )),
  inference(forward_demodulation,[],[f7863,f7440])).
fof(f7863,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(X47,divide(divide(divide(divide(X50,X51),divide(X52,X51)),X48),divide(X49,X48))))) )),
  inference(forward_demodulation,[],[f7862,f7816])).
fof(f7862,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(X47,divide(divide(divide(divide(X50,X51),divide(X52,X51)),X46),divide(divide(X49,X48),divide(X46,X48)))))) )),
  inference(forward_demodulation,[],[f7861,f7711])).
fof(f7711,plain,(
  ( ! [X114,X112,X117,X113] : (divide(divide(X113,X112),divide(X117,inverse(X114))) = divide(divide(divide(X113,X112),X114),X117)) )),
  inference(forward_demodulation,[],[f7707,f5548])).
fof(f7707,plain,(
  ( ! [X118,X114,X112,X120,X119,X117,X113] : (divide(divide(X113,X112),divide(X117,inverse(X114))) = divide(divide(inverse(X118),divide(divide(X119,X120),divide(divide(X118,divide(X114,divide(X113,X112))),X117))),divide(X120,X119))) )),
  inference(backward_demodulation,[],[f7703,f7553])).
fof(f7553,plain,(
  ( ! [X111,X109,X118,X114,X112,X120,X110,X119,X117,X113] : (divide(divide(X113,X112),divide(X117,inverse(X114))) = divide(divide(inverse(X118),divide(divide(X119,X120),divide(divide(X118,divide(divide(X111,X110),divide(inverse(X109),divide(divide(X110,X111),divide(divide(X109,divide(X112,X113)),X114))))),X117))),divide(X120,X119))) )),
  inference(forward_demodulation,[],[f7552,f6721])).
fof(f7552,plain,(
  ( ! [X111,X109,X118,X114,X112,X120,X110,X119,X117,X113] : (divide(divide(X113,X112),divide(X117,inverse(X114))) = divide(divide(inverse(X118),divide(divide(X119,X120),divide(divide(X118,divide(divide(X111,X110),divide(inverse(X109),divide(divide(X110,X111),divide(inverse(divide(divide(X112,X113),X109)),X114))))),X117))),divide(X120,X119))) )),
  inference(forward_demodulation,[],[f6779,f6721])).
fof(f6779,plain,(
  ( ! [X111,X109,X118,X114,X112,X120,X110,X119,X117,X113] : (divide(divide(inverse(X118),divide(divide(X119,X120),divide(inverse(divide(divide(divide(X111,X110),divide(inverse(X109),divide(divide(X110,X111),divide(inverse(divide(divide(X112,X113),X109)),X114)))),X118)),X117))),divide(X120,X119)) = divide(divide(X113,X112),divide(X117,inverse(X114)))) )),
  inference(backward_demodulation,[],[f6721,f3617])).
fof(f3617,plain,(
  ( ! [X111,X109,X118,X114,X112,X120,X110,X119,X117,X113] : (divide(divide(inverse(X118),divide(divide(X119,X120),divide(inverse(divide(divide(divide(X111,X110),divide(inverse(X109),divide(divide(X110,X111),divide(inverse(divide(divide(X112,X113),X109)),X114)))),X118)),X117))),divide(X120,X119)) = divide(inverse(divide(X112,X113)),divide(X117,inverse(X114)))) )),
  inference(backward_demodulation,[],[f3464,f2565])).
fof(f2565,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X120,X110,X119,X117,X115,X113] : (divide(divide(inverse(X118),divide(divide(X119,X120),divide(inverse(divide(divide(divide(X111,X110),divide(inverse(X109),divide(divide(X110,X111),divide(inverse(divide(divide(X112,X113),X109)),X114)))),X118)),X117))),divide(X120,X119)) = divide(divide(divide(divide(divide(X113,X112),X115),inverse(divide(X115,X114))),X116),inverse(divide(X116,X117)))) )),
  inference(superposition,[],[f21,f21])).
fof(f7861,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(X47,divide(divide(divide(X50,X51),divide(X52,X51)),divide(divide(divide(X49,X48),divide(X46,X48)),inverse(X46)))))) )),
  inference(forward_demodulation,[],[f7860,f7723])).
fof(f7860,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(X47,divide(divide(divide(divide(X50,X51),divide(X52,X51)),divide(divide(X48,X46),inverse(X46))),divide(X49,X48))))) )),
  inference(forward_demodulation,[],[f7773,f7723])).
fof(f7773,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(divide(X47,divide(divide(inverse(X46),divide(X48,X46)),divide(X49,X48))),divide(divide(X50,X51),divide(X52,X51))))) )),
  inference(backward_demodulation,[],[f7723,f7486])).
fof(f7486,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(divide(divide(X47,divide(divide(X46,X48),divide(X49,X48))),inverse(X46)),divide(divide(X50,X51),divide(X52,X51))))) )),
  inference(forward_demodulation,[],[f6731,f6721])).
fof(f6731,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (divide(divide(inverse(X52),X49),X47) = divide(inverse(X50),divide(divide(divide(X47,divide(divide(X46,X48),divide(X49,X48))),inverse(X46)),divide(inverse(divide(X51,X50)),divide(X52,X51))))) )),
  inference(backward_demodulation,[],[f6721,f717])).
fof(f717,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (divide(inverse(X50),divide(divide(divide(X47,divide(inverse(divide(X48,X46)),divide(X49,X48))),inverse(X46)),divide(inverse(divide(X51,X50)),divide(X52,X51)))) = divide(divide(inverse(X52),X49),X47)) )),
  inference(superposition,[],[f379,f166])).
fof(f7946,plain,(
  ( ! [X28,X26,X27,X25] : (divide(divide(X27,divide(X25,X26)),X28) = divide(X27,divide(X28,divide(X26,X25)))) )),
  inference(forward_demodulation,[],[f7733,f7709])).
fof(f7733,plain,(
  ( ! [X30,X28,X26,X27,X25] : (divide(divide(X27,divide(X25,X26)),X28) = divide(divide(X27,X30),divide(divide(X28,divide(X26,X25)),X30))) )),
  inference(backward_demodulation,[],[f7709,f7504])).
fof(f7504,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (divide(divide(X27,X30),divide(divide(divide(X28,X29),divide(divide(X26,X25),X29)),X30)) = divide(divide(X27,divide(X25,X26)),X28)) )),
  inference(forward_demodulation,[],[f7503,f6721])).
fof(f7503,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (divide(inverse(divide(divide(X25,X26),X27)),X28) = divide(divide(X27,X30),divide(divide(divide(X28,X29),divide(divide(X26,X25),X29)),X30))) )),
  inference(forward_demodulation,[],[f6747,f6721])).
fof(f6747,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (divide(inverse(divide(divide(X25,X26),X27)),X28) = divide(inverse(divide(X30,X27)),divide(divide(divide(X28,X29),divide(divide(X26,X25),X29)),X30))) )),
  inference(backward_demodulation,[],[f6721,f1364])).
fof(f1364,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (divide(inverse(divide(divide(X25,X26),X27)),X28) = divide(inverse(divide(X30,X27)),divide(divide(inverse(divide(X29,X28)),divide(divide(X26,X25),X29)),X30))) )),
  inference(superposition,[],[f968,f7])).
fof(f8057,plain,(
  ( ! [X109,X107] : (divide(inverse(inverse(X107)),X109) = divide(X107,inverse(inverse(X109)))) )),
  inference(backward_demodulation,[],[f8056,f8027])).
fof(f8056,plain,(
  ( ! [X8,X7,X9] : (divide(divide(X8,X7),divide(X9,X7)) = divide(X8,inverse(inverse(X9)))) )),
  inference(forward_demodulation,[],[f8016,f6761])).
fof(f6761,plain,(
  ( ! [X61,X59,X62,X58] : (divide(X61,X62) = divide(divide(X59,X58),divide(divide(X62,X61),divide(X58,X59)))) )),
  inference(backward_demodulation,[],[f6721,f2344])).
fof(f2344,plain,(
  ( ! [X61,X59,X62,X58] : (divide(X61,X62) = divide(divide(X59,X58),divide(inverse(divide(X61,X62)),divide(X58,X59)))) )),
  inference(superposition,[],[f49,f2013])).
fof(f8016,plain,(
  ( ! [X14,X8,X7,X13,X9] : (divide(divide(X8,X7),divide(X9,X7)) = divide(divide(X13,X14),divide(divide(inverse(inverse(X9)),X8),divide(X14,X13)))) )),
  inference(backward_demodulation,[],[f8006,f7607])).
fof(f7607,plain,(
  ( ! [X14,X8,X7,X15,X13,X9] : (divide(divide(X8,X7),divide(X9,X7)) = divide(divide(X13,X14),divide(inverse(X15),divide(divide(X14,X13),divide(divide(X15,inverse(X9)),X8))))) )),
  inference(forward_demodulation,[],[f6818,f6721])).
fof(f6818,plain,(
  ( ! [X14,X8,X7,X15,X13,X9] : (divide(divide(X13,X14),divide(inverse(X15),divide(divide(X14,X13),divide(inverse(divide(inverse(X9),X15)),X8)))) = divide(divide(X8,X7),divide(X9,X7))) )),
  inference(backward_demodulation,[],[f6721,f4344])).
fof(f4344,plain,(
  ( ! [X14,X8,X7,X15,X13,X9] : (divide(inverse(divide(X7,X8)),divide(X9,X7)) = divide(divide(X13,X14),divide(inverse(X15),divide(divide(X14,X13),divide(inverse(divide(inverse(X9),X15)),X8))))) )),
  inference(backward_demodulation,[],[f4271,f173])).
fof(f173,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (divide(inverse(divide(X7,X8)),divide(X9,X7)) = divide(divide(X13,X14),divide(inverse(X15),divide(divide(X14,X13),divide(inverse(divide(divide(divide(X10,X11),divide(inverse(divide(X12,X9)),divide(divide(X11,X10),X12))),X15)),X8))))) )),
  inference(superposition,[],[f49,f6])).
fof(f10273,plain,(
  ( ! [X8,X9] : (divide(X8,X9) = divide(X8,inverse(inverse(X9)))) )),
  inference(backward_demodulation,[],[f10272,f8056])).
fof(f10272,plain,(
  ( ! [X94,X88,X93] : (divide(divide(X88,X93),divide(X94,X93)) = divide(X88,X94)) )),
  inference(forward_demodulation,[],[f10271,f4174])).
fof(f10271,plain,(
  ( ! [X94,X88,X97,X93,X96] : (divide(divide(X88,X93),divide(X94,X93)) = divide(divide(divide(X88,X94),divide(X96,X97)),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f10270,f7816])).
fof(f10270,plain,(
  ( ! [X94,X88,X97,X87,X93,X96,X86] : (divide(divide(X88,X93),divide(X94,X93)) = divide(divide(divide(X88,divide(X86,X87)),divide(divide(X96,X97),divide(divide(X86,X87),X94))),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f10269,f7946])).
fof(f10269,plain,(
  ( ! [X94,X88,X97,X87,X93,X96,X86] : (divide(divide(X88,X93),divide(X94,X93)) = divide(divide(divide(divide(X88,divide(X86,X87)),divide(X94,divide(X86,X87))),divide(X96,X97)),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f10268,f7946])).
fof(f10268,plain,(
  ( ! [X94,X88,X97,X87,X93,X96,X86] : (divide(divide(X88,X93),divide(X94,X93)) = divide(divide(divide(divide(divide(X88,divide(X86,X87)),divide(X87,X86)),X94),divide(X96,X97)),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f10267,f6721])).
fof(f10267,plain,(
  ( ! [X94,X88,X97,X87,X93,X96,X86] : (divide(divide(X88,X93),divide(X94,X93)) = divide(divide(divide(inverse(divide(divide(X87,X86),divide(X88,divide(X86,X87)))),X94),divide(X96,X97)),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f10266,f8006])).
fof(f10266,plain,(
  ( ! [X94,X88,X97,X87,X95,X93,X96,X86] : (divide(divide(X88,X93),divide(X94,X93)) = divide(divide(inverse(X95),divide(divide(X96,X97),divide(divide(X95,divide(divide(X87,X86),divide(X88,divide(X86,X87)))),X94))),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f7147,f6721])).
fof(f7147,plain,(
  ( ! [X94,X88,X97,X87,X95,X93,X96,X86] : (divide(divide(X88,X93),divide(X94,X93)) = divide(divide(inverse(X95),divide(divide(X96,X97),divide(inverse(divide(divide(divide(X87,X86),divide(X88,divide(X86,X87))),X95)),X94))),divide(X97,X96))) )),
  inference(backward_demodulation,[],[f6721,f6684])).
fof(f6684,plain,(
  ( ! [X94,X88,X97,X87,X95,X93,X96,X86] : (divide(divide(X88,X93),inverse(divide(X93,X94))) = divide(divide(inverse(X95),divide(divide(X96,X97),divide(inverse(divide(divide(divide(X87,X86),divide(X88,divide(X86,X87))),X95)),X94))),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f6683,f1633])).
fof(f6683,plain,(
  ( ! [X94,X90,X88,X97,X87,X95,X93,X89,X96,X86] : (divide(divide(X88,X93),inverse(divide(X93,X94))) = divide(divide(inverse(X95),divide(divide(X96,X97),divide(inverse(divide(divide(divide(X87,X86),divide(inverse(divide(divide(X89,X90),divide(X88,divide(X86,X87)))),divide(X90,X89))),X95)),X94))),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f6682,f6259])).
fof(f6259,plain,(
  ( ! [X158,X154,X159,X157,X153] : (divide(divide(divide(X158,X159),divide(X153,X154)),inverse(X157)) = divide(divide(X158,X159),divide(inverse(X157),divide(X154,X153)))) )),
  inference(backward_demodulation,[],[f5503,f4010])).
fof(f4010,plain,(
  ( ! [X158,X154,X161,X159,X157,X153,X162,X160] : (divide(divide(divide(X158,X159),divide(X153,X154)),inverse(X157)) = divide(divide(inverse(X160),divide(divide(X161,X162),divide(inverse(divide(divide(X159,X158),X160)),divide(inverse(X157),divide(X154,X153))))),divide(X162,X161))) )),
  inference(backward_demodulation,[],[f3467,f2609])).
fof(f2609,plain,(
  ( ! [X158,X156,X154,X161,X159,X157,X155,X153,X162,X160] : (divide(divide(inverse(X160),divide(divide(X161,X162),divide(inverse(divide(divide(X159,X158),X160)),divide(inverse(X155),divide(divide(X154,X153),divide(inverse(divide(X156,X155)),divide(X157,X156))))))),divide(X162,X161)) = divide(divide(divide(X158,X159),divide(X153,X154)),inverse(X157))) )),
  inference(superposition,[],[f21,f166])).
fof(f3467,plain,(
  ( ! [X80,X78,X76,X79,X77] : (divide(inverse(X78),divide(divide(X77,X76),divide(inverse(divide(X79,X78)),divide(X80,X79)))) = divide(inverse(X80),divide(X77,X76))) )),
  inference(superposition,[],[f1633,f166])).
fof(f5503,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (divide(divide(inverse(X38),divide(divide(X39,X40),divide(inverse(divide(divide(X37,X36),X38)),divide(X34,X35)))),divide(X40,X39)) = divide(divide(X36,X37),divide(X34,X35))) )),
  inference(backward_demodulation,[],[f4950,f2920])).
fof(f6682,plain,(
  ( ! [X94,X90,X88,X97,X87,X95,X93,X89,X96,X86] : (divide(divide(X88,X93),inverse(divide(X93,X94))) = divide(divide(inverse(X95),divide(divide(X96,X97),divide(inverse(divide(divide(divide(divide(X87,X86),divide(X89,X90)),inverse(divide(divide(X89,X90),divide(X88,divide(X86,X87))))),X95)),X94))),divide(X97,X96))) )),
  inference(forward_demodulation,[],[f5529,f6407])).
fof(f6407,plain,(
  ( ! [X466,X469,X465,X467,X468] : (divide(inverse(divide(divide(X465,X466),X467)),divide(X468,X469)) = inverse(divide(divide(X468,X469),divide(X467,divide(X465,X466))))) )),
  inference(forward_demodulation,[],[f5515,f1478])).
fof(f1478,plain,(
  ( ! [X476,X478,X480,X475,X477,X479,X481] : (inverse(divide(X475,divide(X476,divide(X477,X478)))) = divide(divide(X479,X480),divide(inverse(X481),divide(divide(X480,X479),divide(inverse(divide(divide(divide(X478,X477),X475),X481)),X476))))) )),
  inference(superposition,[],[f166,f968])).
fof(f5515,plain,(
  ( ! [X466,X472,X474,X469,X465,X467,X473,X468] : (divide(inverse(divide(divide(X465,X466),X467)),divide(X468,X469)) = divide(divide(X472,X473),divide(inverse(X474),divide(divide(X473,X472),divide(inverse(divide(divide(divide(X466,X465),divide(X468,X469)),X474)),X467))))) )),
  inference(backward_demodulation,[],[f4950,f2967])).
fof(f2967,plain,(
  ( ! [X466,X472,X474,X469,X465,X467,X473,X468] : (divide(inverse(divide(divide(X465,X466),X467)),divide(X468,X469)) = divide(divide(X472,X473),divide(inverse(X474),divide(divide(X473,X472),divide(inverse(divide(divide(divide(X466,X465),inverse(divide(X469,X468))),X474)),X467))))) )),
  inference(backward_demodulation,[],[f2921,f1300])).
fof(f1300,plain,(
  ( ! [X466,X472,X474,X469,X471,X465,X467,X473,X468,X470] : (divide(inverse(divide(divide(X465,X466),X467)),divide(X468,X469)) = divide(divide(X472,X473),divide(inverse(X474),divide(divide(X473,X472),divide(inverse(divide(divide(divide(divide(X466,X465),divide(inverse(divide(X470,X471)),divide(divide(X469,X468),X470))),inverse(X471)),X474)),X467))))) )),
  inference(superposition,[],[f166,f17])).
fof(f17,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (divide(divide(inverse(divide(divide(X4,X5),X6)),divide(X3,X2)),divide(divide(divide(X5,X4),divide(inverse(divide(X0,X1)),divide(divide(X2,X3),X0))),inverse(X1))) = X6) )),
  inference(superposition,[],[f7,f1])).
fof(f5529,plain,(
  ( ! [X94,X90,X88,X97,X87,X95,X93,X89,X96,X86] : (divide(divide(X88,X93),inverse(divide(X93,X94))) = divide(divide(inverse(X95),divide(divide(X96,X97),divide(inverse(divide(divide(divide(divide(X87,X86),divide(X89,X90)),divide(inverse(divide(divide(X86,X87),X88)),divide(X89,X90))),X95)),X94))),divide(X97,X96))) )),
  inference(backward_demodulation,[],[f4950,f3015])).
fof(f3015,plain,(
  ( ! [X94,X90,X88,X97,X87,X95,X93,X89,X96,X86] : (divide(divide(X88,X93),inverse(divide(X93,X94))) = divide(divide(inverse(X95),divide(divide(X96,X97),divide(inverse(divide(divide(divide(divide(X87,X86),inverse(divide(X90,X89))),divide(inverse(divide(divide(X86,X87),X88)),divide(X89,X90))),X95)),X94))),divide(X97,X96))) )),
  inference(backward_demodulation,[],[f2921,f2563])).
fof(f2563,plain,(
  ( ! [X94,X92,X90,X88,X97,X87,X95,X93,X91,X89,X96,X86] : (divide(divide(inverse(X95),divide(divide(X96,X97),divide(inverse(divide(divide(divide(divide(divide(X87,X86),divide(inverse(divide(X91,X92)),divide(divide(X90,X89),X91))),inverse(X92)),divide(inverse(divide(divide(X86,X87),X88)),divide(X89,X90))),X95)),X94))),divide(X97,X96)) = divide(divide(X88,X93),inverse(divide(X93,X94)))) )),
  inference(superposition,[],[f21,f17])).
fof(f4,plain,(
  a2 != divide(divide(inverse(b2),inverse(b2)),inverse(a2))),
  inference(definition_unfolding,[],[f3,f2,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (multiply(X0,X1) = divide(X0,inverse(X1))) )),
  file('GRP473-1.tptp',unknown)).
fof(f3,axiom,(
  a2 != multiply(multiply(inverse(b2),b2),a2)),
  file('GRP473-1.tptp',unknown)).
% SZS output end Proof for GRP473-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 27121
% Time elapsed: 0.597 s
% ------------------------------
% ------------------------------
