% Refutation found. Thanks to Tanya!
% SZS status Theorem for GRP621+1
% SZS output start Proof for GRP621+1
fof(f55522,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f470,f475,f575,f580,f585,f590,f595,f600,f605,f610,f619,f660,f5102,f5104,f54050,f54296,f55521])).
fof(f55521,plain,(
  spl28_3 | ~spl28_5 | ~spl28_6 | spl28_12 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16 | ~spl28_221),
  inference(avatar_contradiction_clause,[],[f55520])).
fof(f55520,plain,(
  $false | (spl28_3 | ~spl28_5 | ~spl28_6 | spl28_12 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16 | ~spl28_221)),
  inference(subsumption_resolution,[],[f55515,f4847])).
fof(f4847,plain,(
  ~m1_subset_1(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | (spl28_3 | ~spl28_5 | ~spl28_6 | spl28_12 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(backward_demodulation,[],[f4845,f579])).
fof(f579,plain,(
  ~m1_subset_1(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2),k1_autgroup(sK1)) | spl28_12),
  inference(avatar_component_clause,[],[f577])).
fof(f577,plain,(
  spl28_12 <=> m1_subset_1(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2),k1_autgroup(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_12])])).
fof(f4845,plain,(
  k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2) = k5_relat_1(sK3,sK2) | (spl28_3 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f4844,f474])).
fof(f474,plain,(
  v1_funct_1(sK3) | ~spl28_6),
  inference(avatar_component_clause,[],[f472])).
fof(f472,plain,(
  spl28_6 <=> v1_funct_1(sK3)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_6])])).
fof(f4844,plain,(
  k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2) = k5_relat_1(sK3,sK2) | ~v1_funct_1(sK3) | (spl28_3 | ~spl28_5 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f4820,f913])).
fof(f913,plain,(
  m1_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~spl28_13),
  inference(resolution,[],[f584,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (~m2_relset_1(X2,X0,X1) | m1_relset_1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((m2_relset_1(X2,X0,X1) | ~m1_relset_1(X2,X0,X1)) & (m1_relset_1(X2,X0,X1) | ~m2_relset_1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,axiom,(
  ! [X0,X1,X2] : (m2_relset_1(X2,X0,X1) <=> m1_relset_1(X2,X0,X1))),
  file('GRP621+1.tptp',unknown)).
fof(f584,plain,(
  m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~spl28_13),
  inference(avatar_component_clause,[],[f582])).
fof(f582,plain,(
  spl28_13 <=> m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_13])])).
fof(f4820,plain,(
  ~m1_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2) = k5_relat_1(sK3,sK2) | ~v1_funct_1(sK3) | (spl28_3 | ~spl28_5 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(resolution,[],[f1137,f594])).
fof(f594,plain,(
  v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~spl28_15),
  inference(avatar_component_clause,[],[f592])).
fof(f592,plain,(
  spl28_15 <=> v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_15])])).
fof(f1137,plain,(
  ( ! [X2,X3] : (~v1_funct_2(X3,X2,u1_struct_0(sK1)) | ~m1_relset_1(X3,X2,u1_struct_0(sK1)) | k7_funct_2(X2,u1_struct_0(sK1),u1_struct_0(sK1),X3,sK2) = k5_relat_1(X3,sK2) | ~v1_funct_1(X3)) ) | (spl28_3 | ~spl28_5 | ~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1136,f460])).
fof(f460,plain,(
  ~v1_xboole_0(u1_struct_0(sK1)) | spl28_3),
  inference(avatar_component_clause,[],[f459])).
fof(f459,plain,(
  spl28_3 <=> v1_xboole_0(u1_struct_0(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_3])])).
fof(f1136,plain,(
  ( ! [X2,X3] : (k7_funct_2(X2,u1_struct_0(sK1),u1_struct_0(sK1),X3,sK2) = k5_relat_1(X3,sK2) | ~m1_relset_1(X3,X2,u1_struct_0(sK1)) | ~v1_funct_2(X3,X2,u1_struct_0(sK1)) | ~v1_funct_1(X3) | v1_xboole_0(u1_struct_0(sK1))) ) | (~spl28_5 | ~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1135,f469])).
fof(f469,plain,(
  v1_funct_1(sK2) | ~spl28_5),
  inference(avatar_component_clause,[],[f467])).
fof(f467,plain,(
  spl28_5 <=> v1_funct_1(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_5])])).
fof(f1135,plain,(
  ( ! [X2,X3] : (k7_funct_2(X2,u1_struct_0(sK1),u1_struct_0(sK1),X3,sK2) = k5_relat_1(X3,sK2) | ~v1_funct_1(sK2) | ~m1_relset_1(X3,X2,u1_struct_0(sK1)) | ~v1_funct_2(X3,X2,u1_struct_0(sK1)) | ~v1_funct_1(X3) | v1_xboole_0(u1_struct_0(sK1))) ) | (~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1123,f984])).
fof(f984,plain,(
  m1_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~spl28_14),
  inference(resolution,[],[f589,f339])).
fof(f589,plain,(
  m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~spl28_14),
  inference(avatar_component_clause,[],[f587])).
fof(f587,plain,(
  spl28_14 <=> m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_14])])).
fof(f1123,plain,(
  ( ! [X2,X3] : (~m1_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | k7_funct_2(X2,u1_struct_0(sK1),u1_struct_0(sK1),X3,sK2) = k5_relat_1(X3,sK2) | ~v1_funct_1(sK2) | ~m1_relset_1(X3,X2,u1_struct_0(sK1)) | ~v1_funct_2(X3,X2,u1_struct_0(sK1)) | ~v1_funct_1(X3) | v1_xboole_0(u1_struct_0(sK1))) ) | ~spl28_16),
  inference(resolution,[],[f599,f335])).
fof(f335,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~v1_funct_2(X4,X1,X2) | ~m1_relset_1(X4,X1,X2) | k7_funct_2(X0,X1,X2,X3,X4) = k5_relat_1(X3,X4) | ~v1_funct_1(X4) | ~m1_relset_1(X3,X0,X1) | ~v1_funct_2(X3,X0,X1) | ~v1_funct_1(X3) | v1_xboole_0(X1)) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4] : (k7_funct_2(X0,X1,X2,X3,X4) = k5_relat_1(X3,X4) | ~m1_relset_1(X4,X1,X2) | ~v1_funct_2(X4,X1,X2) | ~v1_funct_1(X4) | ~m1_relset_1(X3,X0,X1) | ~v1_funct_2(X3,X0,X1) | ~v1_funct_1(X3) | v1_xboole_0(X1))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : (k7_funct_2(X0,X1,X2,X3,X4) = k5_relat_1(X3,X4) | (~m1_relset_1(X4,X1,X2) | ~v1_funct_2(X4,X1,X2) | ~v1_funct_1(X4) | ~m1_relset_1(X3,X0,X1) | ~v1_funct_2(X3,X0,X1) | ~v1_funct_1(X3) | v1_xboole_0(X1)))),
  inference(ennf_transformation,[],[f58])).
fof(f58,axiom,(
  ! [X0,X1,X2,X3,X4] : ((m1_relset_1(X4,X1,X2) & v1_funct_2(X4,X1,X2) & v1_funct_1(X4) & m1_relset_1(X3,X0,X1) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3) & ~v1_xboole_0(X1)) => k7_funct_2(X0,X1,X2,X3,X4) = k5_relat_1(X3,X4))),
  file('GRP621+1.tptp',unknown)).
fof(f599,plain,(
  v1_funct_2(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~spl28_16),
  inference(avatar_component_clause,[],[f597])).
fof(f597,plain,(
  spl28_16 <=> v1_funct_2(sK2,u1_struct_0(sK1),u1_struct_0(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_16])])).
fof(f55515,plain,(
  m1_subset_1(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~spl28_221),
  inference(resolution,[],[f8666,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (~r2_hidden(X0,X1) | m1_subset_1(X0,X1)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1))),
  inference(ennf_transformation,[],[f63])).
fof(f63,axiom,(
  ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1))),
  file('GRP621+1.tptp',unknown)).
fof(f8666,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~spl28_221),
  inference(avatar_component_clause,[],[f8664])).
fof(f8664,plain,(
  spl28_221 <=> r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_221])])).
fof(f54296,plain,(
  spl28_221 | ~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16 | ~spl28_21 | ~spl28_163),
  inference(avatar_split_clause,[],[f54295,f5090,f653,f597,f592,f587,f582,f472,f467,f463,f8664])).
fof(f463,plain,(
  spl28_4 <=> m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_4])])).
fof(f653,plain,(
  spl28_21 <=> r2_hidden(sK3,k1_autgroup(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_21])])).
fof(f5090,plain,(
  spl28_163 <=> v4_group_6(sK2,sK1,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_163])])).
fof(f54295,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16 | ~spl28_21 | ~spl28_163)),
  inference(subsumption_resolution,[],[f54294,f584])).
fof(f54294,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16 | ~spl28_21 | ~spl28_163)),
  inference(subsumption_resolution,[],[f54293,f589])).
fof(f54293,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16 | ~spl28_21 | ~spl28_163)),
  inference(subsumption_resolution,[],[f54292,f1045])).
fof(f1045,plain,(
  v4_group_6(sK3,sK1,sK1) | (~spl28_4 | ~spl28_6 | ~spl28_13 | ~spl28_15 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1044,f211])).
fof(f211,plain,(
  ~v3_struct_0(sK1)),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ((~m2_fraenkel(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2),u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1)) & m2_fraenkel(sK3,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))) & m2_fraenkel(sK2,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))) & l1_group_1(sK1) & v4_group_1(sK1) & v3_group_1(sK1) & v1_group_1(sK1) & ~v3_struct_0(sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3])],[f79,f152,f151,f150])).
fof(f150,plain,(
  ? [X0] : (? [X1] : (? [X2] : (~m2_fraenkel(k7_funct_2(u1_struct_0(X0),u1_struct_0(X0),u1_struct_0(X0),X2,X1),u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)) & m2_fraenkel(X2,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0))) & m2_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0))) & l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0)) => (? [X1] : (? [X2] : (~m2_fraenkel(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),X2,X1),u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1)) & m2_fraenkel(X2,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))) & m2_fraenkel(X1,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))) & l1_group_1(sK1) & v4_group_1(sK1) & v3_group_1(sK1) & v1_group_1(sK1) & ~v3_struct_0(sK1))),
  introduced(choice_axiom,[])).
fof(f151,plain,(
  ? [X1] : (? [X2] : (~m2_fraenkel(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),X2,X1),u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1)) & m2_fraenkel(X2,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))) & m2_fraenkel(X1,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))) => (? [X2] : (~m2_fraenkel(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),X2,sK2),u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1)) & m2_fraenkel(X2,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))) & m2_fraenkel(sK2,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1)))),
  introduced(choice_axiom,[])).
fof(f152,plain,(
  ? [X2] : (~m2_fraenkel(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),X2,sK2),u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1)) & m2_fraenkel(X2,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))) => (~m2_fraenkel(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2),u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1)) & m2_fraenkel(sK3,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1)))),
  introduced(choice_axiom,[])).
fof(f79,plain,(
  ? [X0] : (? [X1] : (? [X2] : (~m2_fraenkel(k7_funct_2(u1_struct_0(X0),u1_struct_0(X0),u1_struct_0(X0),X2,X1),u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)) & m2_fraenkel(X2,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0))) & m2_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0))) & l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ? [X0] : (? [X1] : (? [X2] : (~m2_fraenkel(k7_funct_2(u1_struct_0(X0),u1_struct_0(X0),u1_struct_0(X0),X2,X1),u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)) & m2_fraenkel(X2,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0))) & m2_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0))) & (l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f2])).
fof(f2,negated_conjecture,(
  ~! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : (m2_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)) => ! [X2] : (m2_fraenkel(X2,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)) => m2_fraenkel(k7_funct_2(u1_struct_0(X0),u1_struct_0(X0),u1_struct_0(X0),X2,X1),u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)))))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : (m2_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)) => ! [X2] : (m2_fraenkel(X2,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)) => m2_fraenkel(k7_funct_2(u1_struct_0(X0),u1_struct_0(X0),u1_struct_0(X0),X2,X1),u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)))))),
  file('GRP621+1.tptp',unknown)).
fof(f1044,plain,(
  v4_group_6(sK3,sK1,sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_6 | ~spl28_13 | ~spl28_15 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1043,f212])).
fof(f212,plain,(
  v1_group_1(sK1)),
  inference(cnf_transformation,[],[f153])).
fof(f1043,plain,(
  v4_group_6(sK3,sK1,sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_6 | ~spl28_13 | ~spl28_15 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1042,f213])).
fof(f213,plain,(
  v3_group_1(sK1)),
  inference(cnf_transformation,[],[f153])).
fof(f1042,plain,(
  v4_group_6(sK3,sK1,sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_6 | ~spl28_13 | ~spl28_15 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1041,f214])).
fof(f214,plain,(
  v4_group_1(sK1)),
  inference(cnf_transformation,[],[f153])).
fof(f1041,plain,(
  v4_group_6(sK3,sK1,sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_6 | ~spl28_13 | ~spl28_15 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1040,f215])).
fof(f215,plain,(
  l1_group_1(sK1)),
  inference(cnf_transformation,[],[f153])).
fof(f1040,plain,(
  v4_group_6(sK3,sK1,sK1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_6 | ~spl28_13 | ~spl28_15 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1039,f474])).
fof(f1039,plain,(
  v4_group_6(sK3,sK1,sK1) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_13 | ~spl28_15 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1038,f594])).
fof(f1038,plain,(
  v4_group_6(sK3,sK1,sK1) | ~v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_13 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1037,f886])).
fof(f886,plain,(
  v1_group_6(sK3,sK1,sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f885,f211])).
fof(f885,plain,(
  v1_group_6(sK3,sK1,sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f884,f212])).
fof(f884,plain,(
  v1_group_6(sK3,sK1,sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f883,f213])).
fof(f883,plain,(
  v1_group_6(sK3,sK1,sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f882,f214])).
fof(f882,plain,(
  v1_group_6(sK3,sK1,sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f881,f215])).
fof(f881,plain,(
  v1_group_6(sK3,sK1,sK1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f870,f464])).
fof(f464,plain,(
  m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | ~spl28_4),
  inference(avatar_component_clause,[],[f463])).
fof(f870,plain,(
  v1_group_6(sK3,sK1,sK1) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1)),
  inference(resolution,[],[f355,f217])).
fof(f217,plain,(
  m2_fraenkel(sK3,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))),
  inference(cnf_transformation,[],[f153])).
fof(f355,plain,(
  ( ! [X0,X3] : (~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),k1_autgroup(X0)) | v1_group_6(X3,X0,X0) | ~m1_fraenkel(k1_autgroup(X0),u1_struct_0(X0),u1_struct_0(X0)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0)) )),
  inference(equality_resolution,[],[f250])).
fof(f250,plain,(
  ( ! [X0,X3,X1] : (v1_group_6(X3,X0,X0) | ~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1) | k1_autgroup(X0) != X1 | ~m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : (! [X1] : (((k1_autgroup(X0) = X1 | ~sP0(X0,X1) | ((~m2_relset_1(sK5(X0,X1),u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(sK5(X0,X1),X0,X0) | ~v1_funct_2(sK5(X0,X1),u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(sK5(X0,X1))) & m2_fraenkel(sK5(X0,X1),u1_struct_0(X0),u1_struct_0(X0),X1))) & ((sP0(X0,X1) & ! [X3] : ((m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X3,X0,X0) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X3)) | ~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1))) | k1_autgroup(X0) != X1)) | ~m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0))) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f161,f162])).
fof(f162,plain,(
  ! [X1,X0] : (? [X2] : ((~m2_relset_1(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X2,X0,X0) | ~v1_funct_2(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X2)) & m2_fraenkel(X2,u1_struct_0(X0),u1_struct_0(X0),X1)) => ((~m2_relset_1(sK5(X0,X1),u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(sK5(X0,X1),X0,X0) | ~v1_funct_2(sK5(X0,X1),u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(sK5(X0,X1))) & m2_fraenkel(sK5(X0,X1),u1_struct_0(X0),u1_struct_0(X0),X1)))),
  introduced(choice_axiom,[])).
fof(f161,plain,(
  ! [X0] : (! [X1] : (((k1_autgroup(X0) = X1 | ~sP0(X0,X1) | ? [X2] : ((~m2_relset_1(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X2,X0,X0) | ~v1_funct_2(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X2)) & m2_fraenkel(X2,u1_struct_0(X0),u1_struct_0(X0),X1))) & ((sP0(X0,X1) & ! [X3] : ((m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X3,X0,X0) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X3)) | ~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1))) | k1_autgroup(X0) != X1)) | ~m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0))) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(rectify,[],[f160])).
fof(f160,plain,(
  ! [X0] : (! [X1] : (((k1_autgroup(X0) = X1 | ~sP0(X0,X1) | ? [X3] : ((~m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X3,X0,X0) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X3)) & m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1))) & ((sP0(X0,X1) & ! [X3] : ((m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X3,X0,X0) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X3)) | ~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1))) | k1_autgroup(X0) != X1)) | ~m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0))) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : (! [X1] : (((k1_autgroup(X0) = X1 | (~sP0(X0,X1) | ? [X3] : ((~m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X3,X0,X0) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X3)) & m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1)))) & ((sP0(X0,X1) & ! [X3] : ((m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X3,X0,X0) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X3)) | ~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1))) | k1_autgroup(X0) != X1)) | ~m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0))) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0] : (! [X1] : ((k1_autgroup(X0) = X1 <=> (sP0(X0,X1) & ! [X3] : ((m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X3,X0,X0) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X3)) | ~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1)))) | ~m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0))) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(definition_folding,[],[f97,f148])).
fof(f148,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ! [X2] : ((r2_hidden(X2,X1) <=> (v3_group_6(X2,X0,X0) & v2_funct_1(X2))) | ~m2_relset_1(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X2,X0,X0) | ~v1_funct_2(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X2)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f97,plain,(
  ! [X0] : (! [X1] : ((k1_autgroup(X0) = X1 <=> (! [X2] : ((r2_hidden(X2,X1) <=> (v3_group_6(X2,X0,X0) & v2_funct_1(X2))) | ~m2_relset_1(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X2,X0,X0) | ~v1_funct_2(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X2)) & ! [X3] : ((m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X3,X0,X0) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X3)) | ~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1)))) | ~m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0))) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : (! [X1] : ((k1_autgroup(X0) = X1 <=> (! [X2] : ((r2_hidden(X2,X1) <=> (v3_group_6(X2,X0,X0) & v2_funct_1(X2))) | (~m2_relset_1(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X2,X0,X0) | ~v1_funct_2(X2,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X2))) & ! [X3] : ((m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X3,X0,X0) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X3)) | ~m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1)))) | ~m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0))) | (~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : (m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0)) => (k1_autgroup(X0) = X1 <=> (! [X2] : ((m2_relset_1(X2,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X2,X0,X0) & v1_funct_2(X2,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X2)) => (r2_hidden(X2,X1) <=> (v3_group_6(X2,X0,X0) & v2_funct_1(X2)))) & ! [X3] : (m2_fraenkel(X3,u1_struct_0(X0),u1_struct_0(X0),X1) => (m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X3,X0,X0) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X3)))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : (m1_fraenkel(X1,u1_struct_0(X0),u1_struct_0(X0)) => (k1_autgroup(X0) = X1 <=> (! [X2] : ((m2_relset_1(X2,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X2,X0,X0) & v1_funct_2(X2,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X2)) => (r2_hidden(X2,X1) <=> (v3_group_6(X2,X0,X0) & v2_funct_1(X2)))) & ! [X2] : (m2_fraenkel(X2,u1_struct_0(X0),u1_struct_0(X0),X1) => (m2_relset_1(X2,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X2,X0,X0) & v1_funct_2(X2,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X2)))))))),
  file('GRP621+1.tptp',unknown)).
fof(f1037,plain,(
  v4_group_6(sK3,sK1,sK1) | ~v1_group_6(sK3,sK1,sK1) | ~v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | (~spl28_13 | ~spl28_21)),
  inference(subsumption_resolution,[],[f1027,f655])).
fof(f655,plain,(
  r2_hidden(sK3,k1_autgroup(sK1)) | ~spl28_21),
  inference(avatar_component_clause,[],[f653])).
fof(f1027,plain,(
  ~r2_hidden(sK3,k1_autgroup(sK1)) | v4_group_6(sK3,sK1,sK1) | ~v1_group_6(sK3,sK1,sK1) | ~v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_13),
  inference(resolution,[],[f346,f584])).
fof(f346,plain,(
  ( ! [X0,X1] : (~m2_relset_1(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~r2_hidden(X1,k1_autgroup(X0)) | v4_group_6(X1,X0,X0) | ~v1_group_6(X1,X0,X0) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0] : (! [X1] : (((r2_hidden(X1,k1_autgroup(X0)) | ~v4_group_6(X1,X0,X0)) & (v4_group_6(X1,X0,X0) | ~r2_hidden(X1,k1_autgroup(X0)))) | ~m2_relset_1(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X1,X0,X0) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X1)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : (! [X1] : ((r2_hidden(X1,k1_autgroup(X0)) <=> v4_group_6(X1,X0,X0)) | ~m2_relset_1(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X1,X0,X0) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X1)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : (! [X1] : ((r2_hidden(X1,k1_autgroup(X0)) <=> v4_group_6(X1,X0,X0)) | (~m2_relset_1(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_group_6(X1,X0,X0) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X1))) | (~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : ((m2_relset_1(X1,u1_struct_0(X0),u1_struct_0(X0)) & v1_group_6(X1,X0,X0) & v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(X0)) & v1_funct_1(X1)) => (r2_hidden(X1,k1_autgroup(X0)) <=> v4_group_6(X1,X0,X0))))),
  file('GRP621+1.tptp',unknown)).
fof(f54292,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16 | ~spl28_163)),
  inference(subsumption_resolution,[],[f54291,f5091])).
fof(f5091,plain,(
  v4_group_6(sK2,sK1,sK1) | ~spl28_163),
  inference(avatar_component_clause,[],[f5090])).
fof(f54291,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54290,f212])).
fof(f54290,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54289,f880])).
fof(f880,plain,(
  v1_group_6(sK2,sK1,sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f879,f211])).
fof(f879,plain,(
  v1_group_6(sK2,sK1,sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f878,f212])).
fof(f878,plain,(
  v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f877,f213])).
fof(f877,plain,(
  v1_group_6(sK2,sK1,sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f876,f214])).
fof(f876,plain,(
  v1_group_6(sK2,sK1,sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f875,f215])).
fof(f875,plain,(
  v1_group_6(sK2,sK1,sK1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | ~spl28_4),
  inference(subsumption_resolution,[],[f869,f464])).
fof(f869,plain,(
  v1_group_6(sK2,sK1,sK1) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1)),
  inference(resolution,[],[f355,f216])).
fof(f216,plain,(
  m2_fraenkel(sK2,u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))),
  inference(cnf_transformation,[],[f153])).
fof(f54289,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54288,f211])).
fof(f54288,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54287,f213])).
fof(f54287,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54286,f214])).
fof(f54286,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54285,f215])).
fof(f54285,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54284,f474])).
fof(f54284,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54283,f594])).
fof(f54283,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54282,f886])).
fof(f54282,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v1_group_6(sK3,sK1,sK1) | ~v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f54281,f469])).
fof(f54281,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v1_funct_1(sK2) | ~v1_group_6(sK3,sK1,sK1) | ~v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f8433,f599])).
fof(f8433,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v1_funct_2(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK2) | ~v1_group_6(sK3,sK1,sK1) | ~v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(duplicate_literal_removal,[],[f8430])).
fof(f8430,plain,(
  r2_hidden(k5_relat_1(sK3,sK2),k1_autgroup(sK1)) | ~v1_funct_2(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK2) | ~v1_group_6(sK3,sK1,sK1) | ~v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~v1_group_6(sK2,sK1,sK1) | ~v1_group_1(sK1) | ~v4_group_6(sK2,sK1,sK1) | ~v4_group_6(sK3,sK1,sK1) | ~m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(superposition,[],[f2750,f5811])).
fof(f5811,plain,(
  k5_relat_1(sK3,sK2) = k8_group_6(sK1,sK1,sK1,sK3,sK2) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f5810,f211])).
fof(f5810,plain,(
  k5_relat_1(sK3,sK2) = k8_group_6(sK1,sK1,sK1,sK3,sK2) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f5809,f213])).
fof(f5809,plain,(
  k5_relat_1(sK3,sK2) = k8_group_6(sK1,sK1,sK1,sK3,sK2) | ~v3_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f5808,f214])).
fof(f5808,plain,(
  k5_relat_1(sK3,sK2) = k8_group_6(sK1,sK1,sK1,sK3,sK2) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f5807,f215])).
fof(f5807,plain,(
  k5_relat_1(sK3,sK2) = k8_group_6(sK1,sK1,sK1,sK3,sK2) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_5 | ~spl28_6 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f5806,f474])).
fof(f5806,plain,(
  k5_relat_1(sK3,sK2) = k8_group_6(sK1,sK1,sK1,sK3,sK2) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_5 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f5805,f886])).
fof(f5805,plain,(
  ~v1_group_6(sK3,sK1,sK1) | k5_relat_1(sK3,sK2) = k8_group_6(sK1,sK1,sK1,sK3,sK2) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_5 | ~spl28_13 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(subsumption_resolution,[],[f5774,f913])).
fof(f5774,plain,(
  ~m1_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_group_6(sK3,sK1,sK1) | k5_relat_1(sK3,sK2) = k8_group_6(sK1,sK1,sK1,sK3,sK2) | ~v1_funct_1(sK3) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | (~spl28_4 | ~spl28_5 | ~spl28_14 | ~spl28_15 | ~spl28_16)),
  inference(resolution,[],[f1134,f594])).
fof(f1134,plain,(
  ( ! [X0,X1] : (~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~v1_funct_1(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | (~spl28_4 | ~spl28_5 | ~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1133,f211])).
fof(f1133,plain,(
  ( ! [X0,X1] : (k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_funct_1(X1) | v3_struct_0(sK1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | (~spl28_4 | ~spl28_5 | ~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1132,f213])).
fof(f1132,plain,(
  ( ! [X0,X1] : (k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | (~spl28_4 | ~spl28_5 | ~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1131,f214])).
fof(f1131,plain,(
  ( ! [X0,X1] : (k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | (~spl28_4 | ~spl28_5 | ~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1130,f215])).
fof(f1130,plain,(
  ( ! [X0,X1] : (k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | (~spl28_4 | ~spl28_5 | ~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1129,f469])).
fof(f1129,plain,(
  ( ! [X0,X1] : (k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~v1_funct_1(sK2) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | (~spl28_4 | ~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1128,f880])).
fof(f1128,plain,(
  ( ! [X0,X1] : (~v1_group_6(sK2,sK1,sK1) | k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~v1_funct_1(sK2) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | (~spl28_14 | ~spl28_16)),
  inference(subsumption_resolution,[],[f1127,f984])).
fof(f1127,plain,(
  ( ! [X0,X1] : (~m1_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_group_6(sK2,sK1,sK1) | k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~v1_funct_1(sK2) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | ~spl28_16),
  inference(duplicate_literal_removal,[],[f1122])).
fof(f1122,plain,(
  ( ! [X0,X1] : (~m1_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_group_6(sK2,sK1,sK1) | k8_group_6(X0,sK1,sK1,X1,sK2) = k5_relat_1(X1,sK2) | ~v1_funct_1(sK2) | ~m1_relset_1(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_group_6(X1,X0,sK1) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | v3_struct_0(sK1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) ) | ~spl28_16),
  inference(resolution,[],[f599,f336])).
fof(f336,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | k8_group_6(X0,X1,X2,X3,X4) = k5_relat_1(X3,X4) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : (k8_group_6(X0,X1,X2,X3,X4) = k5_relat_1(X3,X4) | ~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : (k8_group_6(X0,X1,X2,X3,X4) = k5_relat_1(X3,X4) | (~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f59])).
fof(f59,axiom,(
  ! [X0,X1,X2,X3,X4] : ((m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) & v1_group_6(X4,X1,X2) & v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) & v1_funct_1(X4) & m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) & v1_group_6(X3,X0,X1) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) & v1_funct_1(X3) & l1_group_1(X2) & v4_group_1(X2) & v3_group_1(X2) & ~v3_struct_0(X2) & l1_group_1(X1) & v4_group_1(X1) & v3_group_1(X1) & ~v3_struct_0(X1) & l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => k8_group_6(X0,X1,X2,X3,X4) = k5_relat_1(X3,X4))),
  file('GRP621+1.tptp',unknown)).
fof(f2750,plain,(
  ( ! [X2,X0,X3,X1] : (r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~v1_group_6(X0,X1,X2) | ~v1_group_1(X2) | ~v4_group_6(X0,X1,X2) | ~v4_group_6(X3,X2,X1) | ~m2_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~m2_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1))) )),
  inference(subsumption_resolution,[],[f2749,f339])).
fof(f2749,plain,(
  ( ! [X2,X0,X3,X1] : (~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~m1_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_group_1(X2) | ~v4_group_6(X0,X1,X2) | ~v4_group_6(X3,X2,X1) | ~m2_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~m2_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1))) )),
  inference(subsumption_resolution,[],[f2748,f339])).
fof(f2748,plain,(
  ( ! [X2,X0,X3,X1] : (~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~m1_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~m1_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_group_1(X2) | ~v4_group_6(X0,X1,X2) | ~v4_group_6(X3,X2,X1) | ~m2_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~m2_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1))) )),
  inference(duplicate_literal_removal,[],[f2747])).
fof(f2747,plain,(
  ( ! [X2,X0,X3,X1] : (~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~m1_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~m1_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_group_1(X2) | ~v4_group_6(X0,X1,X2) | ~v4_group_6(X3,X2,X1) | ~m2_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~m2_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2)) )),
  inference(resolution,[],[f1157,f350])).
fof(f350,plain,(
  ( ! [X4,X2,X0,X3,X1] : (v4_group_6(k8_group_6(X0,X1,X2,X3,X4),X0,X2) | ~v4_group_6(X4,X1,X2) | ~v4_group_6(X3,X0,X1) | ~m2_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (v4_group_6(k8_group_6(X0,X1,X2,X3,X4),X0,X2) | ~v4_group_6(X4,X1,X2) | ~v4_group_6(X3,X0,X1) | ~m2_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4)) | ~m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3)) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2)) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((v4_group_6(k8_group_6(X0,X1,X2,X3,X4),X0,X2) | (~v4_group_6(X4,X1,X2) | ~v4_group_6(X3,X0,X1))) | (~m2_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4))) | (~m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3))) | (~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2))) | (~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1))) | (~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f70])).
fof(f70,axiom,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : ((l1_group_1(X1) & v4_group_1(X1) & v3_group_1(X1) & ~v3_struct_0(X1)) => ! [X2] : ((l1_group_1(X2) & v4_group_1(X2) & v3_group_1(X2) & ~v3_struct_0(X2)) => ! [X3] : ((m2_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) & v1_group_6(X3,X0,X1) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) & v1_funct_1(X3)) => ! [X4] : ((m2_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) & v1_group_6(X4,X1,X2) & v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) & v1_funct_1(X4)) => ((v4_group_6(X4,X1,X2) & v4_group_6(X3,X0,X1)) => v4_group_6(k8_group_6(X0,X1,X2,X3,X4),X0,X2)))))))),
  file('GRP621+1.tptp',unknown)).
fof(f1157,plain,(
  ( ! [X2,X0,X3,X1] : (~v4_group_6(k8_group_6(X2,X1,X2,X3,X0),X2,X2) | ~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~m1_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~m1_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_group_1(X2)) )),
  inference(subsumption_resolution,[],[f1156,f263])).
fof(f263,plain,(
  ( ! [X4,X2,X0,X3,X1] : (v1_funct_2(k8_group_6(X0,X1,X2,X3,X4),u1_struct_0(X0),u1_struct_0(X2)) | ~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4] : ((m2_relset_1(k8_group_6(X0,X1,X2,X3,X4),u1_struct_0(X0),u1_struct_0(X2)) & v1_group_6(k8_group_6(X0,X1,X2,X3,X4),X0,X2) & v1_funct_2(k8_group_6(X0,X1,X2,X3,X4),u1_struct_0(X0),u1_struct_0(X2)) & v1_funct_1(k8_group_6(X0,X1,X2,X3,X4))) | ~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3,X4] : ((m2_relset_1(k8_group_6(X0,X1,X2,X3,X4),u1_struct_0(X0),u1_struct_0(X2)) & v1_group_6(k8_group_6(X0,X1,X2,X3,X4),X0,X2) & v1_funct_2(k8_group_6(X0,X1,X2,X3,X4),u1_struct_0(X0),u1_struct_0(X2)) & v1_funct_1(k8_group_6(X0,X1,X2,X3,X4))) | (~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X4] : ((m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) & v1_group_6(X4,X1,X2) & v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) & v1_funct_1(X4) & m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) & v1_group_6(X3,X0,X1) & v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) & v1_funct_1(X3) & l1_group_1(X2) & v4_group_1(X2) & v3_group_1(X2) & ~v3_struct_0(X2) & l1_group_1(X1) & v4_group_1(X1) & v3_group_1(X1) & ~v3_struct_0(X1) & l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => (m2_relset_1(k8_group_6(X0,X1,X2,X3,X4),u1_struct_0(X0),u1_struct_0(X2)) & v1_group_6(k8_group_6(X0,X1,X2,X3,X4),X0,X2) & v1_funct_2(k8_group_6(X0,X1,X2,X3,X4),u1_struct_0(X0),u1_struct_0(X2)) & v1_funct_1(k8_group_6(X0,X1,X2,X3,X4))))),
  file('GRP621+1.tptp',unknown)).
fof(f1156,plain,(
  ( ! [X2,X0,X3,X1] : (~m1_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~m1_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~v4_group_6(k8_group_6(X2,X1,X2,X3,X0),X2,X2) | r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_funct_2(k8_group_6(X2,X1,X2,X3,X0),u1_struct_0(X2),u1_struct_0(X2)) | ~v1_group_1(X2)) )),
  inference(subsumption_resolution,[],[f1155,f264])).
fof(f264,plain,(
  ( ! [X4,X2,X0,X3,X1] : (v1_group_6(k8_group_6(X0,X1,X2,X3,X4),X0,X2) | ~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f1155,plain,(
  ( ! [X2,X0,X3,X1] : (~m1_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~m1_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~v4_group_6(k8_group_6(X2,X1,X2,X3,X0),X2,X2) | r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_group_6(k8_group_6(X2,X1,X2,X3,X0),X2,X2) | ~v1_funct_2(k8_group_6(X2,X1,X2,X3,X0),u1_struct_0(X2),u1_struct_0(X2)) | ~v1_group_1(X2)) )),
  inference(subsumption_resolution,[],[f1154,f262])).
fof(f262,plain,(
  ( ! [X4,X2,X0,X3,X1] : (v1_funct_1(k8_group_6(X0,X1,X2,X3,X4)) | ~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f1154,plain,(
  ( ! [X2,X0,X3,X1] : (~m1_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~m1_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~v4_group_6(k8_group_6(X2,X1,X2,X3,X0),X2,X2) | r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_group_6(k8_group_6(X2,X1,X2,X3,X0),X2,X2) | ~v1_funct_2(k8_group_6(X2,X1,X2,X3,X0),u1_struct_0(X2),u1_struct_0(X2)) | ~v1_funct_1(k8_group_6(X2,X1,X2,X3,X0)) | ~v1_group_1(X2)) )),
  inference(duplicate_literal_removal,[],[f1143])).
fof(f1143,plain,(
  ( ! [X2,X0,X3,X1] : (~m1_relset_1(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X0,X1,X2) | ~v1_funct_2(X0,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X0) | ~m1_relset_1(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_group_6(X3,X2,X1) | ~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~v4_group_6(k8_group_6(X2,X1,X2,X3,X0),X2,X2) | r2_hidden(k8_group_6(X2,X1,X2,X3,X0),k1_autgroup(X2)) | ~v1_group_6(k8_group_6(X2,X1,X2,X3,X0),X2,X2) | ~v1_funct_2(k8_group_6(X2,X1,X2,X3,X0),u1_struct_0(X2),u1_struct_0(X2)) | ~v1_funct_1(k8_group_6(X2,X1,X2,X3,X0)) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | ~v1_group_1(X2) | v3_struct_0(X2)) )),
  inference(resolution,[],[f265,f347])).
fof(f347,plain,(
  ( ! [X0,X1] : (~m2_relset_1(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v4_group_6(X1,X0,X0) | r2_hidden(X1,k1_autgroup(X0)) | ~v1_group_6(X1,X0,X0) | ~v1_funct_2(X1,u1_struct_0(X0),u1_struct_0(X0)) | ~v1_funct_1(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f265,plain,(
  ( ! [X4,X2,X0,X3,X1] : (m2_relset_1(k8_group_6(X0,X1,X2,X3,X4),u1_struct_0(X0),u1_struct_0(X2)) | ~m1_relset_1(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_group_6(X4,X1,X2) | ~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)) | ~v1_funct_1(X4) | ~m1_relset_1(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_group_6(X3,X0,X1) | ~v1_funct_2(X3,u1_struct_0(X0),u1_struct_0(X1)) | ~v1_funct_1(X3) | ~l1_group_1(X2) | ~v4_group_1(X2) | ~v3_group_1(X2) | v3_struct_0(X2) | ~l1_group_1(X1) | ~v4_group_1(X1) | ~v3_group_1(X1) | v3_struct_0(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f54050,plain,(
  spl28_163 | spl28_3 | ~spl28_4 | ~spl28_25),
  inference(avatar_split_clause,[],[f54049,f686,f463,f459,f5090])).
fof(f686,plain,(
  spl28_25 <=> r2_hidden(sK2,k1_autgroup(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_25])])).
fof(f54049,plain,(
  v4_group_6(sK2,sK1,sK1) | (spl28_3 | ~spl28_4 | ~spl28_25)),
  inference(subsumption_resolution,[],[f53951,f688])).
fof(f688,plain,(
  r2_hidden(sK2,k1_autgroup(sK1)) | ~spl28_25),
  inference(avatar_component_clause,[],[f686])).
fof(f53951,plain,(
  v4_group_6(sK2,sK1,sK1) | ~r2_hidden(sK2,k1_autgroup(sK1)) | (spl28_3 | ~spl28_4)),
  inference(resolution,[],[f18118,f880])).
fof(f18118,plain,(
  ( ! [X0] : (~v1_group_6(X0,sK1,sK1) | v4_group_6(X0,sK1,sK1) | ~r2_hidden(X0,k1_autgroup(sK1))) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f18117,f342])).
fof(f18117,plain,(
  ( ! [X0] : (v4_group_6(X0,sK1,sK1) | ~v1_group_6(X0,sK1,sK1) | ~r2_hidden(X0,k1_autgroup(sK1)) | ~m1_subset_1(X0,k1_autgroup(sK1))) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f18094,f1282])).
fof(f1282,plain,(
  ( ! [X5] : (~m1_subset_1(X5,k1_autgroup(sK1)) | v1_funct_1(X5)) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1281,f460])).
fof(f1281,plain,(
  ( ! [X5] : (~m1_subset_1(X5,k1_autgroup(sK1)) | v1_xboole_0(u1_struct_0(sK1)) | v1_funct_1(X5)) ) | ~spl28_4),
  inference(resolution,[],[f488,f464])).
fof(f488,plain,(
  ( ! [X6,X4,X7,X5] : (~m1_fraenkel(X5,X6,X7) | ~m1_subset_1(X4,X5) | v1_xboole_0(X7) | v1_funct_1(X4)) )),
  inference(duplicate_literal_removal,[],[f487])).
fof(f487,plain,(
  ( ! [X6,X4,X7,X5] : (~m1_subset_1(X4,X5) | ~m1_fraenkel(X5,X6,X7) | v1_xboole_0(X7) | v1_funct_1(X4) | ~m1_fraenkel(X5,X6,X7) | v1_xboole_0(X7)) )),
  inference(resolution,[],[f338,f269])).
fof(f269,plain,(
  ( ! [X2,X0,X3,X1] : (~m2_fraenkel(X3,X0,X1,X2) | v1_funct_1(X3) | ~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (! [X3] : ((m2_relset_1(X3,X0,X1) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3)) | ~m2_fraenkel(X3,X0,X1,X2)) | ~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : (! [X3] : ((m2_relset_1(X3,X0,X1) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3)) | ~m2_fraenkel(X3,X0,X1,X2)) | (~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1)))),
  inference(ennf_transformation,[],[f27])).
fof(f27,axiom,(
  ! [X0,X1,X2] : ((m1_fraenkel(X2,X0,X1) & ~v1_xboole_0(X1)) => ! [X3] : (m2_fraenkel(X3,X0,X1,X2) => (m2_relset_1(X3,X0,X1) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3))))),
  file('GRP621+1.tptp',unknown)).
fof(f338,plain,(
  ( ! [X2,X0,X3,X1] : (m2_fraenkel(X3,X0,X1,X2) | ~m1_subset_1(X3,X2) | ~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1)) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : (! [X3] : ((m2_fraenkel(X3,X0,X1,X2) | ~m1_subset_1(X3,X2)) & (m1_subset_1(X3,X2) | ~m2_fraenkel(X3,X0,X1,X2))) | ~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : (! [X3] : (m2_fraenkel(X3,X0,X1,X2) <=> m1_subset_1(X3,X2)) | ~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : (! [X3] : (m2_fraenkel(X3,X0,X1,X2) <=> m1_subset_1(X3,X2)) | (~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1)))),
  inference(ennf_transformation,[],[f60])).
fof(f60,axiom,(
  ! [X0,X1,X2] : ((m1_fraenkel(X2,X0,X1) & ~v1_xboole_0(X1)) => ! [X3] : (m2_fraenkel(X3,X0,X1,X2) <=> m1_subset_1(X3,X2)))),
  file('GRP621+1.tptp',unknown)).
fof(f18094,plain,(
  ( ! [X0] : (v4_group_6(X0,sK1,sK1) | ~v1_group_6(X0,sK1,sK1) | ~r2_hidden(X0,k1_autgroup(sK1)) | ~v1_funct_1(X0) | ~m1_subset_1(X0,k1_autgroup(sK1))) ) | (spl28_3 | ~spl28_4)),
  inference(resolution,[],[f1569,f1380])).
fof(f1380,plain,(
  ( ! [X5] : (v1_funct_2(X5,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m1_subset_1(X5,k1_autgroup(sK1))) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1379,f460])).
fof(f1379,plain,(
  ( ! [X5] : (v1_funct_2(X5,u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1)) | ~m1_subset_1(X5,k1_autgroup(sK1))) ) | ~spl28_4),
  inference(resolution,[],[f503,f464])).
fof(f503,plain,(
  ( ! [X6,X4,X5,X3] : (~m1_fraenkel(X6,X4,X5) | v1_funct_2(X3,X4,X5) | v1_xboole_0(X5) | ~m1_subset_1(X3,X6)) )),
  inference(duplicate_literal_removal,[],[f502])).
fof(f502,plain,(
  ( ! [X6,X4,X5,X3] : (v1_funct_2(X3,X4,X5) | ~m1_fraenkel(X6,X4,X5) | v1_xboole_0(X5) | ~m1_subset_1(X3,X6) | ~m1_fraenkel(X6,X4,X5) | v1_xboole_0(X5)) )),
  inference(resolution,[],[f270,f338])).
fof(f270,plain,(
  ( ! [X2,X0,X3,X1] : (~m2_fraenkel(X3,X0,X1,X2) | v1_funct_2(X3,X0,X1) | ~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1)) )),
  inference(cnf_transformation,[],[f111])).
fof(f1569,plain,(
  ( ! [X1] : (~v1_funct_2(X1,u1_struct_0(sK1),u1_struct_0(sK1)) | v4_group_6(X1,sK1,sK1) | ~v1_group_6(X1,sK1,sK1) | ~r2_hidden(X1,k1_autgroup(sK1)) | ~v1_funct_1(X1)) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1568,f342])).
fof(f1568,plain,(
  ( ! [X1] : (~m1_subset_1(X1,k1_autgroup(sK1)) | ~r2_hidden(X1,k1_autgroup(sK1)) | v4_group_6(X1,sK1,sK1) | ~v1_group_6(X1,sK1,sK1) | ~v1_funct_2(X1,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(X1)) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1567,f211])).
fof(f1567,plain,(
  ( ! [X1] : (~m1_subset_1(X1,k1_autgroup(sK1)) | ~r2_hidden(X1,k1_autgroup(sK1)) | v4_group_6(X1,sK1,sK1) | ~v1_group_6(X1,sK1,sK1) | ~v1_funct_2(X1,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(X1) | v3_struct_0(sK1)) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1566,f212])).
fof(f1566,plain,(
  ( ! [X1] : (~m1_subset_1(X1,k1_autgroup(sK1)) | ~r2_hidden(X1,k1_autgroup(sK1)) | v4_group_6(X1,sK1,sK1) | ~v1_group_6(X1,sK1,sK1) | ~v1_funct_2(X1,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~v1_group_1(sK1) | v3_struct_0(sK1)) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1565,f213])).
fof(f1565,plain,(
  ( ! [X1] : (~m1_subset_1(X1,k1_autgroup(sK1)) | ~r2_hidden(X1,k1_autgroup(sK1)) | v4_group_6(X1,sK1,sK1) | ~v1_group_6(X1,sK1,sK1) | ~v1_funct_2(X1,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1)) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1564,f214])).
fof(f1564,plain,(
  ( ! [X1] : (~m1_subset_1(X1,k1_autgroup(sK1)) | ~r2_hidden(X1,k1_autgroup(sK1)) | v4_group_6(X1,sK1,sK1) | ~v1_group_6(X1,sK1,sK1) | ~v1_funct_2(X1,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1)) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1548,f215])).
fof(f1548,plain,(
  ( ! [X1] : (~m1_subset_1(X1,k1_autgroup(sK1)) | ~r2_hidden(X1,k1_autgroup(sK1)) | v4_group_6(X1,sK1,sK1) | ~v1_group_6(X1,sK1,sK1) | ~v1_funct_2(X1,u1_struct_0(sK1),u1_struct_0(sK1)) | ~v1_funct_1(X1) | ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1)) ) | (spl28_3 | ~spl28_4)),
  inference(resolution,[],[f1410,f346])).
fof(f1410,plain,(
  ( ! [X5] : (m2_relset_1(X5,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m1_subset_1(X5,k1_autgroup(sK1))) ) | (spl28_3 | ~spl28_4)),
  inference(subsumption_resolution,[],[f1409,f460])).
fof(f1409,plain,(
  ( ! [X5] : (m2_relset_1(X5,u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1)) | ~m1_subset_1(X5,k1_autgroup(sK1))) ) | ~spl28_4),
  inference(resolution,[],[f531,f464])).
fof(f531,plain,(
  ( ! [X6,X4,X5,X3] : (~m1_fraenkel(X6,X4,X5) | m2_relset_1(X3,X4,X5) | v1_xboole_0(X5) | ~m1_subset_1(X3,X6)) )),
  inference(duplicate_literal_removal,[],[f530])).
fof(f530,plain,(
  ( ! [X6,X4,X5,X3] : (m2_relset_1(X3,X4,X5) | ~m1_fraenkel(X6,X4,X5) | v1_xboole_0(X5) | ~m1_subset_1(X3,X6) | ~m1_fraenkel(X6,X4,X5) | v1_xboole_0(X5)) )),
  inference(resolution,[],[f271,f338])).
fof(f271,plain,(
  ( ! [X2,X0,X3,X1] : (~m2_fraenkel(X3,X0,X1,X2) | m2_relset_1(X3,X0,X1) | ~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1)) )),
  inference(cnf_transformation,[],[f111])).
fof(f5104,plain,(
  spl28_25 | spl28_22 | ~spl28_18),
  inference(avatar_split_clause,[],[f738,f607,f657,f686])).
fof(f657,plain,(
  spl28_22 <=> v1_xboole_0(k1_autgroup(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_22])])).
fof(f607,plain,(
  spl28_18 <=> m1_subset_1(sK2,k1_autgroup(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_18])])).
fof(f738,plain,(
  v1_xboole_0(k1_autgroup(sK1)) | r2_hidden(sK2,k1_autgroup(sK1)) | ~spl28_18),
  inference(resolution,[],[f609,f343])).
fof(f343,plain,(
  ( ! [X0,X1] : (~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1))),
  inference(ennf_transformation,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1)))),
  file('GRP621+1.tptp',unknown)).
fof(f609,plain,(
  m1_subset_1(sK2,k1_autgroup(sK1)) | ~spl28_18),
  inference(avatar_component_clause,[],[f607])).
fof(f5102,plain,(
  ~spl28_22 | ~spl28_4),
  inference(avatar_split_clause,[],[f1200,f463,f657])).
fof(f1200,plain,(
  ~v1_xboole_0(k1_autgroup(sK1)) | ~spl28_4),
  inference(resolution,[],[f464,f267])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (~m1_fraenkel(X2,X0,X1) | ~v1_xboole_0(X2)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((v1_fraenkel(X2) & ~v1_xboole_0(X2)) | ~m1_fraenkel(X2,X0,X1))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2] : (m1_fraenkel(X2,X0,X1) => (v1_fraenkel(X2) & ~v1_xboole_0(X2)))),
  file('GRP621+1.tptp',unknown)).
fof(f660,plain,(
  spl28_21 | spl28_22 | ~spl28_17),
  inference(avatar_split_clause,[],[f649,f602,f657,f653])).
fof(f602,plain,(
  spl28_17 <=> m1_subset_1(sK3,k1_autgroup(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_17])])).
fof(f649,plain,(
  v1_xboole_0(k1_autgroup(sK1)) | r2_hidden(sK3,k1_autgroup(sK1)) | ~spl28_17),
  inference(resolution,[],[f604,f343])).
fof(f604,plain,(
  m1_subset_1(sK3,k1_autgroup(sK1)) | ~spl28_17),
  inference(avatar_component_clause,[],[f602])).
fof(f619,plain,(
  ~spl28_3),
  inference(avatar_contradiction_clause,[],[f618])).
fof(f618,plain,(
  $false | ~spl28_3),
  inference(subsumption_resolution,[],[f617,f211])).
fof(f617,plain,(
  v3_struct_0(sK1) | ~spl28_3),
  inference(subsumption_resolution,[],[f611,f364])).
fof(f364,plain,(
  l1_struct_0(sK1)),
  inference(resolution,[],[f266,f215])).
fof(f266,plain,(
  ( ! [X0] : (~l1_group_1(X0) | l1_struct_0(X0)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : (l1_struct_0(X0) | ~l1_group_1(X0))),
  inference(ennf_transformation,[],[f22])).
fof(f22,axiom,(
  ! [X0] : (l1_group_1(X0) => l1_struct_0(X0))),
  file('GRP621+1.tptp',unknown)).
fof(f611,plain,(
  ~l1_struct_0(sK1) | v3_struct_0(sK1) | ~spl28_3),
  inference(resolution,[],[f461,f287])).
fof(f287,plain,(
  ( ! [X0] : (~v1_xboole_0(u1_struct_0(X0)) | ~l1_struct_0(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : (~v1_xboole_0(u1_struct_0(X0)) | ~l1_struct_0(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : (~v1_xboole_0(u1_struct_0(X0)) | (~l1_struct_0(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,axiom,(
  ! [X0] : ((l1_struct_0(X0) & ~v3_struct_0(X0)) => ~v1_xboole_0(u1_struct_0(X0)))),
  file('GRP621+1.tptp',unknown)).
fof(f461,plain,(
  v1_xboole_0(u1_struct_0(sK1)) | ~spl28_3),
  inference(avatar_component_clause,[],[f459])).
fof(f610,plain,(
  spl28_3 | ~spl28_4 | spl28_18),
  inference(avatar_split_clause,[],[f482,f607,f463,f459])).
fof(f482,plain,(
  m1_subset_1(sK2,k1_autgroup(sK1)) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f337,f216])).
fof(f337,plain,(
  ( ! [X2,X0,X3,X1] : (~m2_fraenkel(X3,X0,X1,X2) | m1_subset_1(X3,X2) | ~m1_fraenkel(X2,X0,X1) | v1_xboole_0(X1)) )),
  inference(cnf_transformation,[],[f208])).
fof(f605,plain,(
  spl28_3 | ~spl28_4 | spl28_17),
  inference(avatar_split_clause,[],[f483,f602,f463,f459])).
fof(f483,plain,(
  m1_subset_1(sK3,k1_autgroup(sK1)) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f337,f217])).
fof(f600,plain,(
  spl28_3 | ~spl28_4 | spl28_16),
  inference(avatar_split_clause,[],[f499,f597,f463,f459])).
fof(f499,plain,(
  v1_funct_2(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f270,f216])).
fof(f595,plain,(
  spl28_3 | ~spl28_4 | spl28_15),
  inference(avatar_split_clause,[],[f500,f592,f463,f459])).
fof(f500,plain,(
  v1_funct_2(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f270,f217])).
fof(f590,plain,(
  spl28_3 | ~spl28_4 | spl28_14),
  inference(avatar_split_clause,[],[f527,f587,f463,f459])).
fof(f527,plain,(
  m2_relset_1(sK2,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f271,f216])).
fof(f585,plain,(
  spl28_3 | ~spl28_4 | spl28_13),
  inference(avatar_split_clause,[],[f528,f582,f463,f459])).
fof(f528,plain,(
  m2_relset_1(sK3,u1_struct_0(sK1),u1_struct_0(sK1)) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f271,f217])).
fof(f580,plain,(
  spl28_3 | ~spl28_4 | ~spl28_12),
  inference(avatar_split_clause,[],[f557,f577,f463,f459])).
fof(f557,plain,(
  ~m1_subset_1(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2),k1_autgroup(sK1)) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f218,f338])).
fof(f218,plain,(
  ~m2_fraenkel(k7_funct_2(u1_struct_0(sK1),u1_struct_0(sK1),u1_struct_0(sK1),sK3,sK2),u1_struct_0(sK1),u1_struct_0(sK1),k1_autgroup(sK1))),
  inference(cnf_transformation,[],[f153])).
fof(f575,plain,(
  spl28_4),
  inference(avatar_contradiction_clause,[],[f574])).
fof(f574,plain,(
  $false | spl28_4),
  inference(subsumption_resolution,[],[f573,f211])).
fof(f573,plain,(
  v3_struct_0(sK1) | spl28_4),
  inference(subsumption_resolution,[],[f572,f212])).
fof(f572,plain,(
  ~v1_group_1(sK1) | v3_struct_0(sK1) | spl28_4),
  inference(subsumption_resolution,[],[f571,f213])).
fof(f571,plain,(
  ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | spl28_4),
  inference(subsumption_resolution,[],[f570,f214])).
fof(f570,plain,(
  ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | spl28_4),
  inference(subsumption_resolution,[],[f567,f215])).
fof(f567,plain,(
  ~l1_group_1(sK1) | ~v4_group_1(sK1) | ~v3_group_1(sK1) | ~v1_group_1(sK1) | v3_struct_0(sK1) | spl28_4),
  inference(resolution,[],[f257,f465])).
fof(f465,plain,(
  ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | spl28_4),
  inference(avatar_component_clause,[],[f463])).
fof(f257,plain,(
  ( ! [X0] : (m1_fraenkel(k1_autgroup(X0),u1_struct_0(X0),u1_struct_0(X0)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : (m1_fraenkel(k1_autgroup(X0),u1_struct_0(X0),u1_struct_0(X0)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : (m1_fraenkel(k1_autgroup(X0),u1_struct_0(X0),u1_struct_0(X0)) | (~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | ~v1_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & v1_group_1(X0) & ~v3_struct_0(X0)) => m1_fraenkel(k1_autgroup(X0),u1_struct_0(X0),u1_struct_0(X0)))),
  file('GRP621+1.tptp',unknown)).
fof(f475,plain,(
  spl28_3 | ~spl28_4 | spl28_6),
  inference(avatar_split_clause,[],[f457,f472,f463,f459])).
fof(f457,plain,(
  v1_funct_1(sK3) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f269,f217])).
fof(f470,plain,(
  spl28_3 | ~spl28_4 | spl28_5),
  inference(avatar_split_clause,[],[f456,f467,f463,f459])).
fof(f456,plain,(
  v1_funct_1(sK2) | ~m1_fraenkel(k1_autgroup(sK1),u1_struct_0(sK1),u1_struct_0(sK1)) | v1_xboole_0(u1_struct_0(sK1))),
  inference(resolution,[],[f269,f216])).
% SZS output end Proof for GRP621+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 56672
% Time elapsed: 1.504 s
% ------------------------------
% ------------------------------
