% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP701-1
% SZS output start Proof for GRP701-1
fof(f182915,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f182914])).
fof(f182914,plain,(
  i(b) != i(b)),
  inference(superposition,[],[f11,f179775])).
fof(f179775,plain,(
  ( ! [X92,X93] : (i(X92) = mult(X93,i(mult(X92,X93)))) )),
  inference(backward_demodulation,[],[f179774,f22238])).
fof(f22238,plain,(
  ( ! [X92,X93] : (mult(X93,i(mult(X92,X93))) = ld(ld(i(X93),X92),X93)) )),
  inference(forward_demodulation,[],[f22237,f13323])).
fof(f13323,plain,(
  ( ! [X0,X1] : (mult(mult(X0,X0),ld(X0,X1)) = ld(i(X0),X1)) )),
  inference(backward_demodulation,[],[f13315,f6289])).
fof(f6289,plain,(
  ( ! [X0,X1] : (mult(mult(X0,X0),ld(X0,X1)) = ld(X0,mult(mult(X0,X0),X1))) )),
  inference(superposition,[],[f2,f172])).
fof(f172,plain,(
  ( ! [X2,X3] : (mult(X2,mult(mult(X2,X2),ld(X2,X3))) = mult(mult(X2,X2),X3)) )),
  inference(superposition,[],[f63,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (mult(X0,ld(X0,X1)) = X1) )),
  file('GRP701-1.tptp',unknown)).
fof(f63,plain,(
  ( ! [X0,X1] : (mult(mult(X0,X0),mult(X0,X1)) = mult(X0,mult(mult(X0,X0),X1))) )),
  inference(forward_demodulation,[],[f45,f6])).
fof(f6,axiom,(
  ( ! [X0] : (mult(unit,X0) = X0) )),
  file('GRP701-1.tptp',unknown)).
fof(f45,plain,(
  ( ! [X0,X1] : (mult(X0,mult(mult(mult(unit,X0),X0),X1)) = mult(mult(X0,X0),mult(X0,X1))) )),
  inference(superposition,[],[f7,f5])).
fof(f5,axiom,(
  ( ! [X0] : (mult(X0,unit) = X0) )),
  file('GRP701-1.tptp',unknown)).
fof(f7,axiom,(
  ( ! [X2,X0,X1] : (mult(mult(mult(X0,X1),X0),mult(X0,X2)) = mult(X0,mult(mult(mult(X1,X0),X0),X2))) )),
  file('GRP701-1.tptp',unknown)).
fof(f2,axiom,(
  ( ! [X0,X1] : (ld(X0,mult(X0,X1)) = X1) )),
  file('GRP701-1.tptp',unknown)).
fof(f13315,plain,(
  ( ! [X68,X67] : (ld(X68,mult(mult(X68,X68),X67)) = ld(i(X68),X67)) )),
  inference(backward_demodulation,[],[f13225,f11387])).
fof(f11387,plain,(
  ( ! [X68,X67] : (rd(mult(X68,mult(X67,X68)),X68) = ld(X68,mult(mult(X68,X68),X67))) )),
  inference(forward_demodulation,[],[f11328,f6289])).
fof(f11328,plain,(
  ( ! [X68,X67] : (rd(mult(X68,mult(X67,X68)),X68) = mult(mult(X68,X68),ld(X68,X67))) )),
  inference(backward_demodulation,[],[f10947,f7315])).
fof(f7315,plain,(
  ( ! [X68,X67] : (mult(mult(X68,rd(X67,X68)),X68) = rd(mult(X68,mult(X67,X68)),X68)) )),
  inference(superposition,[],[f251,f3])).
fof(f3,axiom,(
  ( ! [X0,X1] : (mult(rd(X0,X1),X1) = X0) )),
  file('GRP701-1.tptp',unknown)).
fof(f251,plain,(
  ( ! [X2,X3] : (mult(mult(X2,X3),X2) = rd(mult(X2,mult(mult(X3,X2),X2)),X2)) )),
  inference(superposition,[],[f4,f75])).
fof(f75,plain,(
  ( ! [X0,X1] : (mult(mult(mult(X0,X1),X0),X0) = mult(X0,mult(mult(X1,X0),X0))) )),
  inference(forward_demodulation,[],[f54,f5])).
fof(f54,plain,(
  ( ! [X0,X1] : (mult(X0,mult(mult(mult(X1,X0),X0),unit)) = mult(mult(mult(X0,X1),X0),X0)) )),
  inference(superposition,[],[f7,f5])).
fof(f4,axiom,(
  ( ! [X0,X1] : (rd(mult(X0,X1),X1) = X0) )),
  file('GRP701-1.tptp',unknown)).
fof(f10947,plain,(
  ( ! [X103,X101,X102] : (mult(mult(X103,rd(X101,X102)),X102) = mult(mult(X103,X102),ld(X102,X101))) )),
  inference(superposition,[],[f8540,f3])).
fof(f8540,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X2,X1),X0) = mult(mult(X2,X0),ld(X0,mult(X1,X0)))) )),
  inference(superposition,[],[f8249,f1])).
fof(f8249,plain,(
  ( ! [X6,X4,X5] : (mult(mult(X6,mult(X4,X5)),X4) = mult(mult(X6,X4),ld(X4,mult(mult(X4,X5),X4)))) )),
  inference(backward_demodulation,[],[f8094,f88])).
fof(f88,plain,(
  ( ! [X6,X4,X5] : (mult(mult(X6,X4),mult(X4,mult(ld(X4,X5),X4))) = mult(mult(X6,mult(X4,X5)),X4)) )),
  inference(superposition,[],[f8,f1])).
fof(f8,axiom,(
  ( ! [X2,X0,X1] : (mult(mult(X0,X1),mult(X1,mult(X2,X1))) = mult(mult(X0,mult(X1,mult(X1,X2))),X1)) )),
  file('GRP701-1.tptp',unknown)).
fof(f8094,plain,(
  ( ! [X28,X27] : (mult(X27,mult(ld(X27,X28),X27)) = ld(X27,mult(mult(X27,X28),X27))) )),
  inference(superposition,[],[f2,f372])).
fof(f372,plain,(
  ( ! [X2,X3] : (mult(mult(X2,X3),X2) = mult(X2,mult(X2,mult(ld(X2,X3),X2)))) )),
  inference(superposition,[],[f122,f1])).
fof(f122,plain,(
  ( ! [X10,X11] : (mult(mult(X10,mult(X10,X11)),X10) = mult(X10,mult(X10,mult(X11,X10)))) )),
  inference(forward_demodulation,[],[f98,f6])).
fof(f98,plain,(
  ( ! [X10,X11] : (mult(mult(unit,X10),mult(X10,mult(X11,X10))) = mult(mult(X10,mult(X10,X11)),X10)) )),
  inference(superposition,[],[f8,f6])).
fof(f13225,plain,(
  ( ! [X2,X3] : (ld(i(X2),X3) = rd(mult(X2,mult(X3,X2)),X2)) )),
  inference(superposition,[],[f4,f12042])).
fof(f12042,plain,(
  ( ! [X6,X5] : (mult(X5,mult(X6,X5)) = mult(ld(i(X5),X6),X5)) )),
  inference(superposition,[],[f1,f11505])).
fof(f11505,plain,(
  ( ! [X6,X7] : (mult(X7,X6) = ld(X6,mult(ld(i(X6),X7),X6))) )),
  inference(superposition,[],[f11062,f1])).
fof(f11062,plain,(
  ( ! [X114,X115] : (mult(mult(i(X114),X115),X114) = ld(X114,mult(X115,X114))) )),
  inference(forward_demodulation,[],[f10912,f6])).
fof(f10912,plain,(
  ( ! [X114,X115] : (mult(mult(i(X114),X115),X114) = mult(unit,ld(X114,mult(X115,X114)))) )),
  inference(superposition,[],[f8540,f10])).
fof(f10,axiom,(
  ( ! [X0] : (unit = mult(i(X0),X0)) )),
  file('GRP701-1.tptp',unknown)).
fof(f22237,plain,(
  ( ! [X92,X93] : (mult(X93,i(mult(X92,X93))) = ld(mult(mult(X93,X93),ld(X93,X92)),X93)) )),
  inference(forward_demodulation,[],[f22060,f10947])).
fof(f22060,plain,(
  ( ! [X92,X93] : (ld(mult(mult(X93,rd(X92,X93)),X93),X93) = mult(X93,i(mult(X92,X93)))) )),
  inference(superposition,[],[f1155,f3])).
fof(f1155,plain,(
  ( ! [X14,X13] : (mult(X14,i(mult(mult(X13,X14),X14))) = ld(mult(mult(X14,X13),X14),X14)) )),
  inference(forward_demodulation,[],[f1081,f5])).
fof(f1081,plain,(
  ( ! [X14,X13] : (mult(X14,i(mult(mult(X13,X14),X14))) = ld(mult(mult(X14,X13),X14),mult(X14,unit))) )),
  inference(superposition,[],[f62,f9])).
fof(f9,axiom,(
  ( ! [X0] : (unit = mult(X0,i(X0))) )),
  file('GRP701-1.tptp',unknown)).
fof(f62,plain,(
  ( ! [X4,X5,X3] : (mult(X3,X5) = ld(mult(mult(X3,X4),X3),mult(X3,mult(mult(mult(X4,X3),X3),X5)))) )),
  inference(superposition,[],[f2,f7])).
fof(f179774,plain,(
  ( ! [X2,X1] : (i(X2) = ld(ld(i(X1),X2),X1)) )),
  inference(forward_demodulation,[],[f179657,f179074])).
fof(f179074,plain,(
  ( ! [X231,X230] : (i(X230) = mult(ld(X230,X231),i(X231))) )),
  inference(backward_demodulation,[],[f179073,f68371])).
fof(f68371,plain,(
  ( ! [X231,X230] : (ld(X231,i(rd(X230,X231))) = mult(ld(X230,X231),i(X231))) )),
  inference(forward_demodulation,[],[f68139,f2])).
fof(f68139,plain,(
  ( ! [X231,X230] : (ld(X231,i(rd(X230,X231))) = mult(ld(X231,mult(X231,ld(X230,X231))),i(X231))) )),
  inference(superposition,[],[f14141,f28270])).
fof(f28270,plain,(
  ( ! [X0,X1] : (mult(X0,ld(X1,X0)) = mult(i(rd(X1,X0)),X0)) )),
  inference(backward_demodulation,[],[f28087,f25769])).
fof(f25769,plain,(
  ( ! [X0,X1] : (rd(X0,ld(X0,X1)) = mult(i(rd(X1,X0)),X0)) )),
  inference(superposition,[],[f23521,f1])).
fof(f23521,plain,(
  ( ! [X2,X3] : (rd(X2,X3) = mult(i(rd(mult(X2,X3),X2)),X2)) )),
  inference(backward_demodulation,[],[f23024,f20495])).
fof(f20495,plain,(
  ( ! [X2,X3] : (rd(X2,X3) = mult(i(mult(X2,mult(X3,i(X2)))),X2)) )),
  inference(forward_demodulation,[],[f20494,f1])).
fof(f20494,plain,(
  ( ! [X2,X3] : (mult(i(mult(X2,mult(X3,i(X2)))),X2) = rd(X2,mult(X2,ld(X2,X3)))) )),
  inference(forward_demodulation,[],[f20363,f3])).
fof(f20363,plain,(
  ( ! [X2,X3] : (rd(X2,mult(X2,mult(rd(ld(X2,X3),X2),X2))) = mult(i(mult(X2,mult(X3,i(X2)))),X2)) )),
  inference(superposition,[],[f615,f16429])).
fof(f16429,plain,(
  ( ! [X10,X9] : (mult(X10,i(X9)) = mult(X9,rd(ld(X9,X10),X9))) )),
  inference(superposition,[],[f1,f16031])).
fof(f16031,plain,(
  ( ! [X17,X16] : (ld(X16,mult(X17,i(X16))) = rd(ld(X16,X17),X16)) )),
  inference(forward_demodulation,[],[f15843,f12913])).
fof(f12913,plain,(
  ( ! [X76,X75] : (mult(i(X76),rd(X75,X76)) = rd(ld(X76,X75),X76)) )),
  inference(superposition,[],[f11547,f3])).
fof(f11547,plain,(
  ( ! [X64,X65] : (mult(i(X64),X65) = rd(ld(X64,mult(X65,X64)),X64)) )),
  inference(superposition,[],[f4,f11062])).
fof(f15843,plain,(
  ( ! [X17,X16] : (ld(X16,mult(X17,i(X16))) = mult(i(X16),rd(X17,X16))) )),
  inference(backward_demodulation,[],[f15680,f14702])).
fof(f14702,plain,(
  ( ! [X17,X16] : (mult(i(X16),ld(X16,mult(mult(X16,X17),i(X16)))) = ld(X16,mult(X17,i(X16)))) )),
  inference(backward_demodulation,[],[f14544,f14634])).
fof(f14634,plain,(
  ( ! [X17,X16] : (mult(mult(i(X16),X17),i(X16)) = mult(i(X16),ld(X16,mult(mult(X16,X17),i(X16))))) )),
  inference(backward_demodulation,[],[f14633,f12674])).
fof(f12674,plain,(
  ( ! [X17,X16] : (mult(mult(i(X16),X17),i(X16)) = mult(i(X16),mult(ld(X16,ld(i(X16),X17)),i(X16)))) )),
  inference(backward_demodulation,[],[f12609,f11620])).
fof(f11620,plain,(
  ( ! [X17,X16] : (mult(mult(i(X16),X17),i(X16)) = mult(i(X16),mult(mult(X16,mult(i(X16),X17)),i(X16)))) )),
  inference(backward_demodulation,[],[f11617,f11128])).
fof(f11128,plain,(
  ( ! [X17,X16] : (mult(mult(i(X16),X17),i(X16)) = mult(i(X16),mult(mult(mult(X17,X16),i(X16)),i(X16)))) )),
  inference(forward_demodulation,[],[f11077,f11076])).
fof(f11076,plain,(
  ( ! [X6,X5] : (mult(mult(X6,X5),i(X5)) = mult(mult(X6,i(X5)),X5)) )),
  inference(forward_demodulation,[],[f11075,f83])).
fof(f83,plain,(
  ( ! [X1] : (i(i(X1)) = X1) )),
  inference(superposition,[],[f24,f22])).
fof(f22,plain,(
  ( ! [X1] : (i(X1) = ld(X1,unit)) )),
  inference(superposition,[],[f2,f9])).
fof(f24,plain,(
  ( ! [X5] : (ld(i(X5),unit) = X5) )),
  inference(superposition,[],[f2,f10])).
fof(f11075,plain,(
  ( ! [X6,X5] : (mult(mult(X6,X5),i(X5)) = mult(mult(X6,i(X5)),i(i(X5)))) )),
  inference(forward_demodulation,[],[f10921,f22])).
fof(f10921,plain,(
  ( ! [X6,X5] : (mult(mult(X6,X5),i(X5)) = mult(mult(X6,i(X5)),ld(i(X5),unit))) )),
  inference(superposition,[],[f8540,f9])).
fof(f11077,plain,(
  ( ! [X17,X16] : (mult(mult(i(X16),X17),i(X16)) = mult(i(X16),mult(mult(mult(X17,i(X16)),X16),i(X16)))) )),
  inference(backward_demodulation,[],[f11076,f81])).
fof(f81,plain,(
  ( ! [X17,X16] : (mult(mult(i(X16),X17),i(X16)) = mult(i(X16),mult(mult(mult(X17,i(X16)),i(X16)),X16))) )),
  inference(forward_demodulation,[],[f60,f5])).
fof(f60,plain,(
  ( ! [X17,X16] : (mult(i(X16),mult(mult(mult(X17,i(X16)),i(X16)),X16)) = mult(mult(mult(i(X16),X17),i(X16)),unit)) )),
  inference(superposition,[],[f7,f10])).
fof(f11617,plain,(
  ( ! [X4,X5] : (mult(mult(mult(X5,X4),i(X4)),i(X4)) = mult(mult(X4,mult(i(X4),X5)),i(X4))) )),
  inference(backward_demodulation,[],[f11498,f11129])).
fof(f11129,plain,(
  ( ! [X4,X5] : (ld(i(X4),mult(mult(i(X4),X5),i(X4))) = mult(mult(mult(X5,X4),i(X4)),i(X4))) )),
  inference(forward_demodulation,[],[f11078,f11076])).
fof(f11078,plain,(
  ( ! [X4,X5] : (ld(i(X4),mult(mult(i(X4),X5),i(X4))) = mult(mult(mult(X5,i(X4)),X4),i(X4))) )),
  inference(backward_demodulation,[],[f11076,f1867])).
fof(f1867,plain,(
  ( ! [X4,X5] : (mult(mult(mult(X5,i(X4)),i(X4)),X4) = ld(i(X4),mult(mult(i(X4),X5),i(X4)))) )),
  inference(superposition,[],[f2,f81])).
fof(f11498,plain,(
  ( ! [X2,X1] : (ld(i(X1),mult(X2,i(X1))) = mult(mult(X1,X2),i(X1))) )),
  inference(superposition,[],[f11062,f83])).
fof(f12609,plain,(
  ( ! [X24,X23] : (mult(mult(X23,mult(i(X23),X24)),i(X23)) = mult(ld(X23,ld(i(X23),X24)),i(X23))) )),
  inference(backward_demodulation,[],[f12465,f1352])).
fof(f1352,plain,(
  ( ! [X24,X23] : (mult(i(X23),mult(ld(i(X23),X24),i(X23))) = mult(mult(X23,mult(i(X23),X24)),i(X23))) )),
  inference(forward_demodulation,[],[f1351,f1309])).
fof(f1309,plain,(
  ( ! [X6] : (mult(i(X6),mult(X6,X6)) = X6) )),
  inference(forward_demodulation,[],[f1279,f33])).
fof(f33,plain,(
  ( ! [X1] : (rd(unit,i(X1)) = X1) )),
  inference(superposition,[],[f4,f9])).
fof(f1279,plain,(
  ( ! [X6] : (mult(i(X6),mult(X6,X6)) = rd(unit,i(X6))) )),
  inference(superposition,[],[f4,f986])).
fof(f986,plain,(
  ( ! [X29] : (unit = mult(mult(i(X29),mult(X29,X29)),i(X29))) )),
  inference(forward_demodulation,[],[f985,f35])).
fof(f35,plain,(
  ( ! [X6] : (unit = rd(X6,X6)) )),
  inference(superposition,[],[f4,f6])).
fof(f985,plain,(
  ( ! [X30,X29] : (mult(mult(i(X29),mult(X29,X29)),i(X29)) = rd(mult(i(X29),X30),mult(i(X29),X30))) )),
  inference(forward_demodulation,[],[f984,f6])).
fof(f984,plain,(
  ( ! [X30,X29] : (mult(mult(i(X29),mult(X29,X29)),i(X29)) = rd(mult(i(X29),mult(unit,X30)),mult(i(X29),X30))) )),
  inference(forward_demodulation,[],[f936,f9])).
fof(f936,plain,(
  ( ! [X30,X29] : (mult(mult(i(X29),mult(X29,X29)),i(X29)) = rd(mult(i(X29),mult(mult(X29,i(X29)),X30)),mult(i(X29),X30))) )),
  inference(superposition,[],[f61,f457])).
fof(f457,plain,(
  ( ! [X0] : (mult(mult(X0,X0),i(X0)) = X0) )),
  inference(forward_demodulation,[],[f443,f2])).
fof(f443,plain,(
  ( ! [X0] : (mult(mult(X0,X0),i(X0)) = ld(X0,mult(X0,X0))) )),
  inference(superposition,[],[f2,f187])).
fof(f187,plain,(
  ( ! [X1] : (mult(X1,X1) = mult(X1,mult(mult(X1,X1),i(X1)))) )),
  inference(forward_demodulation,[],[f171,f5])).
fof(f171,plain,(
  ( ! [X1] : (mult(X1,mult(mult(X1,X1),i(X1))) = mult(mult(X1,X1),unit)) )),
  inference(superposition,[],[f63,f9])).
fof(f61,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X0,X1),X0) = rd(mult(X0,mult(mult(mult(X1,X0),X0),X2)),mult(X0,X2))) )),
  inference(superposition,[],[f4,f7])).
fof(f1351,plain,(
  ( ! [X24,X23] : (mult(i(X23),mult(ld(i(X23),X24),i(X23))) = mult(mult(mult(i(X23),mult(X23,X23)),mult(i(X23),X24)),i(X23))) )),
  inference(forward_demodulation,[],[f1289,f6])).
fof(f1289,plain,(
  ( ! [X24,X23] : (mult(mult(mult(i(X23),mult(X23,X23)),mult(i(X23),X24)),i(X23)) = mult(unit,mult(i(X23),mult(ld(i(X23),X24),i(X23))))) )),
  inference(superposition,[],[f88,f986])).
fof(f12465,plain,(
  ( ! [X0,X1] : (mult(i(X0),mult(X1,i(X0))) = mult(ld(X0,X1),i(X0))) )),
  inference(superposition,[],[f11850,f1])).
fof(f11850,plain,(
  ( ! [X8,X9] : (mult(X9,i(X8)) = mult(i(X8),mult(mult(X8,X9),i(X8)))) )),
  inference(superposition,[],[f11620,f1])).
fof(f14633,plain,(
  ( ! [X2,X3] : (mult(ld(X2,ld(i(X2),X3)),i(X2)) = ld(X2,mult(mult(X2,X3),i(X2)))) )),
  inference(forward_demodulation,[],[f14632,f11498])).
fof(f14632,plain,(
  ( ! [X2,X3] : (mult(ld(X2,ld(i(X2),X3)),i(X2)) = ld(X2,ld(i(X2),mult(X3,i(X2))))) )),
  inference(forward_demodulation,[],[f14631,f14132])).
fof(f14132,plain,(
  ( ! [X0,X1] : (ld(X0,ld(i(X0),X1)) = ld(i(X0),ld(X0,X1))) )),
  inference(forward_demodulation,[],[f14015,f13593])).
fof(f13593,plain,(
  ( ! [X14,X15] : (mult(i(X14),mult(X14,X15)) = ld(X14,ld(i(X14),X15))) )),
  inference(backward_demodulation,[],[f13485,f11275])).
fof(f11275,plain,(
  ( ! [X14,X15] : (mult(X14,mult(i(X14),X15)) = mult(i(X14),mult(X14,X15))) )),
  inference(forward_demodulation,[],[f11274,f1309])).
fof(f11274,plain,(
  ( ! [X14,X15] : (mult(X14,mult(i(X14),X15)) = mult(i(X14),mult(mult(i(X14),mult(X14,X14)),X15))) )),
  inference(forward_demodulation,[],[f11273,f6])).
fof(f11273,plain,(
  ( ! [X14,X15] : (mult(X14,mult(i(X14),X15)) = mult(i(X14),mult(mult(mult(unit,i(X14)),mult(X14,X14)),X15))) )),
  inference(forward_demodulation,[],[f11272,f9])).
fof(f11272,plain,(
  ( ! [X14,X15] : (mult(X14,mult(i(X14),X15)) = mult(i(X14),mult(mult(mult(mult(X14,i(X14)),i(X14)),mult(X14,X14)),X15))) )),
  inference(forward_demodulation,[],[f11255,f11252])).
fof(f11252,plain,(
  ( ! [X99,X100] : (mult(mult(X100,mult(X99,X99)),i(X99)) = mult(mult(X100,i(X99)),mult(X99,X99))) )),
  inference(forward_demodulation,[],[f10946,f1494])).
fof(f1494,plain,(
  ( ! [X1] : (mult(X1,X1) = ld(i(X1),X1)) )),
  inference(superposition,[],[f2,f1309])).
fof(f10946,plain,(
  ( ! [X99,X100] : (mult(mult(X100,mult(X99,X99)),i(X99)) = mult(mult(X100,i(X99)),ld(i(X99),X99))) )),
  inference(superposition,[],[f8540,f457])).
fof(f11255,plain,(
  ( ! [X14,X15] : (mult(X14,mult(i(X14),X15)) = mult(i(X14),mult(mult(mult(mult(X14,i(X14)),mult(X14,X14)),i(X14)),X15))) )),
  inference(backward_demodulation,[],[f11252,f9720])).
fof(f9720,plain,(
  ( ! [X14,X15] : (mult(i(X14),mult(mult(mult(mult(X14,mult(X14,X14)),i(X14)),i(X14)),X15)) = mult(X14,mult(i(X14),X15))) )),
  inference(forward_demodulation,[],[f9629,f457])).
fof(f9629,plain,(
  ( ! [X14,X15] : (mult(i(X14),mult(mult(mult(mult(X14,mult(X14,X14)),i(X14)),i(X14)),X15)) = mult(mult(mult(X14,X14),i(X14)),mult(i(X14),X15))) )),
  inference(superposition,[],[f7,f8794])).
fof(f8794,plain,(
  ( ! [X0] : (mult(X0,X0) = mult(i(X0),mult(X0,mult(X0,X0)))) )),
  inference(forward_demodulation,[],[f8781,f3223])).
fof(f3223,plain,(
  ( ! [X3] : (i(X3) = mult(i(mult(X3,X3)),X3)) )),
  inference(forward_demodulation,[],[f3195,f1495])).
fof(f1495,plain,(
  ( ! [X2] : (i(X2) = rd(X2,mult(X2,X2))) )),
  inference(superposition,[],[f4,f1309])).
fof(f3195,plain,(
  ( ! [X3] : (mult(i(mult(X3,X3)),X3) = rd(X3,mult(X3,X3))) )),
  inference(superposition,[],[f4,f343])).
fof(f343,plain,(
  ( ! [X14] : (mult(mult(i(mult(X14,X14)),X14),mult(X14,X14)) = X14) )),
  inference(forward_demodulation,[],[f307,f6])).
fof(f307,plain,(
  ( ! [X14] : (mult(unit,X14) = mult(mult(i(mult(X14,X14)),X14),mult(X14,X14))) )),
  inference(superposition,[],[f107,f10])).
fof(f107,plain,(
  ( ! [X0,X1] : (mult(mult(X1,mult(X0,X0)),X0) = mult(mult(X1,X0),mult(X0,X0))) )),
  inference(forward_demodulation,[],[f86,f6])).
fof(f86,plain,(
  ( ! [X0,X1] : (mult(mult(X1,X0),mult(X0,mult(unit,X0))) = mult(mult(X1,mult(X0,X0)),X0)) )),
  inference(superposition,[],[f8,f5])).
fof(f8781,plain,(
  ( ! [X0] : (mult(X0,X0) = mult(mult(i(mult(X0,X0)),X0),mult(X0,mult(X0,X0)))) )),
  inference(superposition,[],[f97,f8492])).
fof(f8492,plain,(
  ( ! [X0] : (i(mult(X0,X0)) = rd(X0,mult(X0,mult(X0,X0)))) )),
  inference(superposition,[],[f34,f7878])).
fof(f7878,plain,(
  ( ! [X3] : (mult(X3,mult(X3,X3)) = ld(i(mult(X3,X3)),X3)) )),
  inference(forward_demodulation,[],[f7877,f185])).
fof(f185,plain,(
  ( ! [X0] : (mult(mult(X0,X0),X0) = mult(X0,mult(X0,X0))) )),
  inference(forward_demodulation,[],[f170,f5])).
fof(f170,plain,(
  ( ! [X0] : (mult(X0,mult(mult(X0,X0),unit)) = mult(mult(X0,X0),X0)) )),
  inference(superposition,[],[f63,f5])).
fof(f7877,plain,(
  ( ! [X3] : (mult(mult(X3,X3),X3) = ld(i(mult(X3,X3)),X3)) )),
  inference(forward_demodulation,[],[f7876,f83])).
fof(f7876,plain,(
  ( ! [X3] : (mult(i(i(mult(X3,X3))),X3) = ld(i(mult(X3,X3)),X3)) )),
  inference(forward_demodulation,[],[f7821,f83])).
fof(f7821,plain,(
  ( ! [X3] : (mult(i(i(mult(X3,X3))),i(i(X3))) = ld(i(mult(X3,X3)),i(i(X3)))) )),
  inference(superposition,[],[f7727,f3662])).
fof(f3662,plain,(
  ( ! [X0] : (mult(i(X0),i(X0)) = i(mult(X0,X0))) )),
  inference(backward_demodulation,[],[f3661,f1646])).
fof(f1646,plain,(
  ( ! [X0] : (mult(i(X0),i(X0)) = ld(X0,i(X0))) )),
  inference(superposition,[],[f1494,f83])).
fof(f3661,plain,(
  ( ! [X0] : (i(mult(X0,X0)) = ld(X0,i(X0))) )),
  inference(backward_demodulation,[],[f3622,f1655])).
fof(f1655,plain,(
  ( ! [X0] : (rd(i(X0),X0) = ld(X0,i(X0))) )),
  inference(backward_demodulation,[],[f1646,f797])).
fof(f797,plain,(
  ( ! [X0] : (mult(i(X0),i(X0)) = rd(i(X0),X0)) )),
  inference(superposition,[],[f553,f83])).
fof(f553,plain,(
  ( ! [X1] : (mult(X1,X1) = rd(X1,i(X1))) )),
  inference(superposition,[],[f4,f457])).
fof(f3622,plain,(
  ( ! [X1] : (rd(i(X1),X1) = i(mult(X1,X1))) )),
  inference(superposition,[],[f4,f3223])).
fof(f7727,plain,(
  ( ! [X4] : (mult(i(mult(X4,X4)),i(X4)) = ld(mult(X4,X4),i(X4))) )),
  inference(backward_demodulation,[],[f7726,f3665])).
fof(f3665,plain,(
  ( ! [X4] : (mult(i(mult(X4,X4)),i(X4)) = mult(i(X4),i(mult(X4,X4)))) )),
  inference(backward_demodulation,[],[f3661,f1658])).
fof(f1658,plain,(
  ( ! [X4] : (mult(ld(X4,i(X4)),i(X4)) = mult(i(X4),ld(X4,i(X4)))) )),
  inference(backward_demodulation,[],[f1655,f905])).
fof(f905,plain,(
  ( ! [X4] : (mult(rd(i(X4),X4),i(X4)) = mult(i(X4),rd(i(X4),X4))) )),
  inference(superposition,[],[f185,f797])).
fof(f7726,plain,(
  ( ! [X4] : (mult(i(X4),i(mult(X4,X4))) = ld(mult(X4,X4),i(X4))) )),
  inference(backward_demodulation,[],[f7725,f7346])).
fof(f7346,plain,(
  ( ! [X4] : (mult(i(X4),i(mult(X4,X4))) = rd(mult(i(X4),i(mult(X4,mult(X4,mult(X4,X4))))),i(mult(X4,X4)))) )),
  inference(forward_demodulation,[],[f7345,f3223])).
fof(f7345,plain,(
  ( ! [X4] : (mult(mult(i(mult(X4,X4)),X4),i(mult(X4,X4))) = rd(mult(i(X4),i(mult(X4,mult(X4,mult(X4,X4))))),i(mult(X4,X4)))) )),
  inference(forward_demodulation,[],[f7344,f185])).
fof(f7344,plain,(
  ( ! [X4] : (mult(mult(i(mult(X4,X4)),X4),i(mult(X4,X4))) = rd(mult(i(X4),i(mult(X4,mult(mult(X4,X4),X4)))),i(mult(X4,X4)))) )),
  inference(forward_demodulation,[],[f7343,f63])).
fof(f7343,plain,(
  ( ! [X4] : (mult(mult(i(mult(X4,X4)),X4),i(mult(X4,X4))) = rd(mult(i(X4),i(mult(mult(X4,X4),mult(X4,X4)))),i(mult(X4,X4)))) )),
  inference(forward_demodulation,[],[f7342,f3662])).
fof(f7342,plain,(
  ( ! [X4] : (mult(mult(i(mult(X4,X4)),X4),i(mult(X4,X4))) = rd(mult(i(X4),mult(i(mult(X4,X4)),i(mult(X4,X4)))),i(mult(X4,X4)))) )),
  inference(forward_demodulation,[],[f7293,f3663])).
fof(f3663,plain,(
  ( ! [X0,X1] : (mult(i(mult(X0,X0)),mult(i(X0),X1)) = mult(i(X0),mult(i(mult(X0,X0)),X1))) )),
  inference(backward_demodulation,[],[f3661,f1656])).
fof(f1656,plain,(
  ( ! [X0,X1] : (mult(ld(X0,i(X0)),mult(i(X0),X1)) = mult(i(X0),mult(ld(X0,i(X0)),X1))) )),
  inference(backward_demodulation,[],[f1655,f903])).
fof(f903,plain,(
  ( ! [X0,X1] : (mult(rd(i(X0),X0),mult(i(X0),X1)) = mult(i(X0),mult(rd(i(X0),X0),X1))) )),
  inference(superposition,[],[f63,f797])).
fof(f7293,plain,(
  ( ! [X4] : (mult(mult(i(mult(X4,X4)),X4),i(mult(X4,X4))) = rd(mult(i(mult(X4,X4)),mult(i(X4),i(mult(X4,X4)))),i(mult(X4,X4)))) )),
  inference(superposition,[],[f251,f4694])).
fof(f4694,plain,(
  ( ! [X3] : (i(X3) = mult(X3,i(mult(X3,X3)))) )),
  inference(superposition,[],[f1,f3661])).
fof(f7725,plain,(
  ( ! [X4] : (rd(mult(i(X4),i(mult(X4,mult(X4,mult(X4,X4))))),i(mult(X4,X4))) = ld(mult(X4,X4),i(X4))) )),
  inference(forward_demodulation,[],[f7724,f6815])).
fof(f6815,plain,(
  ( ! [X21] : (mult(X21,i(mult(X21,mult(X21,mult(X21,X21))))) = ld(mult(X21,X21),i(X21))) )),
  inference(forward_demodulation,[],[f6814,f185])).
fof(f6814,plain,(
  ( ! [X21] : (mult(X21,i(mult(X21,mult(mult(X21,X21),X21)))) = ld(mult(X21,X21),i(X21))) )),
  inference(forward_demodulation,[],[f6813,f63])).
fof(f6813,plain,(
  ( ! [X21] : (mult(X21,i(mult(mult(X21,X21),mult(X21,X21)))) = ld(mult(X21,X21),i(X21))) )),
  inference(forward_demodulation,[],[f6781,f4694])).
fof(f6781,plain,(
  ( ! [X21] : (mult(X21,i(mult(mult(X21,X21),mult(X21,X21)))) = ld(mult(X21,X21),mult(X21,i(mult(X21,X21))))) )),
  inference(superposition,[],[f184,f4694])).
fof(f184,plain,(
  ( ! [X14,X13] : (mult(X13,X14) = ld(mult(X13,X13),mult(X13,mult(mult(X13,X13),X14)))) )),
  inference(superposition,[],[f2,f63])).
fof(f7724,plain,(
  ( ! [X4] : (rd(mult(i(X4),i(mult(X4,mult(X4,mult(X4,X4))))),i(mult(X4,X4))) = mult(X4,i(mult(X4,mult(X4,mult(X4,X4)))))) )),
  inference(forward_demodulation,[],[f7723,f185])).
fof(f7723,plain,(
  ( ! [X4] : (rd(mult(i(X4),i(mult(X4,mult(mult(X4,X4),X4)))),i(mult(X4,X4))) = mult(X4,i(mult(X4,mult(mult(X4,X4),X4))))) )),
  inference(forward_demodulation,[],[f7722,f63])).
fof(f7722,plain,(
  ( ! [X4] : (rd(mult(i(X4),i(mult(mult(X4,X4),mult(X4,X4)))),i(mult(X4,X4))) = mult(X4,i(mult(mult(X4,X4),mult(X4,X4))))) )),
  inference(forward_demodulation,[],[f7668,f3662])).
fof(f7668,plain,(
  ( ! [X4] : (rd(mult(i(X4),mult(i(mult(X4,X4)),i(mult(X4,X4)))),i(mult(X4,X4))) = mult(X4,mult(i(mult(X4,X4)),i(mult(X4,X4))))) )),
  inference(superposition,[],[f312,f4694])).
fof(f312,plain,(
  ( ! [X6,X5] : (mult(X5,mult(X6,X6)) = rd(mult(mult(X5,X6),mult(X6,X6)),X6)) )),
  inference(superposition,[],[f4,f107])).
fof(f34,plain,(
  ( ! [X2,X3] : (rd(X3,ld(X2,X3)) = X2) )),
  inference(superposition,[],[f4,f1])).
fof(f97,plain,(
  ( ! [X8,X7,X9] : (mult(X7,X8) = mult(mult(rd(X7,mult(X8,mult(X8,X9))),X8),mult(X8,mult(X9,X8)))) )),
  inference(superposition,[],[f8,f3])).
fof(f13485,plain,(
  ( ! [X10,X11] : (mult(X10,mult(i(X10),X11)) = ld(X10,ld(i(X10),X11))) )),
  inference(superposition,[],[f12637,f1])).
fof(f12637,plain,(
  ( ! [X6,X5] : (ld(X5,X6) = mult(X5,mult(i(X5),mult(i(X5),X6)))) )),
  inference(forward_demodulation,[],[f12581,f4])).
fof(f12581,plain,(
  ( ! [X6,X5] : (mult(X5,mult(i(X5),mult(i(X5),X6))) = rd(mult(ld(X5,X6),i(X5)),i(X5))) )),
  inference(backward_demodulation,[],[f12465,f1598])).
fof(f1598,plain,(
  ( ! [X6,X5] : (mult(X5,mult(i(X5),mult(i(X5),X6))) = rd(mult(i(X5),mult(X6,i(X5))),i(X5))) )),
  inference(forward_demodulation,[],[f1547,f6])).
fof(f1547,plain,(
  ( ! [X6,X5] : (mult(X5,mult(i(X5),mult(i(X5),X6))) = rd(mult(unit,mult(i(X5),mult(X6,i(X5)))),i(X5))) )),
  inference(superposition,[],[f104,f9])).
fof(f104,plain,(
  ( ! [X12,X13,X11] : (mult(X11,mult(X12,mult(X12,X13))) = rd(mult(mult(X11,X12),mult(X12,mult(X13,X12))),X12)) )),
  inference(superposition,[],[f4,f8])).
fof(f14015,plain,(
  ( ! [X0,X1] : (mult(i(X0),mult(X0,X1)) = ld(i(X0),ld(X0,X1))) )),
  inference(superposition,[],[f13332,f1])).
fof(f13332,plain,(
  ( ! [X45,X44] : (mult(i(X44),mult(X44,mult(X44,X45))) = ld(i(X44),X45)) )),
  inference(backward_demodulation,[],[f13315,f11388])).
fof(f11388,plain,(
  ( ! [X45,X44] : (mult(i(X44),mult(X44,mult(X44,X45))) = ld(X44,mult(mult(X44,X44),X45))) )),
  inference(forward_demodulation,[],[f11329,f6289])).
fof(f11329,plain,(
  ( ! [X45,X44] : (mult(i(X44),mult(X44,mult(X44,X45))) = mult(mult(X44,X44),ld(X44,X45))) )),
  inference(backward_demodulation,[],[f10947,f7373])).
fof(f7373,plain,(
  ( ! [X45,X44] : (mult(i(X44),mult(X44,mult(X44,X45))) = mult(mult(X44,rd(X45,X44)),X44)) )),
  inference(backward_demodulation,[],[f7315,f1604])).
fof(f1604,plain,(
  ( ! [X45,X44] : (mult(i(X44),mult(X44,mult(X44,X45))) = rd(mult(X44,mult(X45,X44)),X44)) )),
  inference(forward_demodulation,[],[f1560,f6])).
fof(f1560,plain,(
  ( ! [X45,X44] : (mult(i(X44),mult(X44,mult(X44,X45))) = rd(mult(unit,mult(X44,mult(X45,X44))),X44)) )),
  inference(superposition,[],[f104,f10])).
fof(f14631,plain,(
  ( ! [X2,X3] : (ld(X2,ld(i(X2),mult(X3,i(X2)))) = mult(ld(i(X2),ld(X2,X3)),i(X2))) )),
  inference(forward_demodulation,[],[f14630,f83])).
fof(f14630,plain,(
  ( ! [X2,X3] : (ld(X2,ld(i(X2),mult(X3,i(X2)))) = mult(ld(i(X2),ld(i(i(X2)),X3)),i(X2))) )),
  inference(forward_demodulation,[],[f14543,f13485])).
fof(f14543,plain,(
  ( ! [X2,X3] : (ld(X2,ld(i(X2),mult(X3,i(X2)))) = mult(mult(i(X2),mult(i(i(X2)),X3)),i(X2))) )),
  inference(superposition,[],[f14142,f11547])).
fof(f14142,plain,(
  ( ! [X12,X13] : (mult(mult(i(X13),rd(X12,i(X13))),i(X13)) = ld(X13,X12)) )),
  inference(backward_demodulation,[],[f14141,f12610])).
fof(f12610,plain,(
  ( ! [X12,X13] : (mult(mult(i(X13),rd(X12,i(X13))),i(X13)) = mult(ld(X13,mult(X12,X13)),i(X13))) )),
  inference(backward_demodulation,[],[f12465,f11123])).
fof(f11123,plain,(
  ( ! [X12,X13] : (mult(mult(i(X13),rd(X12,i(X13))),i(X13)) = mult(i(X13),mult(mult(X12,X13),i(X13)))) )),
  inference(backward_demodulation,[],[f11076,f1857])).
fof(f1857,plain,(
  ( ! [X12,X13] : (mult(mult(i(X13),rd(X12,i(X13))),i(X13)) = mult(i(X13),mult(mult(X12,i(X13)),X13))) )),
  inference(superposition,[],[f81,f3])).
fof(f14544,plain,(
  ( ! [X4,X5] : (ld(X5,mult(X4,i(X5))) = mult(mult(i(X5),X4),i(X5))) )),
  inference(superposition,[],[f14142,f4])).
fof(f15680,plain,(
  ( ! [X2,X3] : (rd(X3,X2) = ld(X2,mult(mult(X2,X3),i(X2)))) )),
  inference(superposition,[],[f4,f14663])).
fof(f14663,plain,(
  ( ! [X88,X87] : (mult(ld(X87,mult(mult(X87,X88),i(X87))),X87) = X88) )),
  inference(backward_demodulation,[],[f14633,f13423])).
fof(f13423,plain,(
  ( ! [X88,X87] : (mult(mult(ld(X87,ld(i(X87),X88)),i(X87)),X87) = X88) )),
  inference(forward_demodulation,[],[f13340,f2])).
fof(f13340,plain,(
  ( ! [X88,X87] : (mult(mult(ld(X87,ld(i(X87),X88)),i(X87)),X87) = ld(i(X87),mult(i(X87),X88))) )),
  inference(backward_demodulation,[],[f13315,f12770])).
fof(f12770,plain,(
  ( ! [X88,X87] : (ld(X87,mult(mult(X87,X87),mult(i(X87),X88))) = mult(mult(ld(X87,ld(i(X87),X88)),i(X87)),X87)) )),
  inference(forward_demodulation,[],[f12769,f12609])).
fof(f12769,plain,(
  ( ! [X88,X87] : (mult(mult(mult(X87,mult(i(X87),X88)),i(X87)),X87) = ld(X87,mult(mult(X87,X87),mult(i(X87),X88)))) )),
  inference(forward_demodulation,[],[f12768,f11376])).
fof(f11376,plain,(
  ( ! [X21,X22] : (mult(X22,mult(mult(X21,X22),i(X22))) = ld(X22,mult(mult(X22,X22),X21))) )),
  inference(forward_demodulation,[],[f11320,f6289])).
fof(f11320,plain,(
  ( ! [X21,X22] : (mult(X22,mult(mult(X21,X22),i(X22))) = mult(mult(X22,X22),ld(X22,X21))) )),
  inference(backward_demodulation,[],[f10947,f483])).
fof(f483,plain,(
  ( ! [X21,X22] : (mult(mult(X22,rd(X21,X22)),X22) = mult(X22,mult(mult(X21,X22),i(X22)))) )),
  inference(superposition,[],[f76,f3])).
fof(f76,plain,(
  ( ! [X2,X3] : (mult(mult(X2,X3),X2) = mult(X2,mult(mult(mult(X3,X2),X2),i(X2)))) )),
  inference(forward_demodulation,[],[f55,f5])).
fof(f55,plain,(
  ( ! [X2,X3] : (mult(X2,mult(mult(mult(X3,X2),X2),i(X2))) = mult(mult(mult(X2,X3),X2),unit)) )),
  inference(superposition,[],[f7,f9])).
fof(f12768,plain,(
  ( ! [X88,X87] : (mult(mult(mult(X87,mult(i(X87),X88)),i(X87)),X87) = mult(X87,mult(mult(mult(i(X87),X88),X87),i(X87)))) )),
  inference(forward_demodulation,[],[f12767,f11076])).
fof(f12767,plain,(
  ( ! [X88,X87] : (mult(mult(mult(X87,mult(i(X87),X88)),i(X87)),X87) = mult(X87,mult(mult(mult(i(X87),X88),i(X87)),X87))) )),
  inference(forward_demodulation,[],[f12501,f83])).
fof(f12501,plain,(
  ( ! [X88,X87] : (mult(mult(mult(X87,mult(i(X87),X88)),i(X87)),i(i(X87))) = mult(i(i(X87)),mult(mult(mult(i(X87),X88),i(X87)),i(i(X87))))) )),
  inference(superposition,[],[f11850,f11620])).
fof(f615,plain,(
  ( ! [X4,X3] : (mult(i(mult(X3,mult(X3,X4))),X3) = rd(X3,mult(X3,mult(X4,X3)))) )),
  inference(superposition,[],[f4,f123])).
fof(f123,plain,(
  ( ! [X12,X13] : (mult(mult(i(mult(X12,mult(X12,X13))),X12),mult(X12,mult(X13,X12))) = X12) )),
  inference(forward_demodulation,[],[f99,f6])).
fof(f99,plain,(
  ( ! [X12,X13] : (mult(mult(i(mult(X12,mult(X12,X13))),X12),mult(X12,mult(X13,X12))) = mult(unit,X12)) )),
  inference(superposition,[],[f8,f10])).
fof(f23024,plain,(
  ( ! [X43,X42] : (mult(X42,mult(X43,i(X42))) = rd(mult(X42,X43),X42)) )),
  inference(superposition,[],[f4,f17832])).
fof(f17832,plain,(
  ( ! [X132,X131] : (mult(X131,X132) = mult(mult(X131,mult(X132,i(X131))),X131)) )),
  inference(forward_demodulation,[],[f17831,f1])).
fof(f17831,plain,(
  ( ! [X132,X131] : (mult(mult(X131,mult(X132,i(X131))),X131) = mult(X131,mult(X131,ld(X131,X132)))) )),
  inference(forward_demodulation,[],[f17670,f3])).
fof(f17670,plain,(
  ( ! [X132,X131] : (mult(X131,mult(X131,mult(rd(ld(X131,X132),X131),X131))) = mult(mult(X131,mult(X132,i(X131))),X131)) )),
  inference(superposition,[],[f122,f16429])).
fof(f28087,plain,(
  ( ! [X30,X29] : (mult(X29,ld(X30,X29)) = rd(X29,ld(X29,X30))) )),
  inference(superposition,[],[f4,f26517])).
fof(f26517,plain,(
  ( ! [X2,X1] : (mult(mult(X1,ld(X2,X1)),ld(X1,X2)) = X1) )),
  inference(superposition,[],[f26157,f34])).
fof(f26157,plain,(
  ( ! [X200,X201] : (mult(mult(X200,X201),ld(X200,rd(X200,X201))) = X200) )),
  inference(forward_demodulation,[],[f26156,f3])).
fof(f26156,plain,(
  ( ! [X200,X201] : (mult(mult(rd(mult(X200,X201),X200),X200),ld(X200,rd(X200,X201))) = X200) )),
  inference(forward_demodulation,[],[f25901,f83])).
fof(f25901,plain,(
  ( ! [X200,X201] : (mult(mult(i(i(rd(mult(X200,X201),X200))),X200),ld(X200,rd(X200,X201))) = X200) )),
  inference(superposition,[],[f20089,f23521])).
fof(f20089,plain,(
  ( ! [X0,X1] : (mult(mult(i(X1),X0),ld(X0,mult(X1,X0))) = X0) )),
  inference(backward_demodulation,[],[f19989,f18740])).
fof(f18740,plain,(
  ( ! [X0,X1] : (mult(mult(i(X1),X0),rd(ld(X0,X1),i(X0))) = X0) )),
  inference(superposition,[],[f16673,f1])).
fof(f16673,plain,(
  ( ! [X0,X1] : (mult(mult(i(mult(X0,X1)),X0),rd(X1,i(X0))) = X0) )),
  inference(backward_demodulation,[],[f16669,f8251])).
fof(f8251,plain,(
  ( ! [X0,X1] : (mult(mult(i(mult(X0,X1)),X0),ld(X0,mult(mult(X0,X1),X0))) = X0) )),
  inference(backward_demodulation,[],[f8094,f578])).
fof(f578,plain,(
  ( ! [X0,X1] : (mult(mult(i(mult(X0,X1)),X0),mult(X0,mult(ld(X0,X1),X0))) = X0) )),
  inference(superposition,[],[f123,f1])).
fof(f16669,plain,(
  ( ! [X2,X1] : (ld(X1,mult(mult(X1,X2),X1)) = rd(X2,i(X1))) )),
  inference(forward_demodulation,[],[f16633,f16433])).
fof(f16433,plain,(
  ( ! [X2,X1] : (mult(mult(X1,X2),X1) = rd(ld(i(X1),X2),i(X1))) )),
  inference(forward_demodulation,[],[f16402,f14354])).
fof(f14354,plain,(
  ( ! [X2,X3] : (mult(mult(X2,X3),X2) = ld(i(X2),mult(X3,X2))) )),
  inference(superposition,[],[f13323,f8540])).
fof(f16402,plain,(
  ( ! [X2,X1] : (ld(i(X1),mult(X2,X1)) = rd(ld(i(X1),X2),i(X1))) )),
  inference(superposition,[],[f16031,f83])).
fof(f16633,plain,(
  ( ! [X2,X1] : (rd(X2,i(X1)) = ld(X1,rd(ld(i(X1),X2),i(X1)))) )),
  inference(superposition,[],[f16245,f83])).
fof(f16245,plain,(
  ( ! [X30,X29] : (rd(X30,X29) = ld(i(X29),rd(ld(X29,X30),X29))) )),
  inference(forward_demodulation,[],[f15863,f16031])).
fof(f15863,plain,(
  ( ! [X30,X29] : (rd(X30,X29) = ld(i(X29),ld(X29,mult(X30,i(X29))))) )),
  inference(backward_demodulation,[],[f15680,f14722])).
fof(f14722,plain,(
  ( ! [X30,X29] : (ld(X29,mult(mult(X29,X30),i(X29))) = ld(i(X29),ld(X29,mult(X30,i(X29))))) )),
  inference(backward_demodulation,[],[f14544,f14659])).
fof(f14659,plain,(
  ( ! [X30,X29] : (ld(i(X29),mult(mult(i(X29),X30),i(X29))) = ld(X29,mult(mult(X29,X30),i(X29)))) )),
  inference(backward_demodulation,[],[f14633,f13018])).
fof(f13018,plain,(
  ( ! [X30,X29] : (mult(ld(X29,ld(i(X29),X30)),i(X29)) = ld(i(X29),mult(mult(i(X29),X30),i(X29)))) )),
  inference(superposition,[],[f2,f12674])).
fof(f19989,plain,(
  ( ! [X0,X1] : (ld(X0,mult(X1,X0)) = rd(ld(X0,X1),i(X0))) )),
  inference(superposition,[],[f16669,f1])).
fof(f14141,plain,(
  ( ! [X2,X3] : (ld(X2,X3) = mult(ld(X2,mult(X3,X2)),i(X2))) )),
  inference(forward_demodulation,[],[f14140,f2])).
fof(f14140,plain,(
  ( ! [X2,X3] : (mult(ld(X2,mult(X3,X2)),i(X2)) = ld(X2,ld(X2,mult(X2,X3)))) )),
  inference(forward_demodulation,[],[f14139,f13629])).
fof(f13629,plain,(
  ( ! [X12,X13] : (mult(X12,X13) = ld(X12,ld(i(X12),ld(i(X12),X13)))) )),
  inference(forward_demodulation,[],[f13628,f2])).
fof(f13628,plain,(
  ( ! [X12,X13] : (ld(X12,ld(i(X12),ld(i(X12),X13))) = ld(X12,mult(X12,mult(X12,X13)))) )),
  inference(forward_demodulation,[],[f13627,f83])).
fof(f13627,plain,(
  ( ! [X12,X13] : (ld(X12,mult(i(i(X12)),mult(i(i(X12)),X13))) = ld(X12,ld(i(X12),ld(i(X12),X13)))) )),
  inference(forward_demodulation,[],[f13486,f13485])).
fof(f13486,plain,(
  ( ! [X12,X13] : (ld(X12,mult(i(i(X12)),mult(i(i(X12)),X13))) = mult(X12,mult(i(X12),ld(i(X12),X13)))) )),
  inference(superposition,[],[f12637,f12637])).
fof(f14139,plain,(
  ( ! [X2,X3] : (mult(ld(X2,mult(X3,X2)),i(X2)) = ld(X2,ld(X2,ld(X2,ld(i(X2),ld(i(X2),X3)))))) )),
  inference(backward_demodulation,[],[f14132,f13392])).
fof(f13392,plain,(
  ( ! [X2,X3] : (mult(ld(X2,mult(X3,X2)),i(X2)) = ld(X2,ld(X2,ld(i(X2),ld(X2,ld(i(X2),X3)))))) )),
  inference(backward_demodulation,[],[f13337,f13126])).
fof(f13126,plain,(
  ( ! [X2,X3] : (mult(ld(X2,mult(X3,X2)),i(X2)) = ld(X2,mult(mult(ld(X2,ld(i(X2),X3)),X2),i(X2)))) )),
  inference(forward_demodulation,[],[f13125,f11076])).
fof(f13125,plain,(
  ( ! [X2,X3] : (ld(X2,mult(mult(ld(X2,ld(i(X2),X3)),i(X2)),X2)) = mult(ld(X2,mult(X3,X2)),i(X2))) )),
  inference(forward_demodulation,[],[f13124,f11062])).
fof(f13124,plain,(
  ( ! [X2,X3] : (ld(X2,mult(mult(ld(X2,ld(i(X2),X3)),i(X2)),X2)) = mult(mult(mult(i(X2),X3),X2),i(X2))) )),
  inference(forward_demodulation,[],[f13009,f11076])).
fof(f13009,plain,(
  ( ! [X2,X3] : (ld(X2,mult(mult(ld(X2,ld(i(X2),X3)),i(X2)),X2)) = mult(mult(mult(i(X2),X3),i(X2)),X2)) )),
  inference(superposition,[],[f11062,f12674])).
fof(f13337,plain,(
  ( ! [X6,X7] : (mult(mult(X7,X6),i(X6)) = ld(X6,ld(i(X6),X7))) )),
  inference(backward_demodulation,[],[f13315,f12797])).
fof(f12797,plain,(
  ( ! [X6,X7] : (mult(mult(X7,X6),i(X6)) = ld(X6,ld(X6,mult(mult(X6,X6),X7)))) )),
  inference(forward_demodulation,[],[f12796,f11076])).
fof(f12796,plain,(
  ( ! [X6,X7] : (mult(mult(X7,i(X6)),X6) = ld(X6,ld(X6,mult(mult(X6,X6),X7)))) )),
  inference(forward_demodulation,[],[f12513,f12773])).
fof(f12773,plain,(
  ( ! [X90,X89] : (mult(mult(mult(X89,X90),i(X89)),X89) = ld(X89,mult(mult(X89,X89),X90))) )),
  inference(forward_demodulation,[],[f12772,f11376])).
fof(f12772,plain,(
  ( ! [X90,X89] : (mult(mult(mult(X89,X90),i(X89)),X89) = mult(X89,mult(mult(X90,X89),i(X89)))) )),
  inference(forward_demodulation,[],[f12771,f11076])).
fof(f12771,plain,(
  ( ! [X90,X89] : (mult(mult(mult(X89,X90),i(X89)),X89) = mult(X89,mult(mult(X90,i(X89)),X89))) )),
  inference(forward_demodulation,[],[f12502,f83])).
fof(f12502,plain,(
  ( ! [X90,X89] : (mult(mult(mult(X89,X90),i(X89)),i(i(X89))) = mult(i(i(X89)),mult(mult(X90,i(X89)),i(i(X89))))) )),
  inference(superposition,[],[f11850,f11850])).
fof(f12513,plain,(
  ( ! [X6,X7] : (mult(mult(X7,i(X6)),X6) = ld(X6,mult(mult(mult(X6,X7),i(X6)),X6))) )),
  inference(superposition,[],[f11062,f11850])).
fof(f179073,plain,(
  ( ! [X0,X1] : (i(X0) = ld(X1,i(rd(X0,X1)))) )),
  inference(forward_demodulation,[],[f179066,f36318])).
fof(f36318,plain,(
  ( ! [X4,X5] : (i(X4) = ld(ld(X4,mult(X5,X4)),ld(X4,X5))) )),
  inference(superposition,[],[f2,f14141])).
fof(f179066,plain,(
  ( ! [X0,X1] : (ld(ld(X0,mult(X1,X0)),ld(X0,X1)) = ld(X1,i(rd(X0,X1)))) )),
  inference(backward_demodulation,[],[f178999,f172337])).
fof(f172337,plain,(
  ( ! [X0,X1] : (ld(X1,i(rd(X0,X1))) = ld(ld(i(ld(X0,X1)),X0),ld(X0,X1))) )),
  inference(forward_demodulation,[],[f172280,f68371])).
fof(f172280,plain,(
  ( ! [X0,X1] : (mult(ld(X0,X1),i(X1)) = ld(ld(i(ld(X0,X1)),X0),ld(X0,X1))) )),
  inference(superposition,[],[f22238,f1])).
fof(f178999,plain,(
  ( ! [X10,X11] : (ld(X10,mult(X11,X10)) = ld(i(ld(X10,X11)),X10)) )),
  inference(superposition,[],[f178825,f14141])).
fof(f178825,plain,(
  ( ! [X17,X16] : (ld(i(mult(X16,i(X17))),X17) = X16) )),
  inference(superposition,[],[f177421,f83])).
fof(f177421,plain,(
  ( ! [X15,X16] : (i(X16) = ld(i(mult(i(X16),i(X15))),X15)) )),
  inference(forward_demodulation,[],[f177395,f85659])).
fof(f85659,plain,(
  ( ! [X26,X25] : (i(mult(i(X25),X26)) = rd(ld(X26,mult(X25,X25)),X25)) )),
  inference(backward_demodulation,[],[f85658,f32180])).
fof(f32180,plain,(
  ( ! [X26,X25] : (i(mult(i(X25),X26)) = rd(ld(mult(mult(i(X25),X26),mult(i(X25),X26)),ld(X25,mult(X26,X25))),X25)) )),
  inference(superposition,[],[f31088,f11546])).
fof(f11546,plain,(
  ( ! [X62,X63] : (ld(mult(i(X62),X63),ld(X62,mult(X63,X62))) = X62) )),
  inference(superposition,[],[f2,f11062])).
fof(f31088,plain,(
  ( ! [X17,X16] : (i(X16) = rd(ld(mult(X16,X16),X17),ld(X16,X17))) )),
  inference(backward_demodulation,[],[f31049,f20353])).
fof(f20353,plain,(
  ( ! [X17,X16] : (i(X16) = rd(ld(X16,mult(i(X16),X17)),ld(X16,X17))) )),
  inference(backward_demodulation,[],[f20347,f17064])).
fof(f17064,plain,(
  ( ! [X17,X16] : (i(X16) = rd(ld(X16,mult(ld(X16,rd(X17,X16)),X16)),ld(X16,X17))) )),
  inference(forward_demodulation,[],[f17063,f15919])).
fof(f15919,plain,(
  ( ! [X10,X9] : (ld(X9,ld(X9,ld(i(X9),X10))) = mult(ld(X9,rd(X10,X9)),X9)) )),
  inference(backward_demodulation,[],[f15841,f13394])).
fof(f13394,plain,(
  ( ! [X10,X9] : (mult(mult(ld(X9,X10),i(X9)),X9) = ld(X9,ld(X9,ld(i(X9),X10)))) )),
  inference(backward_demodulation,[],[f13337,f12059])).
fof(f12059,plain,(
  ( ! [X10,X9] : (ld(X9,mult(mult(X10,X9),i(X9))) = mult(mult(ld(X9,X10),i(X9)),X9)) )),
  inference(forward_demodulation,[],[f12058,f83])).
fof(f12058,plain,(
  ( ! [X10,X9] : (mult(mult(ld(i(i(X9)),X10),i(X9)),X9) = ld(X9,mult(mult(X10,X9),i(X9)))) )),
  inference(forward_demodulation,[],[f12035,f11076])).
fof(f12035,plain,(
  ( ! [X10,X9] : (mult(mult(ld(i(i(X9)),X10),i(X9)),X9) = ld(X9,mult(mult(X10,i(X9)),X9))) )),
  inference(superposition,[],[f11505,f11505])).
fof(f15841,plain,(
  ( ! [X97,X96] : (mult(ld(X96,X97),i(X96)) = ld(X96,rd(X97,X96))) )),
  inference(backward_demodulation,[],[f15680,f14665])).
fof(f14665,plain,(
  ( ! [X97,X96] : (mult(ld(X96,X97),i(X96)) = ld(X96,ld(X96,mult(mult(X96,X97),i(X96))))) )),
  inference(backward_demodulation,[],[f14633,f14143])).
fof(f14143,plain,(
  ( ! [X97,X96] : (ld(X96,mult(ld(X96,ld(i(X96),X97)),i(X96))) = mult(ld(X96,X97),i(X96))) )),
  inference(backward_demodulation,[],[f14141,f12704])).
fof(f12704,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = ld(X96,mult(ld(X96,ld(i(X96),X97)),i(X96)))) )),
  inference(backward_demodulation,[],[f12609,f11658])).
fof(f11658,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = ld(X96,mult(mult(X96,mult(i(X96),X97)),i(X96)))) )),
  inference(backward_demodulation,[],[f11617,f11316])).
fof(f11316,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = ld(X96,mult(mult(mult(X97,X96),i(X96)),i(X96)))) )),
  inference(forward_demodulation,[],[f11315,f11076])).
fof(f11315,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = ld(X96,mult(mult(mult(X97,i(X96)),X96),i(X96)))) )),
  inference(forward_demodulation,[],[f11314,f11076])).
fof(f11314,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = ld(X96,mult(mult(mult(X97,i(X96)),i(X96)),X96))) )),
  inference(forward_demodulation,[],[f11313,f11062])).
fof(f11313,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = mult(mult(i(X96),mult(mult(X97,i(X96)),i(X96))),X96)) )),
  inference(forward_demodulation,[],[f11312,f1309])).
fof(f11312,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = mult(mult(i(X96),mult(mult(X97,i(X96)),i(X96))),mult(i(X96),mult(X96,X96)))) )),
  inference(forward_demodulation,[],[f11311,f6])).
fof(f11311,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = mult(mult(i(X96),mult(mult(X97,i(X96)),i(X96))),mult(i(X96),mult(unit,mult(X96,X96))))) )),
  inference(forward_demodulation,[],[f11264,f9])).
fof(f11264,plain,(
  ( ! [X97,X96] : (mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96)) = mult(mult(i(X96),mult(mult(X97,i(X96)),i(X96))),mult(i(X96),mult(mult(X96,i(X96)),mult(X96,X96))))) )),
  inference(backward_demodulation,[],[f11252,f11174])).
fof(f11174,plain,(
  ( ! [X97,X96] : (mult(mult(i(X96),mult(mult(X97,i(X96)),i(X96))),mult(i(X96),mult(mult(X96,mult(X96,X96)),i(X96)))) = mult(mult(ld(X96,mult(X97,X96)),i(X96)),i(X96))) )),
  inference(forward_demodulation,[],[f11124,f11062])).
fof(f11124,plain,(
  ( ! [X97,X96] : (mult(mult(i(X96),mult(mult(X97,i(X96)),i(X96))),mult(i(X96),mult(mult(X96,mult(X96,X96)),i(X96)))) = mult(mult(mult(mult(i(X96),X97),X96),i(X96)),i(X96))) )),
  inference(backward_demodulation,[],[f11076,f9787])).
fof(f9787,plain,(
  ( ! [X97,X96] : (mult(mult(i(X96),mult(mult(X97,i(X96)),i(X96))),mult(i(X96),mult(mult(X96,mult(X96,X96)),i(X96)))) = mult(mult(mult(mult(i(X96),X97),i(X96)),X96),i(X96))) )),
  inference(forward_demodulation,[],[f9663,f1497])).
fof(f1497,plain,(
  ( ! [X6,X5] : (mult(i(X5),mult(mult(mult(X6,i(X5)),i(X5)),mult(X5,X5))) = mult(mult(mult(i(X5),X6),i(X5)),X5)) )),
  inference(superposition,[],[f7,f1309])).
fof(f9663,plain,(
  ( ! [X97,X96] : (mult(mult(i(X96),mult(mult(X97,i(X96)),i(X96))),mult(i(X96),mult(mult(X96,mult(X96,X96)),i(X96)))) = mult(mult(i(X96),mult(mult(mult(X97,i(X96)),i(X96)),mult(X96,X96))),i(X96))) )),
  inference(superposition,[],[f121,f8794])).
fof(f121,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X0,mult(mult(mult(X1,X0),X0),mult(X0,X2))),X0) = mult(mult(X0,mult(mult(X1,X0),X0)),mult(X0,mult(X2,X0)))) )),
  inference(forward_demodulation,[],[f95,f75])).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (mult(mult(mult(mult(X0,X1),X0),X0),mult(X0,mult(X2,X0))) = mult(mult(X0,mult(mult(mult(X1,X0),X0),mult(X0,X2))),X0)) )),
  inference(superposition,[],[f8,f7])).
fof(f17063,plain,(
  ( ! [X17,X16] : (i(X16) = rd(ld(X16,ld(X16,ld(X16,ld(i(X16),X17)))),ld(X16,X17))) )),
  inference(forward_demodulation,[],[f17062,f17034])).
fof(f17034,plain,(
  ( ! [X0,X1] : (ld(X0,ld(i(X0),X1)) = rd(rd(X1,i(X0)),X0)) )),
  inference(forward_demodulation,[],[f16978,f14132])).
fof(f16978,plain,(
  ( ! [X0,X1] : (ld(i(X0),ld(X0,X1)) = rd(rd(X1,i(X0)),X0)) )),
  inference(superposition,[],[f16245,f16416])).
fof(f16416,plain,(
  ( ! [X26,X27] : (rd(ld(X27,rd(X26,i(X27))),X27) = ld(X27,X26)) )),
  inference(superposition,[],[f16031,f3])).
fof(f17062,plain,(
  ( ! [X17,X16] : (i(X16) = rd(ld(X16,ld(X16,rd(rd(X17,i(X16)),X16))),ld(X16,X17))) )),
  inference(forward_demodulation,[],[f16985,f16439])).
fof(f16439,plain,(
  ( ! [X2,X3] : (ld(X2,ld(X2,rd(X3,X2))) = rd(ld(X2,ld(X2,X3)),X2)) )),
  inference(forward_demodulation,[],[f16438,f83])).
fof(f16438,plain,(
  ( ! [X2,X3] : (rd(ld(X2,ld(i(i(X2)),X3)),X2) = ld(X2,ld(X2,rd(X3,X2)))) )),
  inference(forward_demodulation,[],[f16405,f15917])).
fof(f15917,plain,(
  ( ! [X0,X1] : (mult(i(X0),mult(X1,i(X0))) = ld(X0,rd(X1,X0))) )),
  inference(backward_demodulation,[],[f15841,f12465])).
fof(f16405,plain,(
  ( ! [X2,X3] : (rd(ld(X2,ld(i(i(X2)),X3)),X2) = ld(X2,mult(i(X2),mult(X3,i(X2))))) )),
  inference(superposition,[],[f16031,f12042])).
fof(f16985,plain,(
  ( ! [X17,X16] : (i(X16) = rd(rd(ld(X16,ld(X16,rd(X17,i(X16)))),X16),ld(X16,X17))) )),
  inference(superposition,[],[f16226,f16416])).
fof(f16226,plain,(
  ( ! [X6,X7] : (i(X6) = rd(rd(ld(X6,X7),X6),rd(X7,X6))) )),
  inference(forward_demodulation,[],[f15844,f16031])).
fof(f15844,plain,(
  ( ! [X6,X7] : (i(X6) = rd(ld(X6,mult(X7,i(X6))),rd(X7,X6))) )),
  inference(backward_demodulation,[],[f15680,f14703])).
fof(f14703,plain,(
  ( ! [X6,X7] : (i(X6) = rd(ld(X6,mult(X7,i(X6))),ld(X6,mult(mult(X6,X7),i(X6))))) )),
  inference(backward_demodulation,[],[f14544,f14635])).
fof(f14635,plain,(
  ( ! [X6,X7] : (i(X6) = rd(mult(mult(i(X6),X7),i(X6)),ld(X6,mult(mult(X6,X7),i(X6))))) )),
  inference(backward_demodulation,[],[f14633,f12675])).
fof(f12675,plain,(
  ( ! [X6,X7] : (i(X6) = rd(mult(mult(i(X6),X7),i(X6)),mult(ld(X6,ld(i(X6),X7)),i(X6)))) )),
  inference(backward_demodulation,[],[f12609,f11621])).
fof(f11621,plain,(
  ( ! [X6,X7] : (i(X6) = rd(mult(mult(i(X6),X7),i(X6)),mult(mult(X6,mult(i(X6),X7)),i(X6)))) )),
  inference(backward_demodulation,[],[f11617,f11130])).
fof(f11130,plain,(
  ( ! [X6,X7] : (i(X6) = rd(mult(mult(i(X6),X7),i(X6)),mult(mult(mult(X7,X6),i(X6)),i(X6)))) )),
  inference(forward_demodulation,[],[f11079,f11076])).
fof(f11079,plain,(
  ( ! [X6,X7] : (i(X6) = rd(mult(mult(i(X6),X7),i(X6)),mult(mult(mult(X7,i(X6)),X6),i(X6)))) )),
  inference(backward_demodulation,[],[f11076,f1868])).
fof(f1868,plain,(
  ( ! [X6,X7] : (i(X6) = rd(mult(mult(i(X6),X7),i(X6)),mult(mult(mult(X7,i(X6)),i(X6)),X6))) )),
  inference(superposition,[],[f4,f81])).
fof(f20347,plain,(
  ( ! [X26,X27] : (mult(i(X26),X27) = mult(ld(X26,rd(X27,X26)),X26)) )),
  inference(forward_demodulation,[],[f20346,f4])).
fof(f20346,plain,(
  ( ! [X26,X27] : (rd(mult(mult(i(X26),X27),i(X26)),i(X26)) = mult(ld(X26,rd(X27,X26)),X26)) )),
  inference(forward_demodulation,[],[f20064,f83])).
fof(f20064,plain,(
  ( ! [X26,X27] : (rd(mult(mult(i(X26),X27),i(X26)),i(X26)) = mult(ld(X26,rd(X27,i(i(X26)))),X26)) )),
  inference(superposition,[],[f16749,f16669])).
fof(f16749,plain,(
  ( ! [X41,X40] : (mult(ld(X40,ld(i(X40),X41)),X40) = rd(X41,i(X40))) )),
  inference(backward_demodulation,[],[f16669,f13594])).
fof(f13594,plain,(
  ( ! [X41,X40] : (ld(X40,mult(mult(X40,X41),X40)) = mult(ld(X40,ld(i(X40),X41)),X40)) )),
  inference(backward_demodulation,[],[f13485,f11276])).
fof(f11276,plain,(
  ( ! [X41,X40] : (ld(X40,mult(mult(X40,X41),X40)) = mult(mult(X40,mult(i(X40),X41)),X40)) )),
  inference(backward_demodulation,[],[f11275,f8267])).
fof(f8267,plain,(
  ( ! [X41,X40] : (mult(mult(i(X40),mult(X40,X41)),X40) = ld(X40,mult(mult(X40,X41),X40))) )),
  inference(backward_demodulation,[],[f8094,f1228])).
fof(f1228,plain,(
  ( ! [X41,X40] : (mult(X40,mult(ld(X40,X41),X40)) = mult(mult(i(X40),mult(X40,X41)),X40)) )),
  inference(forward_demodulation,[],[f1186,f6])).
fof(f1186,plain,(
  ( ! [X41,X40] : (mult(mult(i(X40),mult(X40,X41)),X40) = mult(unit,mult(X40,mult(ld(X40,X41),X40)))) )),
  inference(superposition,[],[f88,f10])).
fof(f31049,plain,(
  ( ! [X159,X160] : (ld(X159,mult(i(X159),X160)) = ld(mult(X159,X159),X160)) )),
  inference(backward_demodulation,[],[f30993,f25616])).
fof(f25616,plain,(
  ( ! [X159,X160] : (mult(i(X159),ld(X159,X160)) = ld(X159,mult(i(X159),X160))) )),
  inference(backward_demodulation,[],[f25615,f22722])).
fof(f22722,plain,(
  ( ! [X159,X160] : (mult(i(X159),ld(X159,X160)) = mult(i(mult(X159,X159)),X160)) )),
  inference(forward_demodulation,[],[f22721,f3662])).
fof(f22721,plain,(
  ( ! [X159,X160] : (mult(i(X159),ld(X159,X160)) = mult(mult(i(X159),i(X159)),X160)) )),
  inference(forward_demodulation,[],[f22720,f1494])).
fof(f22720,plain,(
  ( ! [X159,X160] : (mult(ld(i(i(X159)),i(X159)),X160) = mult(i(X159),ld(X159,X160))) )),
  inference(forward_demodulation,[],[f22527,f14430])).
fof(f14430,plain,(
  ( ! [X10,X9] : (ld(X9,X10) = mult(i(mult(X9,X9)),ld(i(X9),X10))) )),
  inference(forward_demodulation,[],[f14330,f83])).
fof(f14330,plain,(
  ( ! [X10,X9] : (ld(i(i(X9)),X10) = mult(i(mult(X9,X9)),ld(i(X9),X10))) )),
  inference(superposition,[],[f13323,f3662])).
fof(f22527,plain,(
  ( ! [X159,X160] : (mult(ld(i(i(X159)),i(X159)),X160) = mult(i(X159),mult(i(mult(X159,X159)),ld(i(X159),X160)))) )),
  inference(superposition,[],[f13325,f3662])).
fof(f13325,plain,(
  ( ! [X37,X35,X36] : (mult(X36,mult(mult(X35,X36),ld(X36,X37))) = mult(ld(i(X36),X35),X37)) )),
  inference(backward_demodulation,[],[f13315,f11377])).
fof(f11377,plain,(
  ( ! [X37,X35,X36] : (mult(X36,mult(mult(X35,X36),ld(X36,X37))) = mult(ld(X36,mult(mult(X36,X36),X35)),X37)) )),
  inference(forward_demodulation,[],[f11321,f6289])).
fof(f11321,plain,(
  ( ! [X37,X35,X36] : (mult(X36,mult(mult(X35,X36),ld(X36,X37))) = mult(mult(mult(X36,X36),ld(X36,X35)),X37)) )),
  inference(backward_demodulation,[],[f10947,f820])).
fof(f820,plain,(
  ( ! [X37,X35,X36] : (mult(mult(mult(X36,rd(X35,X36)),X36),X37) = mult(X36,mult(mult(X35,X36),ld(X36,X37)))) )),
  inference(superposition,[],[f56,f3])).
fof(f56,plain,(
  ( ! [X6,X4,X5] : (mult(X4,mult(mult(mult(X6,X4),X4),ld(X4,X5))) = mult(mult(mult(X4,X6),X4),X5)) )),
  inference(superposition,[],[f7,f1])).
fof(f25615,plain,(
  ( ! [X21,X20] : (mult(i(mult(X20,X20)),X21) = ld(X20,mult(i(X20),X21))) )),
  inference(forward_demodulation,[],[f25614,f2])).
fof(f25614,plain,(
  ( ! [X21,X20] : (mult(i(mult(X20,X20)),X21) = ld(i(X20),mult(i(X20),ld(X20,mult(i(X20),X21))))) )),
  inference(forward_demodulation,[],[f25613,f552])).
fof(f552,plain,(
  ( ! [X0] : (i(X0) = ld(mult(X0,X0),X0)) )),
  inference(superposition,[],[f2,f457])).
fof(f25613,plain,(
  ( ! [X21,X20] : (mult(i(mult(X20,X20)),X21) = ld(ld(mult(X20,X20),X20),mult(i(X20),ld(X20,mult(i(X20),X21))))) )),
  inference(forward_demodulation,[],[f25612,f83])).
fof(f25612,plain,(
  ( ! [X21,X20] : (mult(i(mult(X20,X20)),X21) = ld(ld(i(i(mult(X20,X20))),X20),mult(i(X20),ld(X20,mult(i(X20),X21))))) )),
  inference(forward_demodulation,[],[f25473,f22722])).
fof(f25473,plain,(
  ( ! [X21,X20] : (mult(i(mult(X20,X20)),X21) = ld(ld(i(i(mult(X20,X20))),X20),mult(i(mult(X20,X20)),mult(i(X20),X21)))) )),
  inference(superposition,[],[f13327,f4694])).
fof(f13327,plain,(
  ( ! [X33,X31,X32] : (mult(X32,X33) = ld(ld(i(X32),X31),mult(X32,mult(mult(X31,X32),X33)))) )),
  inference(backward_demodulation,[],[f13315,f11382])).
fof(f11382,plain,(
  ( ! [X33,X31,X32] : (mult(X32,X33) = ld(ld(X32,mult(mult(X32,X32),X31)),mult(X32,mult(mult(X31,X32),X33)))) )),
  inference(forward_demodulation,[],[f11323,f6289])).
fof(f11323,plain,(
  ( ! [X33,X31,X32] : (mult(X32,X33) = ld(mult(mult(X32,X32),ld(X32,X31)),mult(X32,mult(mult(X31,X32),X33)))) )),
  inference(backward_demodulation,[],[f10947,f1068])).
fof(f1068,plain,(
  ( ! [X33,X31,X32] : (mult(X32,X33) = ld(mult(mult(X32,rd(X31,X32)),X32),mult(X32,mult(mult(X31,X32),X33)))) )),
  inference(superposition,[],[f62,f3])).
fof(f30993,plain,(
  ( ! [X8,X7] : (mult(i(X7),ld(X7,X8)) = ld(mult(X7,X7),X8)) )),
  inference(superposition,[],[f1,f13923])).
fof(f13923,plain,(
  ( ! [X21,X22] : (ld(X21,X22) = ld(i(X21),ld(mult(X21,X21),X22))) )),
  inference(superposition,[],[f13315,f1])).
fof(f85658,plain,(
  ( ! [X37,X36] : (ld(mult(mult(i(X36),X37),mult(i(X36),X37)),ld(X36,mult(X37,X36))) = ld(X37,mult(X36,X36))) )),
  inference(forward_demodulation,[],[f85568,f68209])).
fof(f68209,plain,(
  ( ! [X14,X13] : (mult(i(mult(i(X13),X14)),X13) = ld(X14,mult(X13,X13))) )),
  inference(forward_demodulation,[],[f68017,f56377])).
fof(f56377,plain,(
  ( ! [X156,X157,X155] : (mult(X155,ld(ld(X155,mult(X156,X155)),X157)) = ld(X156,mult(X155,X157))) )),
  inference(forward_demodulation,[],[f56376,f1])).
fof(f56376,plain,(
  ( ! [X156,X157,X155] : (mult(X155,ld(ld(X155,mult(X156,X155)),X157)) = ld(mult(mult(X155,X155),ld(mult(X155,X155),X156)),mult(X155,X157))) )),
  inference(forward_demodulation,[],[f56375,f53420])).
fof(f53420,plain,(
  ( ! [X280,X279,X278] : (mult(mult(X280,X278),ld(mult(X278,X278),X279)) = mult(mult(X280,ld(X278,rd(X279,X278))),X278)) )),
  inference(forward_demodulation,[],[f52621,f48490])).
fof(f48490,plain,(
  ( ! [X4,X3] : (ld(X3,rd(X4,X3)) = rd(mult(i(X3),X4),X3)) )),
  inference(forward_demodulation,[],[f48489,f6])).
fof(f48489,plain,(
  ( ! [X4,X3] : (rd(mult(i(X3),X4),X3) = mult(unit,ld(X3,rd(X4,X3)))) )),
  inference(forward_demodulation,[],[f48488,f9])).
fof(f48488,plain,(
  ( ! [X4,X3] : (rd(mult(i(X3),X4),X3) = mult(mult(X3,i(X3)),ld(X3,rd(X4,X3)))) )),
  inference(forward_demodulation,[],[f48487,f15917])).
fof(f48487,plain,(
  ( ! [X4,X3] : (mult(mult(X3,i(X3)),mult(i(X3),mult(X4,i(X3)))) = rd(mult(i(X3),X4),X3)) )),
  inference(forward_demodulation,[],[f48226,f45798])).
fof(f45798,plain,(
  ( ! [X19,X20] : (rd(X20,X19) = mult(X19,rd(mult(i(X19),X20),X19))) )),
  inference(superposition,[],[f15074,f2])).
fof(f15074,plain,(
  ( ! [X41,X40] : (mult(X40,rd(X41,X40)) = rd(ld(i(X40),X41),X40)) )),
  inference(superposition,[],[f4,f14556])).
fof(f14556,plain,(
  ( ! [X2,X1] : (mult(mult(X1,rd(X2,X1)),X1) = ld(i(X1),X2)) )),
  inference(superposition,[],[f14142,f83])).
fof(f48226,plain,(
  ( ! [X4,X3] : (mult(mult(X3,i(X3)),mult(i(X3),mult(X4,i(X3)))) = mult(X3,rd(mult(i(X3),mult(i(X3),X4)),X3))) )),
  inference(superposition,[],[f17223,f8])).
fof(f17223,plain,(
  ( ! [X33,X34] : (mult(X33,rd(X34,X33)) = mult(mult(X33,X34),i(X33))) )),
  inference(backward_demodulation,[],[f17222,f13422])).
fof(f13422,plain,(
  ( ! [X33,X34] : (mult(mult(mult(X33,X33),mult(i(X33),X34)),i(X33)) = mult(mult(X33,X34),i(X33))) )),
  inference(forward_demodulation,[],[f13339,f11498])).
fof(f13339,plain,(
  ( ! [X33,X34] : (mult(mult(mult(X33,X33),mult(i(X33),X34)),i(X33)) = ld(i(X33),mult(X34,i(X33)))) )),
  inference(backward_demodulation,[],[f13315,f11389])).
fof(f11389,plain,(
  ( ! [X33,X34] : (mult(mult(mult(X33,X33),mult(i(X33),X34)),i(X33)) = ld(X33,mult(mult(X33,X33),mult(X34,i(X33))))) )),
  inference(forward_demodulation,[],[f11330,f6289])).
fof(f11330,plain,(
  ( ! [X33,X34] : (mult(mult(mult(X33,X33),mult(i(X33),X34)),i(X33)) = mult(mult(X33,X33),ld(X33,mult(X34,i(X33))))) )),
  inference(backward_demodulation,[],[f10947,f11166])).
fof(f11166,plain,(
  ( ! [X33,X34] : (mult(mult(mult(X33,X33),mult(i(X33),X34)),i(X33)) = mult(mult(X33,rd(mult(X34,i(X33)),X33)),X33)) )),
  inference(backward_demodulation,[],[f11165,f1354])).
fof(f1354,plain,(
  ( ! [X33,X34] : (mult(mult(mult(X33,X33),mult(i(X33),X34)),i(X33)) = mult(X33,mult(mult(X33,mult(i(X33),X34)),i(X33)))) )),
  inference(backward_demodulation,[],[f1352,f1183])).
fof(f1183,plain,(
  ( ! [X33,X34] : (mult(mult(mult(X33,X33),mult(i(X33),X34)),i(X33)) = mult(X33,mult(i(X33),mult(ld(i(X33),X34),i(X33))))) )),
  inference(superposition,[],[f88,f457])).
fof(f11165,plain,(
  ( ! [X94,X93] : (mult(X93,mult(mult(X93,mult(i(X93),X94)),i(X93))) = mult(mult(X93,rd(mult(X94,i(X93)),X93)),X93)) )),
  inference(forward_demodulation,[],[f11112,f483])).
fof(f11112,plain,(
  ( ! [X94,X93] : (mult(X93,mult(mult(X93,mult(i(X93),X94)),i(X93))) = mult(X93,mult(mult(mult(X94,i(X93)),X93),i(X93)))) )),
  inference(backward_demodulation,[],[f11076,f8663])).
fof(f8663,plain,(
  ( ! [X94,X93] : (mult(X93,mult(mult(mult(X94,i(X93)),i(X93)),X93)) = mult(X93,mult(mult(X93,mult(i(X93),X94)),i(X93)))) )),
  inference(forward_demodulation,[],[f8662,f1354])).
fof(f8662,plain,(
  ( ! [X94,X93] : (mult(mult(mult(X93,X93),mult(i(X93),X94)),i(X93)) = mult(X93,mult(mult(mult(X94,i(X93)),i(X93)),X93))) )),
  inference(forward_demodulation,[],[f8529,f1867])).
fof(f8529,plain,(
  ( ! [X94,X93] : (mult(mult(mult(X93,X93),mult(i(X93),X94)),i(X93)) = mult(X93,ld(i(X93),mult(mult(i(X93),X94),i(X93))))) )),
  inference(superposition,[],[f8249,f457])).
fof(f17222,plain,(
  ( ! [X107,X108] : (mult(mult(mult(X107,X107),mult(i(X107),X108)),i(X107)) = mult(X107,rd(X108,X107))) )),
  inference(forward_demodulation,[],[f17100,f83])).
fof(f17100,plain,(
  ( ! [X107,X108] : (mult(mult(mult(X107,X107),mult(i(X107),X108)),i(X107)) = mult(X107,rd(X108,i(i(X107))))) )),
  inference(superposition,[],[f16672,f457])).
fof(f16672,plain,(
  ( ! [X6,X4,X5] : (mult(mult(X6,mult(X4,X5)),X4) = mult(mult(X6,X4),rd(X5,i(X4)))) )),
  inference(backward_demodulation,[],[f16669,f8249])).
fof(f52621,plain,(
  ( ! [X280,X279,X278] : (mult(mult(X280,X278),ld(mult(X278,X278),X279)) = mult(mult(X280,rd(mult(i(X278),X279),X278)),X278)) )),
  inference(superposition,[],[f23510,f51024])).
fof(f51024,plain,(
  ( ! [X17,X18] : (mult(i(X17),X18) = mult(X17,ld(mult(X17,X17),X18))) )),
  inference(backward_demodulation,[],[f51023,f6778])).
fof(f6778,plain,(
  ( ! [X17,X18] : (mult(X17,ld(mult(X17,X17),X18)) = ld(mult(X17,X17),mult(X17,X18))) )),
  inference(superposition,[],[f184,f1])).
fof(f51023,plain,(
  ( ! [X167,X168] : (mult(i(X167),X168) = ld(mult(X167,X167),mult(X167,X168))) )),
  inference(forward_demodulation,[],[f51022,f20350])).
fof(f20350,plain,(
  ( ! [X33,X32] : (mult(i(X32),X33) = rd(rd(ld(X32,X33),X32),i(X32))) )),
  inference(backward_demodulation,[],[f20347,f16931])).
fof(f16931,plain,(
  ( ! [X33,X32] : (rd(rd(ld(X32,X33),X32),i(X32)) = mult(ld(X32,rd(X33,X32)),X32)) )),
  inference(forward_demodulation,[],[f16930,f15919])).
fof(f16930,plain,(
  ( ! [X33,X32] : (rd(rd(ld(X32,X33),X32),i(X32)) = ld(X32,ld(X32,ld(i(X32),X33)))) )),
  inference(forward_demodulation,[],[f16929,f83])).
fof(f16929,plain,(
  ( ! [X33,X32] : (rd(rd(ld(X32,X33),X32),i(X32)) = ld(i(i(X32)),ld(X32,ld(i(X32),X33)))) )),
  inference(forward_demodulation,[],[f16653,f15829])).
fof(f15829,plain,(
  ( ! [X37,X38] : (ld(X37,ld(i(X37),X38)) = rd(rd(X38,X37),i(X37))) )),
  inference(backward_demodulation,[],[f15680,f14134])).
fof(f14134,plain,(
  ( ! [X37,X38] : (ld(X37,ld(i(X37),X38)) = rd(ld(X37,mult(mult(X37,X38),i(X37))),i(X37))) )),
  inference(backward_demodulation,[],[f14132,f13706])).
fof(f13706,plain,(
  ( ! [X37,X38] : (rd(ld(X37,mult(mult(X37,X38),i(X37))),i(X37)) = ld(i(X37),ld(X37,X38))) )),
  inference(forward_demodulation,[],[f13671,f13323])).
fof(f13671,plain,(
  ( ! [X37,X38] : (rd(ld(X37,mult(mult(X37,X38),i(X37))),i(X37)) = mult(mult(X37,X37),ld(X37,ld(X37,X38)))) )),
  inference(backward_demodulation,[],[f13498,f13585])).
fof(f13585,plain,(
  ( ! [X37,X38] : (mult(mult(X37,X37),mult(i(X37),mult(i(X37),X38))) = rd(ld(X37,mult(mult(X37,X38),i(X37))),i(X37))) )),
  inference(backward_demodulation,[],[f13580,f12580])).
fof(f12580,plain,(
  ( ! [X37,X38] : (mult(mult(X37,X37),mult(i(X37),mult(i(X37),X38))) = rd(mult(X37,mult(ld(X37,X38),i(X37))),i(X37))) )),
  inference(backward_demodulation,[],[f12465,f1557])).
fof(f1557,plain,(
  ( ! [X37,X38] : (mult(mult(X37,X37),mult(i(X37),mult(i(X37),X38))) = rd(mult(X37,mult(i(X37),mult(X38,i(X37)))),i(X37))) )),
  inference(superposition,[],[f104,f457])).
fof(f13580,plain,(
  ( ! [X4,X3] : (ld(X3,mult(mult(X3,X4),i(X3))) = mult(X3,mult(ld(X3,X4),i(X3)))) )),
  inference(forward_demodulation,[],[f13480,f12465])).
fof(f13480,plain,(
  ( ! [X4,X3] : (ld(X3,mult(mult(X3,X4),i(X3))) = mult(X3,mult(i(X3),mult(X4,i(X3))))) )),
  inference(superposition,[],[f12637,f11850])).
fof(f13498,plain,(
  ( ! [X19,X20] : (mult(i(X19),mult(i(X19),X20)) = ld(X19,ld(X19,X20))) )),
  inference(superposition,[],[f2,f12637])).
fof(f16653,plain,(
  ( ! [X33,X32] : (rd(rd(ld(X32,X33),X32),i(X32)) = ld(i(i(X32)),rd(rd(X33,X32),i(X32)))) )),
  inference(superposition,[],[f16245,f16245])).
fof(f51022,plain,(
  ( ! [X167,X168] : (rd(rd(ld(X167,X168),X167),i(X167)) = ld(mult(X167,X167),mult(X167,X168))) )),
  inference(forward_demodulation,[],[f51021,f6778])).
fof(f51021,plain,(
  ( ! [X167,X168] : (rd(rd(ld(X167,X168),X167),i(X167)) = mult(X167,ld(mult(X167,X167),X168))) )),
  inference(forward_demodulation,[],[f51020,f83])).
fof(f51020,plain,(
  ( ! [X167,X168] : (rd(rd(ld(X167,X168),X167),i(X167)) = mult(i(i(X167)),ld(mult(X167,X167),X168))) )),
  inference(forward_demodulation,[],[f50816,f31096])).
fof(f31096,plain,(
  ( ! [X28,X27] : (rd(ld(X27,ld(X27,rd(X28,X27))),i(X27)) = ld(mult(X27,X27),X28)) )),
  inference(backward_demodulation,[],[f31049,f25619])).
fof(f25619,plain,(
  ( ! [X28,X27] : (rd(ld(X27,ld(X27,rd(X28,X27))),i(X27)) = ld(X27,mult(i(X27),X28))) )),
  inference(backward_demodulation,[],[f25616,f15934])).
fof(f15934,plain,(
  ( ! [X28,X27] : (mult(i(X27),ld(X27,X28)) = rd(ld(X27,ld(X27,rd(X28,X27))),i(X27))) )),
  inference(backward_demodulation,[],[f15841,f15421])).
fof(f15421,plain,(
  ( ! [X28,X27] : (mult(i(X27),ld(X27,X28)) = rd(ld(X27,mult(ld(X27,X28),i(X27))),i(X27))) )),
  inference(backward_demodulation,[],[f15306,f14750])).
fof(f14750,plain,(
  ( ! [X28,X27] : (mult(i(X27),ld(X27,X28)) = rd(mult(i(X27),ld(X27,mult(X28,i(X27)))),i(X27))) )),
  inference(backward_demodulation,[],[f14693,f13621])).
fof(f13621,plain,(
  ( ! [X28,X27] : (rd(mult(i(mult(X27,X27)),mult(X28,i(X27))),i(X27)) = mult(i(X27),ld(X27,X28))) )),
  inference(backward_demodulation,[],[f13611,f13134])).
fof(f13134,plain,(
  ( ! [X28,X27] : (mult(i(X27),mult(i(X27),ld(X27,ld(i(X27),X28)))) = rd(mult(i(mult(X27,X27)),mult(X28,i(X27))),i(X27))) )),
  inference(forward_demodulation,[],[f13017,f12838])).
fof(f12838,plain,(
  ( ! [X83,X82] : (mult(i(mult(X82,X82)),mult(X83,i(X82))) = mult(i(X82),mult(mult(i(X82),X83),i(X82)))) )),
  inference(forward_demodulation,[],[f12837,f11614])).
fof(f11614,plain,(
  ( ! [X123,X124] : (mult(mult(i(X123),X124),i(X123)) = mult(i(mult(X123,X123)),mult(mult(X123,X124),i(X123)))) )),
  inference(backward_demodulation,[],[f11498,f10916])).
fof(f10916,plain,(
  ( ! [X123,X124] : (mult(mult(i(X123),X124),i(X123)) = mult(i(mult(X123,X123)),ld(i(X123),mult(X124,i(X123))))) )),
  inference(superposition,[],[f8540,f3662])).
fof(f12837,plain,(
  ( ! [X83,X82] : (mult(i(X82),mult(i(mult(X82,X82)),mult(mult(X82,X83),i(X82)))) = mult(i(mult(X82,X82)),mult(X83,i(X82)))) )),
  inference(forward_demodulation,[],[f12539,f3662])).
fof(f12539,plain,(
  ( ! [X83,X82] : (mult(i(X82),mult(mult(i(X82),i(X82)),mult(mult(X82,X83),i(X82)))) = mult(mult(i(X82),i(X82)),mult(X83,i(X82)))) )),
  inference(superposition,[],[f63,f11850])).
fof(f13017,plain,(
  ( ! [X28,X27] : (mult(i(X27),mult(i(X27),ld(X27,ld(i(X27),X28)))) = rd(mult(i(X27),mult(mult(i(X27),X28),i(X27))),i(X27))) )),
  inference(superposition,[],[f392,f12674])).
fof(f392,plain,(
  ( ! [X10,X11] : (mult(X10,mult(X10,X11)) = rd(mult(X10,mult(X10,mult(X11,X10))),X10)) )),
  inference(superposition,[],[f4,f122])).
fof(f13611,plain,(
  ( ! [X33,X34] : (ld(X33,X34) = mult(i(X33),ld(X33,ld(i(X33),X34)))) )),
  inference(backward_demodulation,[],[f13485,f12828])).
fof(f12828,plain,(
  ( ! [X33,X34] : (ld(X33,X34) = mult(i(X33),mult(X33,mult(i(X33),X34)))) )),
  inference(forward_demodulation,[],[f12827,f11275])).
fof(f12827,plain,(
  ( ! [X33,X34] : (mult(i(X33),mult(i(X33),mult(X33,X34))) = ld(X33,X34)) )),
  inference(forward_demodulation,[],[f12826,f4])).
fof(f12826,plain,(
  ( ! [X33,X34] : (mult(i(X33),mult(i(X33),mult(X33,X34))) = rd(mult(ld(X33,X34),i(X33)),i(X33))) )),
  inference(forward_demodulation,[],[f12523,f12465])).
fof(f12523,plain,(
  ( ! [X33,X34] : (mult(i(X33),mult(i(X33),mult(X33,X34))) = rd(mult(i(X33),mult(X34,i(X33))),i(X33))) )),
  inference(superposition,[],[f392,f11850])).
fof(f14693,plain,(
  ( ! [X83,X82] : (mult(i(mult(X82,X82)),mult(X83,i(X82))) = mult(i(X82),ld(X82,mult(X83,i(X82))))) )),
  inference(backward_demodulation,[],[f14544,f12838])).
fof(f15306,plain,(
  ( ! [X0,X1] : (ld(X0,mult(ld(X0,X1),i(X0))) = mult(i(X0),ld(X0,mult(X1,i(X0))))) )),
  inference(superposition,[],[f14702,f1])).
fof(f50816,plain,(
  ( ! [X167,X168] : (rd(rd(ld(X167,X168),X167),i(X167)) = mult(i(i(X167)),rd(ld(X167,ld(X167,rd(X168,X167))),i(X167)))) )),
  inference(superposition,[],[f18530,f16047])).
fof(f16047,plain,(
  ( ! [X0,X1] : (ld(X0,ld(X0,rd(X1,X0))) = mult(i(X0),rd(ld(X0,X1),X0))) )),
  inference(backward_demodulation,[],[f16031,f15931])).
fof(f15931,plain,(
  ( ! [X0,X1] : (mult(i(X0),ld(X0,mult(X1,i(X0)))) = ld(X0,ld(X0,rd(X1,X0)))) )),
  inference(backward_demodulation,[],[f15841,f15306])).
fof(f18530,plain,(
  ( ! [X0,X1] : (rd(X0,X1) = mult(i(X1),rd(mult(X1,X0),X1))) )),
  inference(superposition,[],[f3,f16502])).
fof(f16502,plain,(
  ( ! [X2,X1] : (i(X1) = rd(rd(X2,X1),rd(mult(X1,X2),X1))) )),
  inference(superposition,[],[f16226,f2])).
fof(f23510,plain,(
  ( ! [X8,X7,X9] : (mult(mult(X9,X8),X7) = mult(mult(X9,rd(mult(X8,X7),X8)),X8)) )),
  inference(backward_demodulation,[],[f23024,f17118])).
fof(f17118,plain,(
  ( ! [X8,X7,X9] : (mult(mult(X9,mult(X8,mult(X7,i(X8)))),X8) = mult(mult(X9,X8),X7)) )),
  inference(superposition,[],[f16672,f4])).
fof(f56375,plain,(
  ( ! [X156,X157,X155] : (ld(mult(mult(X155,ld(X155,rd(X156,X155))),X155),mult(X155,X157)) = mult(X155,ld(ld(X155,mult(X156,X155)),X157))) )),
  inference(forward_demodulation,[],[f56092,f11062])).
fof(f56092,plain,(
  ( ! [X156,X157,X155] : (ld(mult(mult(X155,ld(X155,rd(X156,X155))),X155),mult(X155,X157)) = mult(X155,ld(mult(mult(i(X155),X156),X155),X157))) )),
  inference(superposition,[],[f1079,f20347])).
fof(f1079,plain,(
  ( ! [X10,X8,X9] : (mult(X9,ld(mult(mult(X8,X9),X9),X10)) = ld(mult(mult(X9,X8),X9),mult(X9,X10))) )),
  inference(superposition,[],[f62,f1])).
fof(f68017,plain,(
  ( ! [X14,X13] : (mult(X13,ld(ld(X13,mult(X14,X13)),X13)) = mult(i(mult(i(X13),X14)),X13)) )),
  inference(superposition,[],[f28270,f11547])).
fof(f85568,plain,(
  ( ! [X37,X36] : (ld(mult(mult(i(X36),X37),mult(i(X36),X37)),ld(X36,mult(X37,X36))) = mult(i(mult(i(X36),X37)),X36)) )),
  inference(superposition,[],[f30993,f11546])).
fof(f177395,plain,(
  ( ! [X15,X16] : (i(X16) = ld(rd(ld(i(X15),mult(X16,X16)),X16),X15)) )),
  inference(superposition,[],[f48758,f174805])).
fof(f174805,plain,(
  ( ! [X6,X5] : (rd(ld(i(X5),mult(X6,X6)),mult(X6,X6)) = X5) )),
  inference(forward_demodulation,[],[f174804,f13315])).
fof(f174804,plain,(
  ( ! [X6,X5] : (rd(ld(X5,mult(mult(X5,X5),mult(X6,X6))),mult(X6,X6)) = X5) )),
  inference(forward_demodulation,[],[f174803,f33052])).
fof(f33052,plain,(
  ( ! [X15,X16] : (ld(i(mult(X16,X16)),X15) = mult(mult(X16,X16),X15)) )),
  inference(forward_demodulation,[],[f33048,f24999])).
fof(f24999,plain,(
  ( ! [X30,X31,X29] : (mult(mult(X29,X29),X30) = rd(mult(X29,mult(X29,mult(mult(mult(X30,X29),X29),X31))),mult(X29,mult(X29,X31)))) )),
  inference(forward_demodulation,[],[f24823,f13957])).
fof(f13957,plain,(
  ( ! [X10,X9] : (mult(mult(X9,X9),X10) = ld(i(X9),mult(X9,X10))) )),
  inference(forward_demodulation,[],[f13917,f2])).
fof(f13917,plain,(
  ( ! [X10,X9] : (ld(i(X9),mult(X9,X10)) = ld(X9,mult(X9,mult(mult(X9,X9),X10)))) )),
  inference(superposition,[],[f13315,f63])).
fof(f24823,plain,(
  ( ! [X30,X31,X29] : (ld(i(X29),mult(X29,X30)) = rd(mult(X29,mult(X29,mult(mult(mult(X30,X29),X29),X31))),mult(X29,mult(X29,X31)))) )),
  inference(superposition,[],[f13326,f7])).
fof(f13326,plain,(
  ( ! [X33,X31,X32] : (rd(mult(X32,mult(mult(X31,X32),X33)),mult(X32,X33)) = ld(i(X32),X31)) )),
  inference(backward_demodulation,[],[f13315,f11381])).
fof(f11381,plain,(
  ( ! [X33,X31,X32] : (rd(mult(X32,mult(mult(X31,X32),X33)),mult(X32,X33)) = ld(X32,mult(mult(X32,X32),X31))) )),
  inference(forward_demodulation,[],[f11322,f6289])).
fof(f11322,plain,(
  ( ! [X33,X31,X32] : (rd(mult(X32,mult(mult(X31,X32),X33)),mult(X32,X33)) = mult(mult(X32,X32),ld(X32,X31))) )),
  inference(backward_demodulation,[],[f10947,f937])).
fof(f937,plain,(
  ( ! [X33,X31,X32] : (mult(mult(X32,rd(X31,X32)),X32) = rd(mult(X32,mult(mult(X31,X32),X33)),mult(X32,X33))) )),
  inference(superposition,[],[f61,f3])).
fof(f33048,plain,(
  ( ! [X15,X16] : (ld(i(mult(X16,X16)),X15) = rd(mult(X16,mult(X16,mult(mult(mult(X15,X16),X16),X16))),mult(X16,mult(X16,X16)))) )),
  inference(backward_demodulation,[],[f33046,f24992])).
fof(f24992,plain,(
  ( ! [X15,X16] : (ld(i(mult(X16,X16)),X15) = rd(mult(mult(X16,X16),mult(mult(X15,X16),mult(X16,X16))),mult(X16,mult(X16,X16)))) )),
  inference(forward_demodulation,[],[f24818,f185])).
fof(f24818,plain,(
  ( ! [X15,X16] : (ld(i(mult(X16,X16)),X15) = rd(mult(mult(X16,X16),mult(mult(X15,X16),mult(X16,X16))),mult(mult(X16,X16),X16))) )),
  inference(superposition,[],[f13326,f107])).
fof(f33046,plain,(
  ( ! [X4,X5] : (mult(X4,mult(X4,mult(mult(mult(X5,X4),X4),X4))) = mult(mult(X4,X4),mult(mult(X5,X4),mult(X4,X4)))) )),
  inference(forward_demodulation,[],[f33045,f33037])).
fof(f33037,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = mult(mult(X31,X30),mult(X30,X30))) )),
  inference(forward_demodulation,[],[f33036,f21010])).
fof(f21010,plain,(
  ( ! [X74,X75] : (mult(mult(rd(X75,i(X74)),X74),X74) = mult(mult(X75,X74),mult(X74,X74))) )),
  inference(forward_demodulation,[],[f21009,f1])).
fof(f21009,plain,(
  ( ! [X74,X75] : (mult(mult(rd(X75,i(X74)),X74),X74) = mult(X74,ld(X74,mult(mult(X75,X74),mult(X74,X74))))) )),
  inference(forward_demodulation,[],[f21008,f6956])).
fof(f6956,plain,(
  ( ! [X6,X7,X5] : (ld(X5,mult(mult(X6,X5),mult(X5,X7))) = mult(ld(X5,mult(mult(X6,X5),X5)),X7)) )),
  inference(backward_demodulation,[],[f6843,f734])).
fof(f734,plain,(
  ( ! [X6,X7,X5] : (mult(mult(mult(ld(X5,X6),X5),X5),X7) = ld(X5,mult(mult(X6,X5),mult(X5,X7)))) )),
  inference(superposition,[],[f2,f47])).
fof(f47,plain,(
  ( ! [X6,X4,X5] : (mult(X4,mult(mult(mult(ld(X4,X5),X4),X4),X6)) = mult(mult(X5,X4),mult(X4,X6))) )),
  inference(superposition,[],[f7,f1])).
fof(f6843,plain,(
  ( ! [X21,X22] : (mult(mult(ld(X21,X22),X21),X21) = ld(X21,mult(mult(X22,X21),X21))) )),
  inference(superposition,[],[f2,f236])).
fof(f236,plain,(
  ( ! [X2,X3] : (mult(mult(X3,X2),X2) = mult(X2,mult(mult(ld(X2,X3),X2),X2))) )),
  inference(superposition,[],[f75,f1])).
fof(f21008,plain,(
  ( ! [X74,X75] : (mult(mult(rd(X75,i(X74)),X74),X74) = mult(X74,mult(ld(X74,mult(mult(X75,X74),X74)),X74))) )),
  inference(forward_demodulation,[],[f20736,f6843])).
fof(f20736,plain,(
  ( ! [X74,X75] : (mult(X74,mult(mult(mult(ld(X74,X75),X74),X74),X74)) = mult(mult(rd(X75,i(X74)),X74),X74)) )),
  inference(superposition,[],[f75,f16671])).
fof(f16671,plain,(
  ( ! [X28,X27] : (mult(X27,mult(ld(X27,X28),X27)) = rd(X28,i(X27))) )),
  inference(backward_demodulation,[],[f16669,f8094])).
fof(f33036,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = mult(mult(rd(X31,i(X30)),X30),X30)) )),
  inference(forward_demodulation,[],[f33035,f19412])).
fof(f19412,plain,(
  ( ! [X14,X13] : (mult(rd(X14,i(X13)),X13) = rd(mult(X14,X13),i(X13))) )),
  inference(forward_demodulation,[],[f19411,f83])).
fof(f19411,plain,(
  ( ! [X14,X13] : (mult(rd(X14,i(X13)),X13) = rd(mult(X14,i(i(X13))),i(X13))) )),
  inference(forward_demodulation,[],[f19410,f16749])).
fof(f19410,plain,(
  ( ! [X14,X13] : (rd(mult(X14,i(i(X13))),i(X13)) = mult(mult(ld(X13,ld(i(X13),X14)),X13),X13)) )),
  inference(forward_demodulation,[],[f19306,f16980])).
fof(f16980,plain,(
  ( ! [X4,X5] : (mult(ld(X4,X5),X4) = ld(X4,rd(X5,i(X4)))) )),
  inference(superposition,[],[f3,f16416])).
fof(f19306,plain,(
  ( ! [X14,X13] : (rd(mult(X14,i(i(X13))),i(X13)) = mult(ld(X13,rd(ld(i(X13),X14),i(X13))),X13)) )),
  inference(superposition,[],[f16749,f16031])).
fof(f33035,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = mult(rd(mult(X31,X30),i(X30)),X30)) )),
  inference(forward_demodulation,[],[f33034,f19412])).
fof(f33034,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = rd(mult(mult(X31,X30),X30),i(X30))) )),
  inference(forward_demodulation,[],[f33033,f28483])).
fof(f28483,plain,(
  ( ! [X0,X1] : (rd(X1,i(X0)) = ld(mult(X0,X0),mult(X0,mult(X0,mult(X1,X0))))) )),
  inference(superposition,[],[f16678,f122])).
fof(f16678,plain,(
  ( ! [X21,X19,X20] : (ld(mult(X19,X20),mult(mult(X19,mult(X20,X21)),X20)) = rd(X21,i(X20))) )),
  inference(backward_demodulation,[],[f16669,f8257])).
fof(f8257,plain,(
  ( ! [X21,X19,X20] : (ld(mult(X19,X20),mult(mult(X19,mult(X20,X21)),X20)) = ld(X20,mult(mult(X20,X21),X20))) )),
  inference(backward_demodulation,[],[f8094,f1209])).
fof(f1209,plain,(
  ( ! [X21,X19,X20] : (mult(X20,mult(ld(X20,X21),X20)) = ld(mult(X19,X20),mult(mult(X19,mult(X20,X21)),X20))) )),
  inference(superposition,[],[f2,f88])).
fof(f33033,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = ld(mult(X30,X30),mult(X30,mult(X30,mult(mult(mult(X31,X30),X30),X30))))) )),
  inference(forward_demodulation,[],[f33032,f7])).
fof(f33032,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = ld(mult(X30,X30),mult(X30,mult(mult(mult(X30,X31),X30),mult(X30,X30))))) )),
  inference(forward_demodulation,[],[f33001,f107])).
fof(f33001,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = ld(mult(X30,X30),mult(X30,mult(mult(mult(X30,X31),mult(X30,X30)),X30)))) )),
  inference(backward_demodulation,[],[f33000,f32722])).
fof(f32722,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = ld(mult(X30,X30),mult(mult(X30,mult(mult(X30,X31),X30)),mult(X30,X30)))) )),
  inference(forward_demodulation,[],[f32721,f24296])).
fof(f24296,plain,(
  ( ! [X17,X16] : (mult(X16,mult(mult(X16,X17),X16)) = mult(mult(mult(X16,X16),X17),X16)) )),
  inference(forward_demodulation,[],[f24295,f2])).
fof(f24295,plain,(
  ( ! [X17,X16] : (mult(mult(mult(X16,X16),X17),X16) = mult(X16,mult(mult(X16,X17),ld(X16,mult(X16,X16))))) )),
  inference(forward_demodulation,[],[f23828,f10908])).
fof(f10908,plain,(
  ( ! [X103,X101,X102] : (mult(mult(rd(X101,X102),X103),X102) = mult(X101,ld(X102,mult(X103,X102)))) )),
  inference(superposition,[],[f8540,f3])).
fof(f23828,plain,(
  ( ! [X17,X16] : (mult(X16,mult(mult(rd(mult(X16,X17),X16),X16),X16)) = mult(mult(mult(X16,X16),X17),X16)) )),
  inference(superposition,[],[f75,f23510])).
fof(f32721,plain,(
  ( ! [X30,X31] : (mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30)) = ld(mult(X30,X30),mult(mult(mult(mult(X30,X30),X31),X30),mult(X30,X30)))) )),
  inference(forward_demodulation,[],[f32720,f107])).
fof(f32720,plain,(
  ( ! [X30,X31] : (ld(mult(X30,X30),mult(mult(mult(mult(X30,X30),X31),mult(X30,X30)),X30)) = mult(mult(mult(X31,i(X30)),mult(X30,X30)),mult(X30,X30))) )),
  inference(forward_demodulation,[],[f32719,f11252])).
fof(f32719,plain,(
  ( ! [X30,X31] : (ld(mult(X30,X30),mult(mult(mult(mult(X30,X30),X31),mult(X30,X30)),X30)) = mult(mult(mult(X31,mult(X30,X30)),i(X30)),mult(X30,X30))) )),
  inference(forward_demodulation,[],[f32388,f11252])).
fof(f32388,plain,(
  ( ! [X30,X31] : (ld(mult(X30,X30),mult(mult(mult(mult(X30,X30),X31),mult(X30,X30)),X30)) = mult(mult(mult(X31,mult(X30,X30)),mult(X30,X30)),i(X30))) )),
  inference(superposition,[],[f835,f552])).
fof(f835,plain,(
  ( ! [X2,X0,X1] : (mult(mult(mult(X1,X0),X0),ld(X0,X2)) = ld(X0,mult(mult(mult(X0,X1),X0),X2))) )),
  inference(superposition,[],[f2,f56])).
fof(f33000,plain,(
  ( ! [X6,X4,X5] : (mult(X4,mult(mult(mult(X4,X5),mult(X4,X4)),X6)) = mult(mult(X4,mult(mult(X4,X5),X4)),mult(X4,X6))) )),
  inference(forward_demodulation,[],[f32999,f24296])).
fof(f32999,plain,(
  ( ! [X6,X4,X5] : (mult(mult(mult(mult(X4,X4),X5),X4),mult(X4,X6)) = mult(X4,mult(mult(mult(X4,X5),mult(X4,X4)),X6))) )),
  inference(forward_demodulation,[],[f32998,f23181])).
fof(f23181,plain,(
  ( ! [X8,X7] : (mult(mult(X7,X8),mult(X7,X7)) = mult(mult(X7,mult(X8,X7)),X7)) )),
  inference(backward_demodulation,[],[f23180,f20806])).
fof(f20806,plain,(
  ( ! [X8,X7] : (mult(mult(X7,mult(X8,X7)),X7) = mult(X7,mult(rd(X8,i(X7)),X7))) )),
  inference(forward_demodulation,[],[f20687,f18245])).
fof(f18245,plain,(
  ( ! [X12,X13] : (rd(mult(mult(X12,X13),X12),i(X12)) = mult(mult(X12,mult(X13,X12)),X12)) )),
  inference(forward_demodulation,[],[f18117,f18242])).
fof(f18242,plain,(
  ( ! [X10,X9] : (mult(mult(X9,X10),X9) = mult(mult(mult(X9,mult(X10,X9)),X9),i(X9))) )),
  inference(forward_demodulation,[],[f18115,f16433])).
fof(f18115,plain,(
  ( ! [X10,X9] : (rd(ld(i(X9),X10),i(X9)) = mult(mult(mult(X9,mult(X10,X9)),X9),i(X9))) )),
  inference(superposition,[],[f16670,f12042])).
fof(f16670,plain,(
  ( ! [X0,X1] : (mult(mult(mult(X1,X0),X0),i(X0)) = rd(X1,i(X0))) )),
  inference(backward_demodulation,[],[f16669,f492])).
fof(f492,plain,(
  ( ! [X0,X1] : (mult(mult(mult(X1,X0),X0),i(X0)) = ld(X0,mult(mult(X0,X1),X0))) )),
  inference(superposition,[],[f2,f76])).
fof(f18117,plain,(
  ( ! [X12,X13] : (rd(mult(mult(X12,X13),X12),i(X12)) = mult(mult(mult(X12,mult(mult(X13,X12),X12)),X12),i(X12))) )),
  inference(superposition,[],[f16670,f75])).
fof(f20687,plain,(
  ( ! [X8,X7] : (rd(mult(mult(X7,X8),X7),i(X7)) = mult(X7,mult(rd(X8,i(X7)),X7))) )),
  inference(superposition,[],[f16671,f16669])).
fof(f23180,plain,(
  ( ! [X2,X3,X1] : (mult(mult(X1,X2),mult(X1,X3)) = mult(X1,mult(rd(X2,i(X1)),X3))) )),
  inference(forward_demodulation,[],[f23179,f16749])).
fof(f23179,plain,(
  ( ! [X2,X3,X1] : (mult(mult(X1,X2),mult(X1,X3)) = mult(X1,mult(mult(ld(X1,ld(i(X1),X2)),X1),X3))) )),
  inference(forward_demodulation,[],[f23010,f13393])).
fof(f13393,plain,(
  ( ! [X6,X5] : (mult(mult(X6,i(X5)),X5) = ld(X5,ld(i(X5),X6))) )),
  inference(backward_demodulation,[],[f13337,f11076])).
fof(f23010,plain,(
  ( ! [X2,X3,X1] : (mult(X1,mult(mult(mult(mult(X2,i(X1)),X1),X1),X3)) = mult(mult(X1,X2),mult(X1,X3))) )),
  inference(superposition,[],[f7,f17832])).
fof(f32998,plain,(
  ( ! [X6,X4,X5] : (mult(mult(mult(mult(X4,X4),X5),X4),mult(X4,X6)) = mult(X4,mult(mult(mult(X4,mult(X5,X4)),X4),X6))) )),
  inference(forward_demodulation,[],[f32889,f12042])).
fof(f32889,plain,(
  ( ! [X6,X4,X5] : (mult(mult(mult(mult(X4,X4),X5),X4),mult(X4,X6)) = mult(X4,mult(mult(mult(ld(i(X4),X5),X4),X4),X6))) )),
  inference(superposition,[],[f49,f13933])).
fof(f13933,plain,(
  ( ! [X10,X9] : (rd(mult(mult(X9,X9),X10),ld(i(X9),X10)) = X9) )),
  inference(superposition,[],[f34,f13315])).
fof(f49,plain,(
  ( ! [X12,X13,X11] : (mult(rd(X11,X12),mult(mult(mult(X12,rd(X11,X12)),rd(X11,X12)),X13)) = mult(mult(X11,rd(X11,X12)),mult(rd(X11,X12),X13))) )),
  inference(superposition,[],[f7,f3])).
fof(f33045,plain,(
  ( ! [X4,X5] : (mult(mult(X4,X4),mult(mult(mult(X5,i(X4)),mult(X4,X4)),mult(X4,X4))) = mult(X4,mult(X4,mult(mult(mult(X5,X4),X4),X4)))) )),
  inference(forward_demodulation,[],[f33044,f7])).
fof(f33044,plain,(
  ( ! [X4,X5] : (mult(mult(X4,X4),mult(mult(mult(X5,i(X4)),mult(X4,X4)),mult(X4,X4))) = mult(X4,mult(mult(mult(X4,X5),X4),mult(X4,X4)))) )),
  inference(forward_demodulation,[],[f33004,f107])).
fof(f33004,plain,(
  ( ! [X4,X5] : (mult(mult(X4,X4),mult(mult(mult(X5,i(X4)),mult(X4,X4)),mult(X4,X4))) = mult(X4,mult(mult(mult(X4,X5),mult(X4,X4)),X4))) )),
  inference(backward_demodulation,[],[f33000,f24305])).
fof(f24305,plain,(
  ( ! [X4,X5] : (mult(mult(X4,X4),mult(mult(mult(X5,i(X4)),mult(X4,X4)),mult(X4,X4))) = mult(mult(X4,mult(mult(X4,X5),X4)),mult(X4,X4))) )),
  inference(backward_demodulation,[],[f24296,f11317])).
fof(f11317,plain,(
  ( ! [X4,X5] : (mult(mult(mult(mult(X4,X4),X5),X4),mult(X4,X4)) = mult(mult(X4,X4),mult(mult(mult(X5,i(X4)),mult(X4,X4)),mult(X4,X4)))) )),
  inference(forward_demodulation,[],[f11265,f11252])).
fof(f11265,plain,(
  ( ! [X4,X5] : (mult(mult(mult(mult(X4,X4),X5),X4),mult(X4,X4)) = mult(mult(X4,X4),mult(mult(mult(X5,mult(X4,X4)),i(X4)),mult(X4,X4)))) )),
  inference(backward_demodulation,[],[f11252,f563])).
fof(f563,plain,(
  ( ! [X4,X5] : (mult(mult(X4,X4),mult(mult(mult(X5,mult(X4,X4)),mult(X4,X4)),i(X4))) = mult(mult(mult(mult(X4,X4),X5),X4),mult(X4,X4))) )),
  inference(forward_demodulation,[],[f555,f107])).
fof(f555,plain,(
  ( ! [X4,X5] : (mult(mult(X4,X4),mult(mult(mult(X5,mult(X4,X4)),mult(X4,X4)),i(X4))) = mult(mult(mult(mult(X4,X4),X5),mult(X4,X4)),X4)) )),
  inference(superposition,[],[f7,f457])).
fof(f174803,plain,(
  ( ! [X6,X5] : (rd(ld(X5,ld(i(mult(X5,X5)),mult(X6,X6))),mult(X6,X6)) = X5) )),
  inference(forward_demodulation,[],[f174712,f100197])).
fof(f100197,plain,(
  ( ! [X14,X13] : (i(rd(i(mult(X14,X14)),X13)) = ld(i(X13),mult(X14,X14))) )),
  inference(forward_demodulation,[],[f100152,f13225])).
fof(f100152,plain,(
  ( ! [X14,X13] : (rd(mult(X13,mult(mult(X14,X14),X13)),X13) = i(rd(i(mult(X14,X14)),X13))) )),
  inference(superposition,[],[f67261,f33744])).
fof(f33744,plain,(
  ( ! [X24,X25] : (i(mult(X24,X24)) = rd(X25,mult(mult(X24,X24),X25))) )),
  inference(superposition,[],[f34,f33052])).
fof(f67261,plain,(
  ( ! [X12,X13] : (rd(mult(X12,X13),X12) = i(rd(rd(X12,X13),X12))) )),
  inference(superposition,[],[f83,f25834])).
fof(f25834,plain,(
  ( ! [X8,X9] : (i(rd(mult(X8,X9),X8)) = rd(rd(X8,X9),X8)) )),
  inference(superposition,[],[f4,f23521])).
fof(f174712,plain,(
  ( ! [X6,X5] : (rd(ld(X5,i(rd(i(mult(X6,X6)),mult(X5,X5)))),mult(X6,X6)) = X5) )),
  inference(superposition,[],[f174525,f33805])).
fof(f33805,plain,(
  ( ! [X21,X22] : (mult(X22,i(mult(X21,X21))) = rd(X22,mult(X21,X21))) )),
  inference(forward_demodulation,[],[f33804,f83])).
fof(f33804,plain,(
  ( ! [X21,X22] : (rd(X22,i(i(mult(X21,X21)))) = mult(X22,i(mult(X21,X21)))) )),
  inference(forward_demodulation,[],[f33803,f16429])).
fof(f33803,plain,(
  ( ! [X21,X22] : (rd(X22,i(i(mult(X21,X21)))) = mult(mult(X21,X21),rd(ld(mult(X21,X21),X22),mult(X21,X21)))) )),
  inference(forward_demodulation,[],[f33727,f16032])).
fof(f16032,plain,(
  ( ! [X4,X5] : (mult(mult(i(X5),X4),i(X5)) = rd(ld(X5,X4),X5)) )),
  inference(backward_demodulation,[],[f16031,f14544])).
fof(f33727,plain,(
  ( ! [X21,X22] : (rd(X22,i(i(mult(X21,X21)))) = mult(mult(X21,X21),mult(mult(i(mult(X21,X21)),X22),i(mult(X21,X21))))) )),
  inference(superposition,[],[f33052,f16669])).
fof(f174525,plain,(
  ( ! [X52,X51] : (mult(ld(X52,i(rd(X51,mult(X52,X52)))),X51) = X52) )),
  inference(forward_demodulation,[],[f174472,f38141])).
fof(f38141,plain,(
  ( ! [X6,X5] : (rd(mult(X6,i(X5)),X5) = rd(X6,mult(X5,X5))) )),
  inference(backward_demodulation,[],[f38126,f20852])).
fof(f20852,plain,(
  ( ! [X6,X5] : (rd(mult(X6,i(X5)),X5) = mult(rd(X6,X5),i(X5))) )),
  inference(forward_demodulation,[],[f20843,f15916])).
fof(f15916,plain,(
  ( ! [X2,X1] : (mult(X2,i(X1)) = ld(i(X1),ld(X1,rd(X2,X1)))) )),
  inference(backward_demodulation,[],[f15841,f12026])).
fof(f12026,plain,(
  ( ! [X2,X1] : (mult(X2,i(X1)) = ld(i(X1),mult(ld(X1,X2),i(X1)))) )),
  inference(superposition,[],[f11505,f83])).
fof(f20843,plain,(
  ( ! [X6,X5] : (rd(mult(X6,i(X5)),X5) = ld(i(X5),ld(X5,rd(rd(X6,X5),X5)))) )),
  inference(backward_demodulation,[],[f20836,f16637])).
fof(f16637,plain,(
  ( ! [X6,X5] : (rd(mult(X6,i(X5)),X5) = ld(i(X5),rd(rd(ld(X5,X6),X5),X5))) )),
  inference(superposition,[],[f16245,f16031])).
fof(f20836,plain,(
  ( ! [X35,X34] : (ld(X34,rd(rd(X35,X34),X34)) = rd(rd(ld(X34,X35),X34),X34)) )),
  inference(forward_demodulation,[],[f20835,f83])).
fof(f20835,plain,(
  ( ! [X35,X34] : (rd(rd(ld(X34,X35),X34),i(i(X34))) = ld(X34,rd(rd(X35,X34),X34))) )),
  inference(forward_demodulation,[],[f20703,f15917])).
fof(f20703,plain,(
  ( ! [X35,X34] : (rd(rd(ld(X34,X35),X34),i(i(X34))) = mult(i(X34),mult(rd(X35,X34),i(X34)))) )),
  inference(superposition,[],[f16671,f16245])).
fof(f38126,plain,(
  ( ! [X2,X1] : (rd(X2,mult(X1,X1)) = mult(rd(X2,X1),i(X1))) )),
  inference(forward_demodulation,[],[f38125,f33805])).
fof(f38125,plain,(
  ( ! [X2,X1] : (mult(X2,i(mult(X1,X1))) = mult(rd(X2,X1),i(X1))) )),
  inference(forward_demodulation,[],[f38000,f3662])).
fof(f38000,plain,(
  ( ! [X2,X1] : (mult(rd(X2,X1),i(X1)) = mult(X2,mult(i(X1),i(X1)))) )),
  inference(superposition,[],[f37332,f83])).
fof(f37332,plain,(
  ( ! [X14,X13] : (mult(X14,mult(X13,X13)) = mult(rd(X14,i(X13)),X13)) )),
  inference(backward_demodulation,[],[f37330,f19412])).
fof(f37330,plain,(
  ( ! [X39,X38] : (mult(X39,mult(X38,X38)) = rd(mult(X39,X38),i(X38))) )),
  inference(forward_demodulation,[],[f37329,f2])).
fof(f37329,plain,(
  ( ! [X39,X38] : (rd(mult(X39,X38),i(X38)) = ld(X38,mult(X38,mult(X39,mult(X38,X38))))) )),
  inference(forward_demodulation,[],[f37328,f83])).
fof(f37328,plain,(
  ( ! [X39,X38] : (rd(mult(X39,X38),i(X38)) = ld(i(i(X38)),mult(X38,mult(X39,mult(X38,X38))))) )),
  inference(forward_demodulation,[],[f37327,f34056])).
fof(f34056,plain,(
  ( ! [X8,X7] : (mult(mult(X7,mult(X8,X7)),X7) = mult(X7,mult(X8,mult(X7,X7)))) )),
  inference(backward_demodulation,[],[f33958,f23181])).
fof(f33958,plain,(
  ( ! [X173,X172] : (mult(mult(X172,X173),mult(X172,X172)) = mult(X172,mult(X173,mult(X172,X172)))) )),
  inference(backward_demodulation,[],[f33946,f24188])).
fof(f24188,plain,(
  ( ! [X173,X172] : (mult(X172,rd(X173,i(mult(X172,X172)))) = mult(mult(X172,X173),mult(X172,X172))) )),
  inference(forward_demodulation,[],[f24187,f23181])).
fof(f24187,plain,(
  ( ! [X173,X172] : (mult(X172,rd(X173,i(mult(X172,X172)))) = mult(mult(X172,mult(X173,X172)),X172)) )),
  inference(forward_demodulation,[],[f24186,f6799])).
fof(f6799,plain,(
  ( ! [X6,X7] : (mult(mult(X6,X7),X6) = ld(mult(X6,X6),mult(X6,mult(X6,mult(X6,mult(X7,X6)))))) )),
  inference(forward_demodulation,[],[f6798,f372])).
fof(f6798,plain,(
  ( ! [X6,X7] : (mult(X6,mult(X6,mult(ld(X6,X7),X6))) = ld(mult(X6,X6),mult(X6,mult(X6,mult(X6,mult(X7,X6)))))) )),
  inference(forward_demodulation,[],[f6773,f122])).
fof(f6773,plain,(
  ( ! [X6,X7] : (mult(X6,mult(X6,mult(ld(X6,X7),X6))) = ld(mult(X6,X6),mult(X6,mult(mult(X6,mult(X6,X7)),X6)))) )),
  inference(superposition,[],[f184,f88])).
fof(f24186,plain,(
  ( ! [X173,X172] : (mult(X172,rd(X173,i(mult(X172,X172)))) = ld(mult(X172,X172),mult(X172,mult(X172,mult(X172,mult(mult(X173,X172),X172)))))) )),
  inference(forward_demodulation,[],[f24038,f75])).
fof(f24038,plain,(
  ( ! [X173,X172] : (mult(X172,rd(X173,i(mult(X172,X172)))) = ld(mult(X172,X172),mult(X172,mult(X172,mult(mult(mult(X172,X173),X172),X172))))) )),
  inference(backward_demodulation,[],[f24017,f17950])).
fof(f17950,plain,(
  ( ! [X173,X172] : (mult(X172,rd(X173,i(mult(X172,X172)))) = ld(mult(X172,X172),mult(X172,mult(mult(mult(X172,X172),X173),mult(X172,X172))))) )),
  inference(superposition,[],[f184,f17329])).
fof(f17329,plain,(
  ( ! [X121,X120] : (mult(mult(X120,X121),X120) = mult(X120,rd(X121,i(X120)))) )),
  inference(forward_demodulation,[],[f17104,f6])).
fof(f17104,plain,(
  ( ! [X121,X120] : (mult(mult(unit,mult(X120,X121)),X120) = mult(X120,rd(X121,i(X120)))) )),
  inference(superposition,[],[f16672,f6])).
fof(f24017,plain,(
  ( ! [X4,X2,X3] : (mult(mult(mult(X2,X2),X3),mult(X2,X4)) = mult(X2,mult(mult(mult(X2,X3),X2),X4))) )),
  inference(forward_demodulation,[],[f24016,f2])).
fof(f24016,plain,(
  ( ! [X4,X2,X3] : (mult(mult(mult(X2,X2),X3),mult(X2,X4)) = mult(X2,mult(mult(mult(X2,X3),ld(X2,mult(X2,X2))),X4))) )),
  inference(forward_demodulation,[],[f23824,f10908])).
fof(f23824,plain,(
  ( ! [X4,X2,X3] : (mult(X2,mult(mult(mult(rd(mult(X2,X3),X2),X2),X2),X4)) = mult(mult(mult(X2,X2),X3),mult(X2,X4))) )),
  inference(superposition,[],[f7,f23510])).
fof(f33946,plain,(
  ( ! [X8,X9] : (mult(X9,mult(X8,X8)) = rd(X9,i(mult(X8,X8)))) )),
  inference(forward_demodulation,[],[f33735,f2])).
fof(f33735,plain,(
  ( ! [X8,X9] : (rd(X9,i(mult(X8,X8))) = mult(ld(mult(X8,X8),mult(mult(X8,X8),X9)),mult(X8,X8))) )),
  inference(superposition,[],[f16749,f33052])).
fof(f37327,plain,(
  ( ! [X39,X38] : (rd(mult(X39,X38),i(X38)) = ld(i(i(X38)),mult(mult(X38,mult(X39,X38)),X38))) )),
  inference(forward_demodulation,[],[f37326,f17392])).
fof(f17392,plain,(
  ( ! [X30,X31] : (mult(X30,mult(X31,X30)) = rd(mult(X30,X31),i(X30))) )),
  inference(forward_demodulation,[],[f17391,f16758])).
fof(f16758,plain,(
  ( ! [X6,X5] : (mult(X5,mult(X6,X5)) = ld(i(X5),rd(X6,i(X5)))) )),
  inference(backward_demodulation,[],[f16669,f13956])).
fof(f13956,plain,(
  ( ! [X6,X5] : (mult(X5,mult(X6,X5)) = ld(i(X5),ld(X5,mult(mult(X5,X6),X5)))) )),
  inference(forward_demodulation,[],[f13955,f2])).
fof(f13955,plain,(
  ( ! [X6,X5] : (ld(i(X5),ld(X5,mult(mult(X5,X6),X5))) = ld(X5,mult(X5,mult(X5,mult(X6,X5))))) )),
  inference(forward_demodulation,[],[f13915,f122])).
fof(f13915,plain,(
  ( ! [X6,X5] : (ld(i(X5),ld(X5,mult(mult(X5,X6),X5))) = ld(X5,mult(mult(X5,mult(X5,X6)),X5))) )),
  inference(superposition,[],[f13315,f8249])).
fof(f17391,plain,(
  ( ! [X30,X31] : (ld(i(X30),rd(X31,i(X30))) = rd(mult(X30,X31),i(X30))) )),
  inference(forward_demodulation,[],[f17140,f16669])).
fof(f17140,plain,(
  ( ! [X30,X31] : (ld(i(X30),rd(X31,i(X30))) = ld(X30,mult(mult(X30,mult(X30,X31)),X30))) )),
  inference(superposition,[],[f13315,f16672])).
fof(f37326,plain,(
  ( ! [X39,X38] : (rd(mult(X39,X38),i(X38)) = ld(i(i(X38)),mult(rd(mult(X38,X39),i(X38)),X38))) )),
  inference(forward_demodulation,[],[f37216,f19412])).
fof(f37216,plain,(
  ( ! [X39,X38] : (rd(mult(X39,X38),i(X38)) = ld(i(i(X38)),rd(mult(mult(X38,X39),X38),i(X38)))) )),
  inference(superposition,[],[f16245,f14354])).
fof(f174472,plain,(
  ( ! [X52,X51] : (mult(ld(X52,i(rd(mult(X51,i(X52)),X52))),X51) = X52) )),
  inference(superposition,[],[f174236,f4])).
fof(f174236,plain,(
  ( ! [X0,X1] : (mult(ld(X0,i(rd(X1,X0))),rd(X1,i(X0))) = X0) )),
  inference(backward_demodulation,[],[f174235,f16673])).
fof(f174235,plain,(
  ( ! [X0,X1] : (mult(i(mult(X0,X1)),X0) = ld(X0,i(rd(X1,X0)))) )),
  inference(forward_demodulation,[],[f174210,f173867])).
fof(f173867,plain,(
  ( ! [X59,X60] : (ld(ld(X59,X60),i(X59)) = ld(X59,i(rd(X60,X59)))) )),
  inference(forward_demodulation,[],[f173866,f68371])).
fof(f173866,plain,(
  ( ! [X59,X60] : (ld(ld(X59,X60),i(X59)) = mult(ld(X60,X59),i(X59))) )),
  inference(forward_demodulation,[],[f173789,f172952])).
fof(f172952,plain,(
  ( ! [X101,X100] : (ld(X101,X100) = rd(i(X100),ld(X100,rd(X101,X100)))) )),
  inference(forward_demodulation,[],[f172951,f2])).
fof(f172951,plain,(
  ( ! [X101,X100] : (rd(i(X100),ld(X100,rd(X101,X100))) = ld(X100,mult(X100,ld(X101,X100)))) )),
  inference(forward_demodulation,[],[f172949,f28087])).
fof(f172949,plain,(
  ( ! [X101,X100] : (rd(i(X100),ld(X100,rd(X101,X100))) = ld(X100,rd(X100,ld(X100,X101)))) )),
  inference(backward_demodulation,[],[f172932,f137411])).
fof(f137411,plain,(
  ( ! [X101,X100] : (mult(i(ld(X100,ld(X100,X101))),i(X100)) = rd(i(X100),ld(X100,rd(X101,X100)))) )),
  inference(forward_demodulation,[],[f137246,f15917])).
fof(f137246,plain,(
  ( ! [X101,X100] : (rd(i(X100),mult(i(X100),mult(X101,i(X100)))) = mult(i(ld(X100,ld(X100,X101))),i(X100))) )),
  inference(superposition,[],[f615,f13498])).
fof(f172932,plain,(
  ( ! [X532,X533] : (mult(i(ld(X532,X533)),i(X532)) = ld(X532,rd(X532,X533))) )),
  inference(backward_demodulation,[],[f172871,f172786])).
fof(f172786,plain,(
  ( ! [X532,X533] : (mult(i(ld(X532,X533)),i(X532)) = ld(X532,rd(ld(rd(X533,X532),X532),X532))) )),
  inference(superposition,[],[f27898,f172431])).
fof(f172431,plain,(
  ( ! [X261,X262] : (ld(rd(X261,X262),X262) = mult(X262,i(ld(X262,X261)))) )),
  inference(superposition,[],[f172363,f3])).
fof(f172363,plain,(
  ( ! [X62,X63] : (ld(X63,X62) = mult(X62,i(ld(X62,mult(X63,X62))))) )),
  inference(forward_demodulation,[],[f172287,f2])).
fof(f172287,plain,(
  ( ! [X62,X63] : (ld(ld(i(X62),mult(i(X62),X63)),X62) = mult(X62,i(ld(X62,mult(X63,X62))))) )),
  inference(superposition,[],[f22238,f11062])).
fof(f27898,plain,(
  ( ! [X21,X20] : (mult(X21,i(X20)) = ld(X20,rd(mult(X20,X21),X20))) )),
  inference(forward_demodulation,[],[f27897,f3669])).
fof(f3669,plain,(
  ( ! [X10] : (i(X10) = rd(i(mult(X10,X10)),i(X10))) )),
  inference(backward_demodulation,[],[f3661,f1662])).
fof(f1662,plain,(
  ( ! [X10] : (i(X10) = rd(ld(X10,i(X10)),i(X10))) )),
  inference(backward_demodulation,[],[f1655,f911])).
fof(f911,plain,(
  ( ! [X10] : (i(X10) = rd(rd(i(X10),X10),i(X10))) )),
  inference(superposition,[],[f4,f797])).
fof(f27897,plain,(
  ( ! [X21,X20] : (mult(X21,rd(i(mult(X20,X20)),i(X20))) = ld(X20,rd(mult(X20,X21),X20))) )),
  inference(forward_demodulation,[],[f27896,f16635])).
fof(f16635,plain,(
  ( ! [X2,X1] : (ld(i(X1),rd(X2,X1)) = rd(mult(X1,X2),X1)) )),
  inference(superposition,[],[f16245,f2])).
fof(f27896,plain,(
  ( ! [X21,X20] : (mult(X21,rd(i(mult(X20,X20)),i(X20))) = ld(X20,ld(i(X20),rd(X21,X20)))) )),
  inference(forward_demodulation,[],[f27681,f13393])).
fof(f27681,plain,(
  ( ! [X21,X20] : (mult(X21,rd(i(mult(X20,X20)),i(X20))) = mult(mult(rd(X21,X20),i(X20)),X20)) )),
  inference(superposition,[],[f16677,f4694])).
fof(f16677,plain,(
  ( ! [X37,X35,X36] : (mult(mult(rd(X35,X36),mult(X36,X37)),X36) = mult(X35,rd(X37,i(X36)))) )),
  inference(backward_demodulation,[],[f16669,f8256])).
fof(f8256,plain,(
  ( ! [X37,X35,X36] : (mult(mult(rd(X35,X36),mult(X36,X37)),X36) = mult(X35,ld(X36,mult(mult(X36,X37),X36)))) )),
  inference(backward_demodulation,[],[f8094,f1184])).
fof(f1184,plain,(
  ( ! [X37,X35,X36] : (mult(mult(rd(X35,X36),mult(X36,X37)),X36) = mult(X35,mult(X36,mult(ld(X36,X37),X36)))) )),
  inference(superposition,[],[f88,f3])).
fof(f172871,plain,(
  ( ! [X4,X5] : (rd(X5,X4) = rd(ld(rd(X4,X5),X5),X5)) )),
  inference(superposition,[],[f34,f172835])).
fof(f172835,plain,(
  ( ! [X629,X628] : (ld(rd(ld(rd(X629,X628),X628),X628),X628) = X629) )),
  inference(forward_demodulation,[],[f172834,f1])).
fof(f172834,plain,(
  ( ! [X629,X628] : (ld(rd(ld(rd(X629,X628),X628),X628),X628) = mult(X628,ld(X628,X629))) )),
  inference(forward_demodulation,[],[f172793,f83])).
fof(f172793,plain,(
  ( ! [X629,X628] : (mult(X628,i(i(ld(X628,X629)))) = ld(rd(ld(rd(X629,X628),X628),X628),X628)) )),
  inference(superposition,[],[f172373,f172431])).
fof(f172373,plain,(
  ( ! [X253,X254] : (mult(X254,i(X253)) = ld(rd(mult(X254,X253),X254),X254)) )),
  inference(forward_demodulation,[],[f172291,f16635])).
fof(f172291,plain,(
  ( ! [X253,X254] : (ld(ld(i(X254),rd(X253,X254)),X254) = mult(X254,i(X253))) )),
  inference(superposition,[],[f22238,f3])).
fof(f173789,plain,(
  ( ! [X59,X60] : (ld(ld(X59,X60),i(X59)) = mult(rd(i(X59),ld(X59,rd(X60,X59))),i(X59))) )),
  inference(superposition,[],[f173540,f15841])).
fof(f173540,plain,(
  ( ! [X21,X22] : (ld(X21,X22) = mult(rd(X22,mult(X21,X22)),X22)) )),
  inference(superposition,[],[f3,f173002])).
fof(f173002,plain,(
  ( ! [X4,X5] : (rd(ld(X4,X5),X5) = rd(X5,mult(X4,X5))) )),
  inference(superposition,[],[f34,f172851])).
fof(f172851,plain,(
  ( ! [X68,X67] : (mult(X67,X68) = ld(rd(ld(X67,X68),X68),X68)) )),
  inference(superposition,[],[f172835,f4])).
fof(f174210,plain,(
  ( ! [X0,X1] : (mult(i(mult(X0,X1)),X0) = ld(ld(X0,X1),i(X0))) )),
  inference(superposition,[],[f174124,f1])).
fof(f174124,plain,(
  ( ! [X4,X3] : (mult(i(mult(X3,mult(X3,X4))),X3) = ld(X4,i(X3))) )),
  inference(backward_demodulation,[],[f174123,f615])).
fof(f174123,plain,(
  ( ! [X2,X1] : (rd(X1,mult(X1,mult(X2,X1))) = ld(X2,i(X1))) )),
  inference(forward_demodulation,[],[f174088,f16758])).
fof(f174088,plain,(
  ( ! [X2,X1] : (ld(X2,i(X1)) = rd(X1,ld(i(X1),rd(X2,i(X1))))) )),
  inference(superposition,[],[f172952,f83])).
fof(f48758,plain,(
  ( ! [X208,X209] : (i(X208) = ld(rd(X209,X208),rd(X209,mult(X208,X208)))) )),
  inference(forward_demodulation,[],[f48757,f2])).
fof(f48757,plain,(
  ( ! [X208,X209] : (i(X208) = ld(rd(ld(X208,mult(X208,X209)),X208),rd(X209,mult(X208,X208)))) )),
  inference(forward_demodulation,[],[f48756,f16032])).
fof(f48756,plain,(
  ( ! [X208,X209] : (i(X208) = ld(mult(mult(i(X208),mult(X208,X209)),i(X208)),rd(X209,mult(X208,X208)))) )),
  inference(forward_demodulation,[],[f48755,f38126])).
fof(f48755,plain,(
  ( ! [X208,X209] : (i(X208) = ld(mult(mult(i(X208),mult(X208,X209)),i(X208)),mult(rd(X209,X208),i(X208)))) )),
  inference(forward_demodulation,[],[f48754,f27898])).
fof(f48754,plain,(
  ( ! [X208,X209] : (i(X208) = ld(mult(mult(i(X208),mult(X208,X209)),i(X208)),ld(X208,rd(mult(X208,rd(X209,X208)),X208)))) )),
  inference(forward_demodulation,[],[f48309,f15917])).
fof(f48309,plain,(
  ( ! [X208,X209] : (i(X208) = ld(mult(mult(i(X208),mult(X208,X209)),i(X208)),mult(i(X208),mult(mult(X208,rd(X209,X208)),i(X208))))) )),
  inference(superposition,[],[f250,f17223])).
fof(f250,plain,(
  ( ! [X0,X1] : (ld(mult(mult(X0,X1),X0),mult(X0,mult(mult(X1,X0),X0))) = X0) )),
  inference(superposition,[],[f2,f75])).
fof(f179657,plain,(
  ( ! [X2,X1] : (ld(ld(i(X1),X2),X1) = mult(ld(X2,X1),i(X1))) )),
  inference(backward_demodulation,[],[f179649,f174175])).
fof(f174175,plain,(
  ( ! [X2,X1] : (mult(i(mult(i(X1),X2)),i(X1)) = ld(ld(i(X1),X2),X1)) )),
  inference(backward_demodulation,[],[f174171,f161256])).
fof(f161256,plain,(
  ( ! [X2,X1] : (mult(i(mult(i(X1),X2)),i(X1)) = rd(i(X1),rd(X2,X1))) )),
  inference(superposition,[],[f18810,f83])).
fof(f18810,plain,(
  ( ! [X12,X13] : (mult(i(mult(X12,X13)),X12) = rd(X12,rd(X13,i(X12)))) )),
  inference(superposition,[],[f4,f16673])).
fof(f174171,plain,(
  ( ! [X52,X53] : (ld(ld(i(X52),X53),X52) = rd(i(X52),rd(X53,X52))) )),
  inference(forward_demodulation,[],[f174098,f2])).
fof(f174098,plain,(
  ( ! [X52,X53] : (ld(ld(i(X52),X53),X52) = rd(i(X52),ld(X52,mult(X52,rd(X53,X52))))) )),
  inference(superposition,[],[f172952,f15074])).
fof(f179649,plain,(
  ( ! [X2,X3] : (ld(X3,X2) = i(mult(i(X2),X3))) )),
  inference(forward_demodulation,[],[f179566,f179265])).
fof(f179265,plain,(
  ( ! [X259,X258] : (ld(X258,X259) = mult(ld(X259,rd(X259,X258)),X259)) )),
  inference(forward_demodulation,[],[f179264,f2])).
fof(f179264,plain,(
  ( ! [X259,X258] : (mult(ld(X259,rd(X259,X258)),X259) = ld(ld(X259,mult(X259,X258)),X259)) )),
  inference(forward_demodulation,[],[f179126,f175185])).
fof(f175185,plain,(
  ( ! [X57,X56] : (ld(ld(X56,X57),X56) = mult(i(X57),mult(X56,X56))) )),
  inference(superposition,[],[f2,f174550])).
fof(f174550,plain,(
  ( ! [X2,X1] : (mult(ld(X2,X1),mult(i(X1),mult(X2,X2))) = X2) )),
  inference(superposition,[],[f174499,f83])).
fof(f174499,plain,(
  ( ! [X76,X75] : (mult(ld(X76,i(X75)),mult(X75,mult(X76,X76))) = X76) )),
  inference(forward_demodulation,[],[f174454,f37330])).
fof(f174454,plain,(
  ( ! [X76,X75] : (mult(ld(X76,i(X75)),rd(mult(X75,X76),i(X76))) = X76) )),
  inference(superposition,[],[f174236,f4])).
fof(f179126,plain,(
  ( ! [X259,X258] : (mult(ld(X259,rd(X259,X258)),X259) = mult(i(mult(X259,X258)),mult(X259,X259))) )),
  inference(backward_demodulation,[],[f179118,f173654])).
fof(f173654,plain,(
  ( ! [X259,X258] : (mult(ld(X258,i(X259)),mult(X259,X259)) = mult(ld(X259,rd(X259,X258)),X259)) )),
  inference(forward_demodulation,[],[f173568,f172948])).
fof(f172948,plain,(
  ( ! [X428,X429] : (rd(i(X428),mult(X429,i(X428))) = ld(X428,rd(X428,X429))) )),
  inference(backward_demodulation,[],[f172932,f131498])).
fof(f131498,plain,(
  ( ! [X428,X429] : (mult(i(ld(X428,X429)),i(X428)) = rd(i(X428),mult(X429,i(X428)))) )),
  inference(forward_demodulation,[],[f131497,f27898])).
fof(f131497,plain,(
  ( ! [X428,X429] : (mult(i(ld(X428,X429)),i(X428)) = rd(i(X428),ld(X428,rd(mult(X428,X429),X428)))) )),
  inference(forward_demodulation,[],[f131496,f16635])).
fof(f131496,plain,(
  ( ! [X428,X429] : (mult(i(ld(X428,X429)),i(X428)) = rd(i(X428),ld(X428,ld(i(X428),rd(X429,X428))))) )),
  inference(forward_demodulation,[],[f131383,f14132])).
fof(f131383,plain,(
  ( ! [X428,X429] : (mult(i(ld(X428,X429)),i(X428)) = rd(i(X428),ld(i(X428),ld(X428,rd(X429,X428))))) )),
  inference(superposition,[],[f25185,f15841])).
fof(f25185,plain,(
  ( ! [X23,X22] : (mult(i(X22),X23) = rd(X23,ld(X23,mult(X22,X23)))) )),
  inference(superposition,[],[f4,f20089])).
fof(f173568,plain,(
  ( ! [X259,X258] : (mult(ld(X258,i(X259)),mult(X259,X259)) = mult(rd(i(X259),mult(X258,i(X259))),X259)) )),
  inference(superposition,[],[f37332,f173002])).
fof(f179118,plain,(
  ( ! [X8,X9] : (ld(X8,i(X9)) = i(mult(X9,X8))) )),
  inference(backward_demodulation,[],[f179074,f98422])).
fof(f98422,plain,(
  ( ! [X8,X9] : (ld(X8,i(X9)) = mult(ld(mult(X9,X8),X8),i(X8))) )),
  inference(superposition,[],[f14141,f63279])).
fof(f63279,plain,(
  ( ! [X52,X51] : (ld(X52,mult(i(X51),X52)) = ld(mult(X51,X52),X52)) )),
  inference(superposition,[],[f2,f25091])).
fof(f25091,plain,(
  ( ! [X2,X1] : (mult(mult(X1,X2),ld(X2,mult(i(X1),X2))) = X2) )),
  inference(superposition,[],[f20089,f83])).
fof(f179566,plain,(
  ( ! [X2,X3] : (mult(ld(X2,rd(X2,X3)),X2) = i(mult(i(X2),X3))) )),
  inference(backward_demodulation,[],[f179562,f179317])).
fof(f179317,plain,(
  ( ! [X2,X3] : (mult(i(ld(X2,ld(i(X2),X3))),X2) = i(mult(i(X2),X3))) )),
  inference(backward_demodulation,[],[f179301,f175859])).
fof(f175859,plain,(
  ( ! [X2,X3] : (mult(i(ld(X2,ld(i(X2),X3))),X2) = rd(i(X3),i(X2))) )),
  inference(backward_demodulation,[],[f175857,f160696])).
fof(f160696,plain,(
  ( ! [X2,X3] : (mult(i(ld(X2,ld(i(X2),X3))),X2) = ld(mult(X2,X3),mult(X2,X2))) )),
  inference(backward_demodulation,[],[f160695,f25926])).
fof(f25926,plain,(
  ( ! [X2,X3] : (rd(X2,mult(ld(X2,X3),X2)) = mult(i(ld(X2,ld(i(X2),X3))),X2)) )),
  inference(forward_demodulation,[],[f25770,f17034])).
fof(f25770,plain,(
  ( ! [X2,X3] : (rd(X2,mult(ld(X2,X3),X2)) = mult(i(rd(rd(X3,i(X2)),X2)),X2)) )),
  inference(superposition,[],[f23521,f16671])).
fof(f160695,plain,(
  ( ! [X123,X122] : (rd(X122,mult(ld(X122,X123),X122)) = ld(mult(X122,X123),mult(X122,X122))) )),
  inference(forward_demodulation,[],[f160694,f78822])).
fof(f78822,plain,(
  ( ! [X8,X7] : (ld(i(mult(i(X8),X7)),X7) = ld(X8,mult(X7,X7))) )),
  inference(forward_demodulation,[],[f78754,f56377])).
fof(f78754,plain,(
  ( ! [X8,X7] : (mult(X7,ld(ld(X7,mult(X8,X7)),X7)) = ld(i(mult(i(X8),X7)),X7)) )),
  inference(superposition,[],[f43606,f25185])).
fof(f43606,plain,(
  ( ! [X2,X3] : (mult(X2,ld(X3,X2)) = ld(i(rd(X2,X3)),X2)) )),
  inference(forward_demodulation,[],[f43372,f2])).
fof(f43372,plain,(
  ( ! [X2,X3] : (ld(i(rd(X2,X3)),X2) = ld(rd(X2,X3),mult(rd(X2,X3),mult(X2,ld(X3,X2))))) )),
  inference(superposition,[],[f13315,f11058])).
fof(f11058,plain,(
  ( ! [X12,X13] : (mult(mult(rd(X12,X13),rd(X12,X13)),X12) = mult(rd(X12,X13),mult(X12,ld(X13,X12)))) )),
  inference(forward_demodulation,[],[f11051,f3])).
fof(f11051,plain,(
  ( ! [X12,X13] : (mult(mult(rd(X12,X13),rd(X12,X13)),X12) = mult(rd(X12,X13),mult(X12,ld(X13,mult(rd(X12,X13),X13))))) )),
  inference(backward_demodulation,[],[f10908,f176])).
fof(f176,plain,(
  ( ! [X12,X13] : (mult(rd(X12,X13),mult(mult(rd(X12,X13),rd(X12,X13)),X13)) = mult(mult(rd(X12,X13),rd(X12,X13)),X12)) )),
  inference(superposition,[],[f63,f3])).
fof(f160694,plain,(
  ( ! [X123,X122] : (ld(i(mult(i(mult(X122,X123)),X122)),X122) = rd(X122,mult(ld(X122,X123),X122))) )),
  inference(forward_demodulation,[],[f160618,f30033])).
fof(f30033,plain,(
  ( ! [X10,X9] : (rd(X9,mult(ld(X9,X10),X9)) = mult(X9,ld(rd(X10,i(X9)),X9))) )),
  inference(superposition,[],[f28087,f16980])).
fof(f160618,plain,(
  ( ! [X123,X122] : (ld(i(mult(i(mult(X122,X123)),X122)),X122) = mult(X122,ld(rd(X123,i(X122)),X122))) )),
  inference(superposition,[],[f78751,f18809])).
fof(f18809,plain,(
  ( ! [X10,X11] : (rd(X11,i(X10)) = ld(mult(i(mult(X10,X11)),X10),X10)) )),
  inference(superposition,[],[f2,f16673])).
fof(f78751,plain,(
  ( ! [X2,X1] : (ld(i(X2),X1) = mult(X1,ld(ld(X2,X1),X1))) )),
  inference(superposition,[],[f43606,f34])).
fof(f175857,plain,(
  ( ! [X37,X38] : (rd(i(X38),i(X37)) = ld(mult(X37,X38),mult(X37,X37))) )),
  inference(forward_demodulation,[],[f175748,f3])).
fof(f175748,plain,(
  ( ! [X37,X38] : (rd(i(X38),i(X37)) = ld(mult(rd(mult(X37,X38),X37),X37),mult(X37,X37))) )),
  inference(superposition,[],[f16678,f172666])).
fof(f172666,plain,(
  ( ! [X0,X1] : (mult(rd(mult(X0,X1),X0),mult(X0,i(X1))) = X0) )),
  inference(superposition,[],[f1,f172373])).
fof(f179301,plain,(
  ( ! [X10,X9] : (i(mult(i(X9),X10)) = rd(i(X10),i(X9))) )),
  inference(backward_demodulation,[],[f179298,f175861])).
fof(f175861,plain,(
  ( ! [X10,X9] : (mult(X9,ld(rd(X10,i(X9)),X9)) = rd(i(X10),i(X9))) )),
  inference(backward_demodulation,[],[f175857,f160698])).
fof(f160698,plain,(
  ( ! [X10,X9] : (mult(X9,ld(rd(X10,i(X9)),X9)) = ld(mult(X9,X10),mult(X9,X9))) )),
  inference(backward_demodulation,[],[f160695,f30033])).
fof(f179298,plain,(
  ( ! [X391,X390] : (i(mult(i(X390),X391)) = mult(X390,ld(rd(X391,i(X390)),X390))) )),
  inference(backward_demodulation,[],[f179159,f174182])).
fof(f174182,plain,(
  ( ! [X391,X390] : (mult(X390,mult(i(mult(X391,X390)),X390)) = i(mult(i(X390),X391))) )),
  inference(backward_demodulation,[],[f174174,f172320])).
fof(f172320,plain,(
  ( ! [X391,X390] : (mult(X390,mult(i(mult(X391,X390)),X390)) = rd(ld(ld(i(X390),X391),X390),i(X390))) )),
  inference(superposition,[],[f17392,f22238])).
fof(f174174,plain,(
  ( ! [X323,X322] : (i(mult(i(X322),X323)) = rd(ld(ld(i(X322),X323),X322),i(X322))) )),
  inference(backward_demodulation,[],[f174171,f82023])).
fof(f82023,plain,(
  ( ! [X323,X322] : (rd(rd(i(X322),rd(X323,X322)),i(X322)) = i(mult(i(X322),X323))) )),
  inference(forward_demodulation,[],[f82022,f20348])).
fof(f20348,plain,(
  ( ! [X10,X9] : (mult(i(X9),X10) = ld(X9,ld(X9,ld(i(X9),X10)))) )),
  inference(backward_demodulation,[],[f20347,f15919])).
fof(f82022,plain,(
  ( ! [X323,X322] : (rd(rd(i(X322),rd(X323,X322)),i(X322)) = i(ld(X322,ld(X322,ld(i(X322),X323))))) )),
  inference(forward_demodulation,[],[f82021,f14132])).
fof(f82021,plain,(
  ( ! [X323,X322] : (rd(rd(i(X322),rd(X323,X322)),i(X322)) = i(ld(X322,ld(i(X322),ld(X322,X323))))) )),
  inference(forward_demodulation,[],[f81864,f15829])).
fof(f81864,plain,(
  ( ! [X323,X322] : (rd(rd(i(X322),rd(X323,X322)),i(X322)) = i(rd(rd(ld(X322,X323),X322),i(X322)))) )),
  inference(superposition,[],[f25834,f12913])).
fof(f179159,plain,(
  ( ! [X292,X291] : (ld(rd(X291,i(X292)),X292) = mult(i(mult(X291,X292)),X292)) )),
  inference(backward_demodulation,[],[f179118,f174686])).
fof(f174686,plain,(
  ( ! [X292,X291] : (ld(rd(X291,i(X292)),X292) = mult(ld(X292,i(X291)),X292)) )),
  inference(backward_demodulation,[],[f174591,f173804])).
fof(f173804,plain,(
  ( ! [X292,X291] : (ld(rd(X291,i(X292)),X292) = mult(rd(X292,mult(X291,mult(X292,X292))),X292)) )),
  inference(superposition,[],[f173540,f37332])).
fof(f174591,plain,(
  ( ! [X74,X73] : (ld(X73,i(X74)) = rd(X73,mult(X74,mult(X73,X73)))) )),
  inference(superposition,[],[f4,f174499])).
fof(f179562,plain,(
  ( ! [X90,X89] : (i(ld(X89,ld(i(X89),X90))) = ld(X89,rd(X89,X90))) )),
  inference(forward_demodulation,[],[f179249,f179379])).
fof(f179379,plain,(
  ( ! [X37,X38] : (ld(X38,i(mult(X37,i(X38)))) = ld(X38,rd(X38,X37))) )),
  inference(backward_demodulation,[],[f179330,f175331])).
fof(f175331,plain,(
  ( ! [X37,X38] : (ld(X38,i(mult(X37,i(X38)))) = rd(X38,mult(X37,X38))) )),
  inference(superposition,[],[f174591,f37408])).
fof(f37408,plain,(
  ( ! [X99,X100] : (mult(X100,X99) = mult(mult(X100,i(X99)),mult(X99,X99))) )),
  inference(backward_demodulation,[],[f37374,f11252])).
fof(f37374,plain,(
  ( ! [X21,X20] : (mult(X21,X20) = mult(mult(X21,mult(X20,X20)),i(X20))) )),
  inference(backward_demodulation,[],[f37332,f29479])).
fof(f29479,plain,(
  ( ! [X21,X20] : (mult(X21,X20) = mult(mult(rd(X21,i(X20)),X20),i(X20))) )),
  inference(forward_demodulation,[],[f29294,f3669])).
fof(f29294,plain,(
  ( ! [X21,X20] : (mult(X21,X20) = mult(mult(rd(X21,i(X20)),X20),rd(i(mult(X20,X20)),i(X20)))) )),
  inference(superposition,[],[f16693,f4694])).
fof(f16693,plain,(
  ( ! [X2,X0,X1] : (mult(X2,X0) = mult(mult(rd(X2,mult(X0,X1)),X0),rd(X1,i(X0)))) )),
  inference(backward_demodulation,[],[f16669,f8273])).
fof(f8273,plain,(
  ( ! [X2,X0,X1] : (mult(X2,X0) = mult(mult(rd(X2,mult(X0,X1)),X0),ld(X0,mult(mult(X0,X1),X0)))) )),
  inference(backward_demodulation,[],[f8094,f1359])).
fof(f1359,plain,(
  ( ! [X2,X0,X1] : (mult(X2,X0) = mult(mult(rd(X2,mult(X0,X1)),X0),mult(X0,mult(ld(X0,X1),X0)))) )),
  inference(superposition,[],[f97,f1])).
fof(f179330,plain,(
  ( ! [X138,X137] : (rd(X138,mult(X137,X138)) = ld(X138,rd(X138,X137))) )),
  inference(forward_demodulation,[],[f179329,f26567])).
fof(f26567,plain,(
  ( ! [X39,X40] : (ld(X39,rd(X39,X40)) = ld(mult(X39,X40),X39)) )),
  inference(superposition,[],[f2,f26157])).
fof(f179329,plain,(
  ( ! [X138,X137] : (rd(X138,mult(X137,X138)) = ld(mult(X138,X137),X138)) )),
  inference(forward_demodulation,[],[f179163,f83])).
fof(f179163,plain,(
  ( ! [X138,X137] : (rd(X138,mult(X137,X138)) = ld(i(i(mult(X138,X137))),X138)) )),
  inference(backward_demodulation,[],[f179118,f174953])).
fof(f174953,plain,(
  ( ! [X138,X137] : (rd(X138,mult(X137,X138)) = ld(i(ld(X137,i(X138))),X138)) )),
  inference(forward_demodulation,[],[f174952,f1])).
fof(f174952,plain,(
  ( ! [X138,X137] : (ld(i(ld(X137,i(X138))),X138) = mult(X138,ld(X138,rd(X138,mult(X137,X138))))) )),
  inference(forward_demodulation,[],[f174915,f26567])).
fof(f174915,plain,(
  ( ! [X138,X137] : (ld(i(ld(X137,i(X138))),X138) = mult(X138,ld(mult(X138,mult(X137,X138)),X138))) )),
  inference(superposition,[],[f78751,f174296])).
fof(f174296,plain,(
  ( ! [X0,X1] : (mult(X0,mult(X1,X0)) = ld(ld(X1,i(X0)),X0)) )),
  inference(forward_demodulation,[],[f174224,f17392])).
fof(f174224,plain,(
  ( ! [X0,X1] : (rd(mult(X0,X1),i(X0)) = ld(ld(X1,i(X0)),X0)) )),
  inference(superposition,[],[f18809,f174124])).
fof(f179249,plain,(
  ( ! [X90,X89] : (i(ld(X89,ld(i(X89),X90))) = ld(X89,i(mult(X90,i(X89))))) )),
  inference(backward_demodulation,[],[f179118,f175648])).
fof(f175648,plain,(
  ( ! [X90,X89] : (i(ld(X89,ld(i(X89),X90))) = ld(X89,ld(i(X89),i(X90)))) )),
  inference(forward_demodulation,[],[f175647,f175372])).
fof(f175372,plain,(
  ( ! [X103,X102] : (ld(X102,rd(X102,rd(X103,X102))) = ld(i(X102),i(X103))) )),
  inference(backward_demodulation,[],[f175371,f172950])).
fof(f172950,plain,(
  ( ! [X103,X102] : (rd(i(X102),rd(X103,mult(X102,X102))) = ld(X102,rd(X102,rd(X103,X102)))) )),
  inference(backward_demodulation,[],[f172932,f38148])).
fof(f38148,plain,(
  ( ! [X103,X102] : (mult(i(ld(X102,rd(X103,X102))),i(X102)) = rd(i(X102),rd(X103,mult(X102,X102)))) )),
  inference(backward_demodulation,[],[f38126,f27920])).
fof(f27920,plain,(
  ( ! [X103,X102] : (mult(i(ld(X102,rd(X103,X102))),i(X102)) = rd(i(X102),mult(rd(X103,X102),i(X102)))) )),
  inference(backward_demodulation,[],[f27898,f20576])).
fof(f20576,plain,(
  ( ! [X103,X102] : (mult(i(ld(X102,rd(X103,X102))),i(X102)) = rd(i(X102),ld(X102,rd(mult(X102,rd(X103,X102)),X102)))) )),
  inference(forward_demodulation,[],[f20575,f15917])).
fof(f20575,plain,(
  ( ! [X103,X102] : (rd(i(X102),mult(i(X102),mult(mult(X102,rd(X103,X102)),i(X102)))) = mult(i(ld(X102,rd(X103,X102))),i(X102))) )),
  inference(forward_demodulation,[],[f20406,f15917])).
fof(f20406,plain,(
  ( ! [X103,X102] : (rd(i(X102),mult(i(X102),mult(mult(X102,rd(X103,X102)),i(X102)))) = mult(i(mult(i(X102),mult(X103,i(X102)))),i(X102))) )),
  inference(superposition,[],[f615,f17225])).
fof(f17225,plain,(
  ( ! [X8,X9] : (mult(X9,i(X8)) = mult(i(X8),mult(X8,rd(X9,X8)))) )),
  inference(backward_demodulation,[],[f17223,f11850])).
fof(f175371,plain,(
  ( ! [X35,X36] : (ld(i(X35),i(X36)) = rd(i(X35),rd(X36,mult(X35,X35)))) )),
  inference(forward_demodulation,[],[f175324,f33805])).
fof(f175324,plain,(
  ( ! [X35,X36] : (ld(i(X35),i(X36)) = rd(i(X35),mult(X36,i(mult(X35,X35))))) )),
  inference(superposition,[],[f174591,f3662])).
fof(f175647,plain,(
  ( ! [X90,X89] : (i(ld(X89,ld(i(X89),X90))) = ld(X89,ld(X89,rd(X89,rd(X90,X89))))) )),
  inference(forward_demodulation,[],[f175606,f26567])).
fof(f175606,plain,(
  ( ! [X90,X89] : (ld(X89,ld(mult(X89,rd(X90,X89)),X89)) = i(ld(X89,ld(i(X89),X90)))) )),
  inference(superposition,[],[f172455,f14556])).
fof(f172455,plain,(
  ( ! [X37,X36] : (i(ld(X36,mult(X37,X36))) = ld(X36,ld(X37,X36))) )),
  inference(superposition,[],[f2,f172363])).
fof(f11,axiom,(
  mult(a,i(mult(b,a))) != i(b)),
  file('GRP701-1.tptp',unknown)).
% SZS output end Proof for GRP701-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 233471
% Time elapsed: 10.692 s
% ------------------------------
% ------------------------------
