% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP305-1
% SZS output start Proof for GRP305-1
fof(f821,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f43,f48,f53,f58,f63,f65,f66,f72,f73,f74,f75,f76,f81,f82,f83,f84,f85,f90,f91,f92,f93,f94,f104,f166,f231,f241,f260,f273,f286,f400,f741,f760,f779,f796,f815,f820])).
fof(f820,plain,(
  ~spl0_1 | spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_10),
  inference(avatar_contradiction_clause,[],[f819])).
fof(f819,plain,(
  $false | (~spl0_1 | spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(subsumption_resolution,[],[f818,f507])).
fof(f507,plain,(
  identity = sk_c4 | (~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f505,f2])).
fof(f2,axiom,(
  ( ! [X0] : (identity = multiply(inverse(X0),X0)) )),
  file('GRP305-1.tptp',unknown)).
fof(f505,plain,(
  sk_c4 = multiply(inverse(sk_c6),sk_c6) | (~spl0_9 | ~spl0_10)),
  inference(superposition,[],[f116,f310])).
fof(f310,plain,(
  sk_c6 = multiply(sk_c6,sk_c4) | (~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f308,f89])).
fof(f89,plain,(
  sk_c6 = inverse(sk_c1) | ~spl0_10),
  inference(avatar_component_clause,[],[f87])).
fof(f87,plain,(
  spl0_10 <=> sk_c6 = inverse(sk_c1)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f308,plain,(
  sk_c6 = multiply(inverse(sk_c1),sk_c4) | ~spl0_9),
  inference(superposition,[],[f116,f80])).
fof(f80,plain,(
  sk_c4 = multiply(sk_c1,sk_c6) | ~spl0_9),
  inference(avatar_component_clause,[],[f78])).
fof(f78,plain,(
  spl0_9 <=> sk_c4 = multiply(sk_c1,sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f116,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = X7) )),
  inference(forward_demodulation,[],[f109,f1])).
fof(f1,axiom,(
  ( ! [X0] : (multiply(identity,X0) = X0) )),
  file('GRP305-1.tptp',unknown)).
fof(f109,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = multiply(identity,X7)) )),
  inference(superposition,[],[f3,f2])).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (multiply(multiply(X0,X1),X2) = multiply(X0,multiply(X1,X2))) )),
  file('GRP305-1.tptp',unknown)).
fof(f818,plain,(
  identity != sk_c4 | (~spl0_1 | spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f817,f1])).
fof(f817,plain,(
  sk_c4 != multiply(identity,identity) | (~spl0_1 | spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f816,f598])).
fof(f598,plain,(
  identity = sk_c5 | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f521,f1])).
fof(f521,plain,(
  sk_c5 = multiply(identity,identity) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f507,f299])).
fof(f299,plain,(
  sk_c5 = multiply(sk_c4,sk_c4) | (~spl0_1 | ~spl0_8)),
  inference(forward_demodulation,[],[f297,f71])).
fof(f71,plain,(
  sk_c4 = inverse(sk_c6) | ~spl0_8),
  inference(avatar_component_clause,[],[f69])).
fof(f69,plain,(
  spl0_8 <=> sk_c4 = inverse(sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f297,plain,(
  sk_c5 = multiply(inverse(sk_c6),sk_c4) | ~spl0_1),
  inference(superposition,[],[f116,f33])).
fof(f33,plain,(
  multiply(sk_c6,sk_c5) = sk_c4 | ~spl0_1),
  inference(avatar_component_clause,[],[f31])).
fof(f31,plain,(
  spl0_1 <=> multiply(sk_c6,sk_c5) = sk_c4),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f816,plain,(
  sk_c4 != multiply(sk_c5,identity) | (spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f36,f607])).
fof(f607,plain,(
  identity = sk_c6 | (~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f526,f2])).
fof(f526,plain,(
  sk_c6 = multiply(inverse(identity),identity) | (~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f507,f318])).
fof(f318,plain,(
  sk_c6 = multiply(inverse(sk_c4),identity) | ~spl0_8),
  inference(superposition,[],[f116,f294])).
fof(f294,plain,(
  identity = multiply(sk_c4,sk_c6) | ~spl0_8),
  inference(superposition,[],[f2,f71])).
fof(f36,plain,(
  sk_c4 != multiply(sk_c5,sk_c6) | spl0_2),
  inference(avatar_component_clause,[],[f35])).
fof(f35,plain,(
  spl0_2 <=> sk_c4 = multiply(sk_c5,sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f815,plain,(
  ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11),
  inference(avatar_contradiction_clause,[],[f814])).
fof(f814,plain,(
  $false | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(subsumption_resolution,[],[f813,f689])).
fof(f689,plain,(
  identity = inverse(identity) | (~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f512,f607])).
fof(f512,plain,(
  identity = inverse(sk_c6) | (~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f507,f71])).
fof(f813,plain,(
  identity != inverse(identity) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(forward_demodulation,[],[f809,f689])).
fof(f809,plain,(
  identity != inverse(inverse(identity)) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(trivial_inequality_removal,[],[f808])).
fof(f808,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(superposition,[],[f782,f2])).
fof(f782,plain,(
  ( ! [X3] : (identity != multiply(X3,identity) | identity != inverse(X3)) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(forward_demodulation,[],[f781,f507])).
fof(f781,plain,(
  ( ! [X3] : (sk_c4 != multiply(X3,identity) | identity != inverse(X3)) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(forward_demodulation,[],[f780,f607])).
fof(f780,plain,(
  ( ! [X3] : (identity != inverse(X3) | sk_c4 != multiply(X3,sk_c6)) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(forward_demodulation,[],[f97,f607])).
fof(f97,plain,(
  ( ! [X3] : (sk_c6 != inverse(X3) | sk_c4 != multiply(X3,sk_c6)) ) | ~spl0_11),
  inference(avatar_component_clause,[],[f96])).
fof(f96,plain,(
  spl0_11 <=> ! [X3] : (sk_c6 != inverse(X3) | sk_c4 != multiply(X3,sk_c6))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f796,plain,(
  ~spl0_1 | spl0_3 | ~spl0_4 | ~spl0_8 | ~spl0_9 | ~spl0_10),
  inference(avatar_contradiction_clause,[],[f795])).
fof(f795,plain,(
  $false | (~spl0_1 | spl0_3 | ~spl0_4 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(subsumption_resolution,[],[f792,f689])).
fof(f792,plain,(
  identity != inverse(identity) | (~spl0_1 | spl0_3 | ~spl0_4 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f783,f699])).
fof(f699,plain,(
  identity != inverse(sk_c2) | (spl0_3 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f41,f607])).
fof(f41,plain,(
  sk_c6 != inverse(sk_c2) | spl0_3),
  inference(avatar_component_clause,[],[f40])).
fof(f40,plain,(
  spl0_3 <=> sk_c6 = inverse(sk_c2)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f783,plain,(
  identity = sk_c2 | (~spl0_1 | ~spl0_4 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(superposition,[],[f713,f2])).
fof(f713,plain,(
  ( ! [X0] : (multiply(inverse(sk_c2),X0) = X0) ) | (~spl0_1 | ~spl0_4 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f712,f1])).
fof(f712,plain,(
  ( ! [X0] : (multiply(identity,X0) = multiply(inverse(sk_c2),X0)) ) | (~spl0_1 | ~spl0_4 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f711,f598])).
fof(f711,plain,(
  ( ! [X0] : (multiply(sk_c5,X0) = multiply(inverse(sk_c2),X0)) ) | (~spl0_4 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f710,f1])).
fof(f710,plain,(
  ( ! [X0] : (multiply(sk_c5,X0) = multiply(inverse(sk_c2),multiply(identity,X0))) ) | (~spl0_4 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f340,f607])).
fof(f340,plain,(
  ( ! [X0] : (multiply(sk_c5,X0) = multiply(inverse(sk_c2),multiply(sk_c6,X0))) ) | ~spl0_4),
  inference(superposition,[],[f116,f112])).
fof(f112,plain,(
  ( ! [X10] : (multiply(sk_c2,multiply(sk_c5,X10)) = multiply(sk_c6,X10)) ) | ~spl0_4),
  inference(superposition,[],[f3,f47])).
fof(f47,plain,(
  sk_c6 = multiply(sk_c2,sk_c5) | ~spl0_4),
  inference(avatar_component_clause,[],[f45])).
fof(f45,plain,(
  spl0_4 <=> sk_c6 = multiply(sk_c2,sk_c5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f779,plain,(
  ~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_12),
  inference(avatar_contradiction_clause,[],[f778])).
fof(f778,plain,(
  $false | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_12)),
  inference(subsumption_resolution,[],[f777,f689])).
fof(f777,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_12)),
  inference(forward_demodulation,[],[f773,f689])).
fof(f773,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_12)),
  inference(trivial_inequality_removal,[],[f768])).
fof(f768,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_12)),
  inference(superposition,[],[f763,f2])).
fof(f763,plain,(
  ( ! [X4] : (identity != multiply(X4,identity) | identity != inverse(X4)) ) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_12)),
  inference(forward_demodulation,[],[f762,f607])).
fof(f762,plain,(
  ( ! [X4] : (identity != multiply(X4,identity) | sk_c6 != inverse(X4)) ) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_12)),
  inference(forward_demodulation,[],[f761,f607])).
fof(f761,plain,(
  ( ! [X4] : (sk_c6 != multiply(X4,identity) | sk_c6 != inverse(X4)) ) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_12)),
  inference(forward_demodulation,[],[f100,f598])).
fof(f100,plain,(
  ( ! [X4] : (sk_c6 != multiply(X4,sk_c5) | sk_c6 != inverse(X4)) ) | ~spl0_12),
  inference(avatar_component_clause,[],[f99])).
fof(f99,plain,(
  spl0_12 <=> ! [X4] : (sk_c6 != multiply(X4,sk_c5) | sk_c6 != inverse(X4))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f760,plain,(
  ~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_13),
  inference(avatar_contradiction_clause,[],[f759])).
fof(f759,plain,(
  $false | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_13)),
  inference(subsumption_resolution,[],[f758,f689])).
fof(f758,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_13)),
  inference(forward_demodulation,[],[f754,f689])).
fof(f754,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_13)),
  inference(trivial_inequality_removal,[],[f751])).
fof(f751,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_13)),
  inference(superposition,[],[f736,f2])).
fof(f736,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_13)),
  inference(forward_demodulation,[],[f735,f598])).
fof(f735,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | sk_c5 != inverse(X5)) ) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_13)),
  inference(forward_demodulation,[],[f734,f598])).
fof(f734,plain,(
  ( ! [X5] : (sk_c5 != multiply(X5,identity) | sk_c5 != inverse(X5)) ) | (~spl0_9 | ~spl0_10 | ~spl0_13)),
  inference(forward_demodulation,[],[f103,f507])).
fof(f103,plain,(
  ( ! [X5] : (sk_c5 != multiply(X5,sk_c4) | sk_c5 != inverse(X5)) ) | ~spl0_13),
  inference(avatar_component_clause,[],[f102])).
fof(f102,plain,(
  spl0_13 <=> ! [X5] : (sk_c5 != multiply(X5,sk_c4) | sk_c5 != inverse(X5))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f741,plain,(
  ~spl0_1 | ~spl0_5 | spl0_6 | ~spl0_8 | ~spl0_9 | ~spl0_10),
  inference(avatar_contradiction_clause,[],[f740])).
fof(f740,plain,(
  $false | (~spl0_1 | ~spl0_5 | spl0_6 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(subsumption_resolution,[],[f739,f598])).
fof(f739,plain,(
  identity != sk_c5 | (~spl0_1 | ~spl0_5 | spl0_6 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f738,f1])).
fof(f738,plain,(
  sk_c5 != multiply(identity,identity) | (~spl0_1 | ~spl0_5 | spl0_6 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f737,f670])).
fof(f670,plain,(
  identity = sk_c3 | (~spl0_1 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f668,f1])).
fof(f668,plain,(
  sk_c3 = multiply(identity,identity) | (~spl0_1 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f607,f615])).
fof(f615,plain,(
  sk_c3 = multiply(sk_c6,identity) | (~spl0_1 | ~spl0_5 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f531,f1])).
fof(f531,plain,(
  multiply(sk_c6,identity) = multiply(identity,sk_c3) | (~spl0_1 | ~spl0_5 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f507,f496])).
fof(f496,plain,(
  multiply(sk_c6,identity) = multiply(sk_c4,sk_c3) | (~spl0_1 | ~spl0_5)),
  inference(superposition,[],[f298,f106])).
fof(f106,plain,(
  identity = multiply(sk_c5,sk_c3) | ~spl0_5),
  inference(superposition,[],[f2,f52])).
fof(f52,plain,(
  sk_c5 = inverse(sk_c3) | ~spl0_5),
  inference(avatar_component_clause,[],[f50])).
fof(f50,plain,(
  spl0_5 <=> sk_c5 = inverse(sk_c3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f298,plain,(
  ( ! [X0] : (multiply(sk_c4,X0) = multiply(sk_c6,multiply(sk_c5,X0))) ) | ~spl0_1),
  inference(superposition,[],[f3,f33])).
fof(f737,plain,(
  sk_c5 != multiply(sk_c3,identity) | (spl0_6 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f56,f507])).
fof(f56,plain,(
  sk_c5 != multiply(sk_c3,sk_c4) | spl0_6),
  inference(avatar_component_clause,[],[f55])).
fof(f55,plain,(
  spl0_6 <=> sk_c5 = multiply(sk_c3,sk_c4)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f400,plain,(
  spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10),
  inference(avatar_contradiction_clause,[],[f399])).
fof(f399,plain,(
  $false | (spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(subsumption_resolution,[],[f387,f370])).
fof(f370,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = X0) ) | (~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f367,f366])).
fof(f366,plain,(
  ( ! [X0] : (multiply(sk_c5,X0) = X0) ) | (~spl0_3 | ~spl0_4 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f355,f136])).
fof(f136,plain,(
  ( ! [X0] : (multiply(sk_c5,X0) = multiply(sk_c6,multiply(sk_c6,X0))) ) | (~spl0_3 | ~spl0_4)),
  inference(superposition,[],[f3,f132])).
fof(f132,plain,(
  sk_c5 = multiply(sk_c6,sk_c6) | (~spl0_3 | ~spl0_4)),
  inference(forward_demodulation,[],[f129,f42])).
fof(f42,plain,(
  sk_c6 = inverse(sk_c2) | ~spl0_3),
  inference(avatar_component_clause,[],[f40])).
fof(f129,plain,(
  sk_c5 = multiply(inverse(sk_c2),sk_c6) | ~spl0_4),
  inference(superposition,[],[f116,f47])).
fof(f355,plain,(
  ( ! [X0] : (multiply(sk_c6,multiply(sk_c6,X0)) = X0) ) | (~spl0_3 | ~spl0_4 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f343,f293])).
fof(f293,plain,(
  ( ! [X0] : (multiply(sk_c4,multiply(sk_c6,X0)) = X0) ) | ~spl0_8),
  inference(superposition,[],[f116,f71])).
fof(f343,plain,(
  sk_c6 = sk_c4 | (~spl0_3 | ~spl0_4 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f342,f328])).
fof(f328,plain,(
  sk_c4 = multiply(sk_c2,sk_c6) | (~spl0_3 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f324,f80])).
fof(f324,plain,(
  sk_c2 = sk_c1 | (~spl0_3 | ~spl0_8 | ~spl0_10)),
  inference(forward_demodulation,[],[f323,f287])).
fof(f287,plain,(
  sk_c2 = multiply(sk_c4,identity) | (~spl0_3 | ~spl0_8)),
  inference(forward_demodulation,[],[f126,f71])).
fof(f126,plain,(
  sk_c2 = multiply(inverse(sk_c6),identity) | ~spl0_3),
  inference(superposition,[],[f116,f105])).
fof(f105,plain,(
  identity = multiply(sk_c6,sk_c2) | ~spl0_3),
  inference(superposition,[],[f2,f42])).
fof(f323,plain,(
  sk_c1 = multiply(sk_c4,identity) | (~spl0_8 | ~spl0_10)),
  inference(forward_demodulation,[],[f321,f71])).
fof(f321,plain,(
  sk_c1 = multiply(inverse(sk_c6),identity) | ~spl0_10),
  inference(superposition,[],[f116,f296])).
fof(f296,plain,(
  identity = multiply(sk_c6,sk_c1) | ~spl0_10),
  inference(superposition,[],[f2,f89])).
fof(f342,plain,(
  sk_c6 = multiply(sk_c2,sk_c6) | (~spl0_4 | ~spl0_7 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f338,f310])).
fof(f338,plain,(
  multiply(sk_c2,sk_c6) = multiply(sk_c6,sk_c4) | (~spl0_4 | ~spl0_7)),
  inference(superposition,[],[f112,f62])).
fof(f62,plain,(
  sk_c6 = multiply(sk_c5,sk_c4) | ~spl0_7),
  inference(avatar_component_clause,[],[f60])).
fof(f60,plain,(
  spl0_7 <=> sk_c6 = multiply(sk_c5,sk_c4)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f367,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = multiply(sk_c5,X0)) ) | (~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f366,f352])).
fof(f352,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = multiply(sk_c5,multiply(sk_c5,X0))) ) | (~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f343,f145])).
fof(f145,plain,(
  ( ! [X0] : (multiply(sk_c4,X0) = multiply(sk_c5,multiply(sk_c5,X0))) ) | (~spl0_5 | ~spl0_6)),
  inference(superposition,[],[f3,f133])).
fof(f133,plain,(
  sk_c4 = multiply(sk_c5,sk_c5) | (~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f130,f52])).
fof(f130,plain,(
  sk_c4 = multiply(inverse(sk_c3),sk_c5) | ~spl0_6),
  inference(superposition,[],[f116,f57])).
fof(f57,plain,(
  sk_c5 = multiply(sk_c3,sk_c4) | ~spl0_6),
  inference(avatar_component_clause,[],[f55])).
fof(f387,plain,(
  sk_c6 != multiply(sk_c6,sk_c6) | (spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f373,f345])).
fof(f345,plain,(
  sk_c6 != multiply(sk_c5,sk_c6) | (spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f343,f36])).
fof(f373,plain,(
  sk_c6 = sk_c5 | (~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f370,f132])).
fof(f286,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_12),
  inference(avatar_contradiction_clause,[],[f285])).
fof(f285,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_12)),
  inference(subsumption_resolution,[],[f284,f194])).
fof(f194,plain,(
  identity = inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f192,f178])).
fof(f178,plain,(
  identity = inverse(sk_c2) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f177,f42])).
fof(f177,plain,(
  identity = sk_c6 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f161,f2])).
fof(f161,plain,(
  sk_c6 = multiply(inverse(sk_c6),sk_c6) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f146,f135])).
fof(f135,plain,(
  sk_c6 = multiply(inverse(sk_c6),sk_c5) | (~spl0_3 | ~spl0_4)),
  inference(superposition,[],[f116,f132])).
fof(f146,plain,(
  sk_c6 = sk_c5 | (~spl0_2 | ~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f144,f127])).
fof(f127,plain,(
  sk_c6 = multiply(inverse(sk_c5),sk_c4) | ~spl0_2),
  inference(superposition,[],[f116,f37])).
fof(f37,plain,(
  sk_c4 = multiply(sk_c5,sk_c6) | ~spl0_2),
  inference(avatar_component_clause,[],[f35])).
fof(f144,plain,(
  sk_c5 = multiply(inverse(sk_c5),sk_c4) | (~spl0_5 | ~spl0_6)),
  inference(superposition,[],[f116,f133])).
fof(f192,plain,(
  identity = sk_c2 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f181,f2])).
fof(f181,plain,(
  sk_c2 = multiply(inverse(identity),identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f177,f126])).
fof(f284,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_12)),
  inference(forward_demodulation,[],[f280,f194])).
fof(f280,plain,(
  identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_12)),
  inference(trivial_inequality_removal,[],[f279])).
fof(f279,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_12)),
  inference(superposition,[],[f276,f2])).
fof(f276,plain,(
  ( ! [X4] : (identity != multiply(X4,identity) | identity != inverse(X4)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_12)),
  inference(forward_demodulation,[],[f275,f177])).
fof(f275,plain,(
  ( ! [X4] : (identity != multiply(X4,identity) | sk_c6 != inverse(X4)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_12)),
  inference(forward_demodulation,[],[f274,f177])).
fof(f274,plain,(
  ( ! [X4] : (sk_c6 != multiply(X4,identity) | sk_c6 != inverse(X4)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_12)),
  inference(forward_demodulation,[],[f100,f183])).
fof(f183,plain,(
  identity = sk_c5 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f177,f146])).
fof(f273,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13),
  inference(avatar_contradiction_clause,[],[f272])).
fof(f272,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13)),
  inference(subsumption_resolution,[],[f271,f194])).
fof(f271,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13)),
  inference(forward_demodulation,[],[f267,f194])).
fof(f267,plain,(
  identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13)),
  inference(trivial_inequality_removal,[],[f266])).
fof(f266,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13)),
  inference(superposition,[],[f263,f2])).
fof(f263,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13)),
  inference(forward_demodulation,[],[f262,f183])).
fof(f262,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | sk_c5 != inverse(X5)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13)),
  inference(forward_demodulation,[],[f261,f183])).
fof(f261,plain,(
  ( ! [X5] : (sk_c5 != multiply(X5,identity) | sk_c5 != inverse(X5)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13)),
  inference(forward_demodulation,[],[f103,f202])).
fof(f202,plain,(
  identity = sk_c4 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f201,f182])).
fof(f182,plain,(
  sk_c4 = multiply(sk_c4,identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f177,f141])).
fof(f141,plain,(
  sk_c4 = multiply(sk_c4,sk_c6) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f138,f133])).
fof(f138,plain,(
  multiply(sk_c5,sk_c5) = multiply(sk_c4,sk_c6) | (~spl0_2 | ~spl0_3 | ~spl0_4)),
  inference(superposition,[],[f111,f132])).
fof(f111,plain,(
  ( ! [X9] : (multiply(sk_c5,multiply(sk_c6,X9)) = multiply(sk_c4,X9)) ) | ~spl0_2),
  inference(superposition,[],[f3,f37])).
fof(f201,plain,(
  ( ! [X9] : (multiply(sk_c4,X9) = X9) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f200,f1])).
fof(f200,plain,(
  ( ! [X9] : (multiply(sk_c4,X9) = multiply(identity,X9)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f185,f1])).
fof(f185,plain,(
  ( ! [X9] : (multiply(sk_c4,X9) = multiply(identity,multiply(identity,X9))) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f177,f153])).
fof(f153,plain,(
  ( ! [X9] : (multiply(sk_c4,X9) = multiply(sk_c6,multiply(sk_c6,X9))) ) | (~spl0_2 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f146,f111])).
fof(f260,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_11),
  inference(avatar_contradiction_clause,[],[f259])).
fof(f259,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_11)),
  inference(subsumption_resolution,[],[f258,f194])).
fof(f258,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_11)),
  inference(forward_demodulation,[],[f254,f194])).
fof(f254,plain,(
  identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_11)),
  inference(trivial_inequality_removal,[],[f253])).
fof(f253,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_11)),
  inference(superposition,[],[f244,f2])).
fof(f244,plain,(
  ( ! [X3] : (identity != multiply(X3,identity) | identity != inverse(X3)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_11)),
  inference(forward_demodulation,[],[f243,f202])).
fof(f243,plain,(
  ( ! [X3] : (sk_c4 != multiply(X3,identity) | identity != inverse(X3)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_11)),
  inference(forward_demodulation,[],[f242,f177])).
fof(f242,plain,(
  ( ! [X3] : (identity != inverse(X3) | sk_c4 != multiply(X3,sk_c6)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_11)),
  inference(forward_demodulation,[],[f97,f177])).
fof(f241,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_7),
  inference(avatar_contradiction_clause,[],[f240])).
fof(f240,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_7)),
  inference(subsumption_resolution,[],[f239,f177])).
fof(f239,plain,(
  identity != sk_c6 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_7)),
  inference(forward_demodulation,[],[f238,f1])).
fof(f238,plain,(
  sk_c6 != multiply(identity,identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_7)),
  inference(forward_demodulation,[],[f237,f183])).
fof(f237,plain,(
  sk_c6 != multiply(sk_c5,identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_7)),
  inference(forward_demodulation,[],[f61,f202])).
fof(f61,plain,(
  sk_c6 != multiply(sk_c5,sk_c4) | spl0_7),
  inference(avatar_component_clause,[],[f60])).
fof(f231,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_8),
  inference(avatar_contradiction_clause,[],[f230])).
fof(f230,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_8)),
  inference(subsumption_resolution,[],[f229,f202])).
fof(f229,plain,(
  identity != sk_c4 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_8)),
  inference(forward_demodulation,[],[f228,f194])).
fof(f228,plain,(
  sk_c4 != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_8)),
  inference(forward_demodulation,[],[f70,f177])).
fof(f70,plain,(
  sk_c4 != inverse(sk_c6) | spl0_8),
  inference(avatar_component_clause,[],[f69])).
fof(f166,plain,(
  spl0_1 | ~spl0_2 | ~spl0_5 | ~spl0_6),
  inference(avatar_contradiction_clause,[],[f165])).
fof(f165,plain,(
  $false | (spl0_1 | ~spl0_2 | ~spl0_5 | ~spl0_6)),
  inference(subsumption_resolution,[],[f148,f147])).
fof(f147,plain,(
  sk_c4 != multiply(sk_c6,sk_c6) | (spl0_1 | ~spl0_2 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f146,f32])).
fof(f32,plain,(
  multiply(sk_c6,sk_c5) != sk_c4 | spl0_1),
  inference(avatar_component_clause,[],[f31])).
fof(f148,plain,(
  sk_c4 = multiply(sk_c6,sk_c6) | (~spl0_2 | ~spl0_5 | ~spl0_6)),
  inference(backward_demodulation,[],[f146,f37])).
fof(f104,plain,(
  ~spl0_1 | ~spl0_7 | ~spl0_8 | spl0_11 | ~spl0_2 | spl0_12 | spl0_13),
  inference(avatar_split_clause,[],[f29,f102,f99,f35,f96,f69,f60,f31])).
fof(f29,axiom,(
  ( ! [X4,X5,X3] : (sk_c5 != multiply(X5,sk_c4) | sk_c5 != inverse(X5) | sk_c6 != multiply(X4,sk_c5) | sk_c6 != inverse(X4) | sk_c4 != multiply(sk_c5,sk_c6) | sk_c6 != inverse(X3) | sk_c4 != multiply(X3,sk_c6) | sk_c4 != inverse(sk_c6) | sk_c6 != multiply(sk_c5,sk_c4) | multiply(sk_c6,sk_c5) != sk_c4) )),
  file('GRP305-1.tptp',unknown)).
fof(f94,plain,(
  spl0_10 | spl0_6),
  inference(avatar_split_clause,[],[f28,f55,f87])).
fof(f28,axiom,(
  sk_c5 = multiply(sk_c3,sk_c4) | sk_c6 = inverse(sk_c1)),
  file('GRP305-1.tptp',unknown)).
fof(f93,plain,(
  spl0_10 | spl0_5),
  inference(avatar_split_clause,[],[f27,f50,f87])).
fof(f27,axiom,(
  sk_c5 = inverse(sk_c3) | sk_c6 = inverse(sk_c1)),
  file('GRP305-1.tptp',unknown)).
fof(f92,plain,(
  spl0_10 | spl0_4),
  inference(avatar_split_clause,[],[f26,f45,f87])).
fof(f26,axiom,(
  sk_c6 = multiply(sk_c2,sk_c5) | sk_c6 = inverse(sk_c1)),
  file('GRP305-1.tptp',unknown)).
fof(f91,plain,(
  spl0_10 | spl0_3),
  inference(avatar_split_clause,[],[f25,f40,f87])).
fof(f25,axiom,(
  sk_c6 = inverse(sk_c2) | sk_c6 = inverse(sk_c1)),
  file('GRP305-1.tptp',unknown)).
fof(f90,plain,(
  spl0_10 | spl0_2),
  inference(avatar_split_clause,[],[f24,f35,f87])).
fof(f24,axiom,(
  sk_c4 = multiply(sk_c5,sk_c6) | sk_c6 = inverse(sk_c1)),
  file('GRP305-1.tptp',unknown)).
fof(f85,plain,(
  spl0_9 | spl0_6),
  inference(avatar_split_clause,[],[f23,f55,f78])).
fof(f23,axiom,(
  sk_c5 = multiply(sk_c3,sk_c4) | sk_c4 = multiply(sk_c1,sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f84,plain,(
  spl0_9 | spl0_5),
  inference(avatar_split_clause,[],[f22,f50,f78])).
fof(f22,axiom,(
  sk_c5 = inverse(sk_c3) | sk_c4 = multiply(sk_c1,sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f83,plain,(
  spl0_9 | spl0_4),
  inference(avatar_split_clause,[],[f21,f45,f78])).
fof(f21,axiom,(
  sk_c6 = multiply(sk_c2,sk_c5) | sk_c4 = multiply(sk_c1,sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f82,plain,(
  spl0_9 | spl0_3),
  inference(avatar_split_clause,[],[f20,f40,f78])).
fof(f20,axiom,(
  sk_c6 = inverse(sk_c2) | sk_c4 = multiply(sk_c1,sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f81,plain,(
  spl0_9 | spl0_2),
  inference(avatar_split_clause,[],[f19,f35,f78])).
fof(f19,axiom,(
  sk_c4 = multiply(sk_c5,sk_c6) | sk_c4 = multiply(sk_c1,sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f76,plain,(
  spl0_8 | spl0_6),
  inference(avatar_split_clause,[],[f18,f55,f69])).
fof(f18,axiom,(
  sk_c5 = multiply(sk_c3,sk_c4) | sk_c4 = inverse(sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f75,plain,(
  spl0_8 | spl0_5),
  inference(avatar_split_clause,[],[f17,f50,f69])).
fof(f17,axiom,(
  sk_c5 = inverse(sk_c3) | sk_c4 = inverse(sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f74,plain,(
  spl0_8 | spl0_4),
  inference(avatar_split_clause,[],[f16,f45,f69])).
fof(f16,axiom,(
  sk_c6 = multiply(sk_c2,sk_c5) | sk_c4 = inverse(sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f73,plain,(
  spl0_8 | spl0_3),
  inference(avatar_split_clause,[],[f15,f40,f69])).
fof(f15,axiom,(
  sk_c6 = inverse(sk_c2) | sk_c4 = inverse(sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f72,plain,(
  spl0_8 | spl0_2),
  inference(avatar_split_clause,[],[f14,f35,f69])).
fof(f14,axiom,(
  sk_c4 = multiply(sk_c5,sk_c6) | sk_c4 = inverse(sk_c6)),
  file('GRP305-1.tptp',unknown)).
fof(f66,plain,(
  spl0_7 | spl0_5),
  inference(avatar_split_clause,[],[f12,f50,f60])).
fof(f12,axiom,(
  sk_c5 = inverse(sk_c3) | sk_c6 = multiply(sk_c5,sk_c4)),
  file('GRP305-1.tptp',unknown)).
fof(f65,plain,(
  spl0_7 | spl0_4),
  inference(avatar_split_clause,[],[f11,f45,f60])).
fof(f11,axiom,(
  sk_c6 = multiply(sk_c2,sk_c5) | sk_c6 = multiply(sk_c5,sk_c4)),
  file('GRP305-1.tptp',unknown)).
fof(f63,plain,(
  spl0_7 | spl0_2),
  inference(avatar_split_clause,[],[f9,f35,f60])).
fof(f9,axiom,(
  sk_c4 = multiply(sk_c5,sk_c6) | sk_c6 = multiply(sk_c5,sk_c4)),
  file('GRP305-1.tptp',unknown)).
fof(f58,plain,(
  spl0_1 | spl0_6),
  inference(avatar_split_clause,[],[f8,f55,f31])).
fof(f8,axiom,(
  sk_c5 = multiply(sk_c3,sk_c4) | multiply(sk_c6,sk_c5) = sk_c4),
  file('GRP305-1.tptp',unknown)).
fof(f53,plain,(
  spl0_1 | spl0_5),
  inference(avatar_split_clause,[],[f7,f50,f31])).
fof(f7,axiom,(
  sk_c5 = inverse(sk_c3) | multiply(sk_c6,sk_c5) = sk_c4),
  file('GRP305-1.tptp',unknown)).
fof(f48,plain,(
  spl0_1 | spl0_4),
  inference(avatar_split_clause,[],[f6,f45,f31])).
fof(f6,axiom,(
  sk_c6 = multiply(sk_c2,sk_c5) | multiply(sk_c6,sk_c5) = sk_c4),
  file('GRP305-1.tptp',unknown)).
fof(f43,plain,(
  spl0_1 | spl0_3),
  inference(avatar_split_clause,[],[f5,f40,f31])).
fof(f5,axiom,(
  sk_c6 = inverse(sk_c2) | multiply(sk_c6,sk_c5) = sk_c4),
  file('GRP305-1.tptp',unknown)).
% SZS output end Proof for GRP305-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
