% Refutation found. Thanks to Tanya!
% SZS status Theorem for GRP665+1
% SZS output start Proof for GRP665+1
fof(f361,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f46,f341,f357,f360])).
fof(f360,plain,(
  spl2_3),
  inference(avatar_contradiction_clause,[],[f359])).
fof(f359,plain,(
  $false | spl2_3),
  inference(subsumption_resolution,[],[f358,f96])).
fof(f96,plain,(
  ( ! [X6,X5] : (mult(X5,mult(op_c,X6)) = mult(op_c,mult(X5,X6))) )),
  inference(forward_demodulation,[],[f79,f24])).
fof(f24,plain,(
  ( ! [X0,X1] : (rd(mult(X1,X0),X0) = X1) )),
  inference(cnf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : rd(mult(X1,X0),X0) = X1),
  file('GRP665+1.tptp',unknown)).
fof(f79,plain,(
  ( ! [X6,X5] : (mult(X5,mult(op_c,X6)) = mult(rd(mult(op_c,X5),X5),mult(X5,X6))) )),
  inference(superposition,[],[f27,f29])).
fof(f29,plain,(
  ( ! [X0] : (mult(op_c,X0) = mult(X0,op_c)) )),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0] : mult(op_c,X0) = mult(X0,op_c)),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X1] : mult(op_c,X1) = mult(X1,op_c)),
  file('GRP665+1.tptp',unknown)).
fof(f27,plain,(
  ( ! [X2,X0,X1] : (mult(X2,mult(X1,X0)) = mult(rd(mult(X2,X1),X2),mult(X2,X0))) )),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2] : mult(X2,mult(X1,X0)) = mult(rd(mult(X2,X1),X2),mult(X2,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X2,X0,X1] : mult(X1,mult(X0,X2)) = mult(rd(mult(X1,X0),X1),mult(X1,X2))),
  file('GRP665+1.tptp',unknown)).
fof(f358,plain,(
  mult(sK0,mult(op_c,sK1)) != mult(op_c,mult(sK0,sK1)) | spl2_3),
  inference(forward_demodulation,[],[f45,f139])).
fof(f139,plain,(
  ( ! [X6,X5] : (mult(op_c,mult(X5,X6)) = mult(mult(op_c,X5),X6)) )),
  inference(forward_demodulation,[],[f138,f29])).
fof(f138,plain,(
  ( ! [X6,X5] : (mult(mult(X5,X6),op_c) = mult(mult(op_c,X5),X6)) )),
  inference(forward_demodulation,[],[f117,f67])).
fof(f67,plain,(
  ( ! [X1] : (ld(op_c,mult(X1,op_c)) = X1) )),
  inference(superposition,[],[f22,f29])).
fof(f22,plain,(
  ( ! [X0,X1] : (ld(X1,mult(X1,X0)) = X0) )),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1] : ld(X1,mult(X1,X0)) = X0),
  file('GRP665+1.tptp',unknown)).
fof(f117,plain,(
  ( ! [X6,X5] : (mult(mult(X5,X6),op_c) = mult(mult(op_c,X5),ld(op_c,mult(X6,op_c)))) )),
  inference(superposition,[],[f28,f29])).
fof(f28,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X2,X1),X0) = mult(mult(X2,X0),ld(X0,mult(X1,X0)))) )),
  inference(cnf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0,X1,X2] : mult(mult(X2,X1),X0) = mult(mult(X2,X0),ld(X0,mult(X1,X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X2,X0,X1] : mult(mult(X1,X0),X2) = mult(mult(X1,X2),ld(X2,mult(X0,X2)))),
  file('GRP665+1.tptp',unknown)).
fof(f45,plain,(
  mult(sK0,mult(op_c,sK1)) != mult(mult(op_c,sK0),sK1) | spl2_3),
  inference(avatar_component_clause,[],[f43])).
fof(f43,plain,(
  spl2_3 <=> mult(sK0,mult(op_c,sK1)) = mult(mult(op_c,sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl2_3])])).
fof(f357,plain,(
  spl2_2),
  inference(avatar_contradiction_clause,[],[f356])).
fof(f356,plain,(
  $false | spl2_2),
  inference(subsumption_resolution,[],[f41,f96])).
fof(f41,plain,(
  mult(sK0,mult(op_c,sK1)) != mult(op_c,mult(sK0,sK1)) | spl2_2),
  inference(avatar_component_clause,[],[f39])).
fof(f39,plain,(
  spl2_2 <=> mult(sK0,mult(op_c,sK1)) = mult(op_c,mult(sK0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_2])])).
fof(f341,plain,(
  spl2_1),
  inference(avatar_contradiction_clause,[],[f340])).
fof(f340,plain,(
  $false | spl2_1),
  inference(trivial_inequality_removal,[],[f324])).
fof(f324,plain,(
  mult(op_c,mult(sK0,sK1)) != mult(op_c,mult(sK0,sK1)) | spl2_1),
  inference(superposition,[],[f37,f139])).
fof(f37,plain,(
  mult(op_c,mult(sK0,sK1)) != mult(mult(op_c,sK0),sK1) | spl2_1),
  inference(avatar_component_clause,[],[f35])).
fof(f35,plain,(
  spl2_1 <=> mult(op_c,mult(sK0,sK1)) = mult(mult(op_c,sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl2_1])])).
fof(f46,plain,(
  ~spl2_1 | ~spl2_2 | ~spl2_3),
  inference(avatar_split_clause,[],[f33,f43,f39,f35])).
fof(f33,plain,(
  mult(sK0,mult(op_c,sK1)) != mult(mult(op_c,sK0),sK1) | mult(sK0,mult(op_c,sK1)) != mult(op_c,mult(sK0,sK1)) | mult(op_c,mult(sK0,sK1)) != mult(mult(op_c,sK0),sK1)),
  inference(forward_demodulation,[],[f32,f29])).
fof(f32,plain,(
  mult(sK0,mult(op_c,sK1)) != mult(op_c,mult(sK0,sK1)) | mult(mult(sK0,op_c),sK1) != mult(sK0,mult(op_c,sK1)) | mult(op_c,mult(sK0,sK1)) != mult(mult(op_c,sK0),sK1)),
  inference(forward_demodulation,[],[f31,f29])).
fof(f31,plain,(
  mult(mult(sK0,sK1),op_c) != mult(sK0,mult(op_c,sK1)) | mult(mult(sK0,op_c),sK1) != mult(sK0,mult(op_c,sK1)) | mult(op_c,mult(sK0,sK1)) != mult(mult(op_c,sK0),sK1)),
  inference(forward_demodulation,[],[f30,f29])).
fof(f30,plain,(
  mult(mult(sK0,sK1),op_c) != mult(sK0,mult(sK1,op_c)) | mult(mult(sK0,op_c),sK1) != mult(sK0,mult(op_c,sK1)) | mult(op_c,mult(sK0,sK1)) != mult(mult(op_c,sK0),sK1)),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  mult(mult(sK0,sK1),op_c) != mult(sK0,mult(sK1,op_c)) | mult(mult(sK0,op_c),sK1) != mult(sK0,mult(op_c,sK1)) | mult(op_c,mult(sK0,sK1)) != mult(mult(op_c,sK0),sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f18,f19])).
fof(f19,plain,(
  ? [X0,X1] : (mult(mult(X0,X1),op_c) != mult(X0,mult(X1,op_c)) | mult(mult(X0,op_c),X1) != mult(X0,mult(op_c,X1)) | mult(op_c,mult(X0,X1)) != mult(mult(op_c,X0),X1)) => (mult(mult(sK0,sK1),op_c) != mult(sK0,mult(sK1,op_c)) | mult(mult(sK0,op_c),sK1) != mult(sK0,mult(op_c,sK1)) | mult(op_c,mult(sK0,sK1)) != mult(mult(op_c,sK0),sK1))),
  introduced(choice_axiom,[])).
fof(f18,plain,(
  ? [X0,X1] : (mult(mult(X0,X1),op_c) != mult(X0,mult(X1,op_c)) | mult(mult(X0,op_c),X1) != mult(X0,mult(op_c,X1)) | mult(op_c,mult(X0,X1)) != mult(mult(op_c,X0),X1))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1] : (mult(mult(X0,X1),op_c) = mult(X0,mult(X1,op_c)) & mult(mult(X0,op_c),X1) = mult(X0,mult(op_c,X1)) & mult(op_c,mult(X0,X1)) = mult(mult(op_c,X0),X1))),
  inference(rectify,[],[f11])).
fof(f11,negated_conjecture,(
  ~! [X3,X4] : (mult(mult(X3,X4),op_c) = mult(X3,mult(X4,op_c)) & mult(mult(X3,op_c),X4) = mult(X3,mult(op_c,X4)) & mult(op_c,mult(X3,X4)) = mult(mult(op_c,X3),X4))),
  inference(negated_conjecture,[],[f10])).
fof(f10,conjecture,(
  ! [X3,X4] : (mult(mult(X3,X4),op_c) = mult(X3,mult(X4,op_c)) & mult(mult(X3,op_c),X4) = mult(X3,mult(op_c,X4)) & mult(op_c,mult(X3,X4)) = mult(mult(op_c,X3),X4))),
  file('GRP665+1.tptp',unknown)).
% SZS output end Proof for GRP665+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
