% Refutation found. Thanks to Tanya!
% SZS status Theorem for GEO196+1
% SZS output start Proof for GEO196+1
fof(f2581,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f68,f1532,f1552,f1561,f1598,f2410,f2580])).
fof(f2580,plain,(
  ~spl4_28),
  inference(avatar_contradiction_clause,[],[f2579])).
fof(f2579,plain,(
  $false | ~spl4_28),
  inference(subsumption_resolution,[],[f2576,f55])).
fof(f55,plain,(
  convergent_lines(sK0,sK1)),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  (apart_point_and_line(intersection_point(sK2,sK3),sK1) | apart_point_and_line(intersection_point(sK2,sK3),sK0)) & ~apart_point_and_line(intersection_point(sK0,sK1),sK3) & ~apart_point_and_line(intersection_point(sK0,sK1),sK2) & convergent_lines(sK2,sK3) & convergent_lines(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f38,f39])).
fof(f39,plain,(
  ? [X0,X1,X2,X3] : ((apart_point_and_line(intersection_point(X2,X3),X1) | apart_point_and_line(intersection_point(X2,X3),X0)) & ~apart_point_and_line(intersection_point(X0,X1),X3) & ~apart_point_and_line(intersection_point(X0,X1),X2) & convergent_lines(X2,X3) & convergent_lines(X0,X1)) => ((apart_point_and_line(intersection_point(sK2,sK3),sK1) | apart_point_and_line(intersection_point(sK2,sK3),sK0)) & ~apart_point_and_line(intersection_point(sK0,sK1),sK3) & ~apart_point_and_line(intersection_point(sK0,sK1),sK2) & convergent_lines(sK2,sK3) & convergent_lines(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f38,plain,(
  ? [X0,X1,X2,X3] : ((apart_point_and_line(intersection_point(X2,X3),X1) | apart_point_and_line(intersection_point(X2,X3),X0)) & ~apart_point_and_line(intersection_point(X0,X1),X3) & ~apart_point_and_line(intersection_point(X0,X1),X2) & convergent_lines(X2,X3) & convergent_lines(X0,X1))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ? [X0,X1,X2,X3] : ((apart_point_and_line(intersection_point(X2,X3),X1) | apart_point_and_line(intersection_point(X2,X3),X0)) & (~apart_point_and_line(intersection_point(X0,X1),X3) & ~apart_point_and_line(intersection_point(X0,X1),X2) & convergent_lines(X2,X3) & convergent_lines(X0,X1)))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2,X3] : ((~apart_point_and_line(intersection_point(X0,X1),X3) & ~apart_point_and_line(intersection_point(X0,X1),X2) & convergent_lines(X2,X3) & convergent_lines(X0,X1)) => (~apart_point_and_line(intersection_point(X2,X3),X1) & ~apart_point_and_line(intersection_point(X2,X3),X0)))),
  inference(rectify,[],[f16])).
fof(f16,negated_conjecture,(
  ~! [X0,X1,X3,X4] : ((~apart_point_and_line(intersection_point(X0,X1),X4) & ~apart_point_and_line(intersection_point(X0,X1),X3) & convergent_lines(X3,X4) & convergent_lines(X0,X1)) => (~apart_point_and_line(intersection_point(X3,X4),X1) & ~apart_point_and_line(intersection_point(X3,X4),X0)))),
  inference(negated_conjecture,[],[f15])).
fof(f15,conjecture,(
  ! [X0,X1,X3,X4] : ((~apart_point_and_line(intersection_point(X0,X1),X4) & ~apart_point_and_line(intersection_point(X0,X1),X3) & convergent_lines(X3,X4) & convergent_lines(X0,X1)) => (~apart_point_and_line(intersection_point(X3,X4),X1) & ~apart_point_and_line(intersection_point(X3,X4),X0)))),
  file('GEO196+1.tptp',unknown)).
fof(f2576,plain,(
  ~convergent_lines(sK0,sK1) | ~spl4_28),
  inference(resolution,[],[f2561,f49])).
fof(f49,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(intersection_point(X0,X1),X0) | ~convergent_lines(X0,X1)) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : (~apart_point_and_line(intersection_point(X0,X1),X0) | ~convergent_lines(X0,X1))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1] : (convergent_lines(X0,X1) => ~apart_point_and_line(intersection_point(X0,X1),X0))),
  file('GEO196+1.tptp',unknown)).
fof(f2561,plain,(
  apart_point_and_line(intersection_point(sK0,sK1),sK0) | ~spl4_28),
  inference(subsumption_resolution,[],[f2555,f57])).
fof(f57,plain,(
  ~apart_point_and_line(intersection_point(sK0,sK1),sK2)),
  inference(cnf_transformation,[],[f40])).
fof(f2555,plain,(
  apart_point_and_line(intersection_point(sK0,sK1),sK0) | apart_point_and_line(intersection_point(sK0,sK1),sK2) | ~spl4_28),
  inference(resolution,[],[f1551,f58])).
fof(f58,plain,(
  ~apart_point_and_line(intersection_point(sK0,sK1),sK3)),
  inference(cnf_transformation,[],[f40])).
fof(f1551,plain,(
  ( ! [X0] : (apart_point_and_line(X0,sK3) | apart_point_and_line(X0,sK0) | apart_point_and_line(X0,sK2)) ) | ~spl4_28),
  inference(avatar_component_clause,[],[f1550])).
fof(f1550,plain,(
  spl4_28 <=> ! [X0] : (apart_point_and_line(X0,sK0) | apart_point_and_line(X0,sK3) | apart_point_and_line(X0,sK2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_28])])).
fof(f2410,plain,(
  ~spl4_7),
  inference(avatar_contradiction_clause,[],[f2409])).
fof(f2409,plain,(
  $false | ~spl4_7),
  inference(subsumption_resolution,[],[f2406,f55])).
fof(f2406,plain,(
  ~convergent_lines(sK0,sK1) | ~spl4_7),
  inference(resolution,[],[f2374,f50])).
fof(f50,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(intersection_point(X0,X1),X1) | ~convergent_lines(X0,X1)) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : (~apart_point_and_line(intersection_point(X0,X1),X1) | ~convergent_lines(X0,X1))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1] : (convergent_lines(X0,X1) => ~apart_point_and_line(intersection_point(X0,X1),X1))),
  file('GEO196+1.tptp',unknown)).
fof(f2374,plain,(
  apart_point_and_line(intersection_point(sK0,sK1),sK1) | ~spl4_7),
  inference(subsumption_resolution,[],[f2368,f57])).
fof(f2368,plain,(
  apart_point_and_line(intersection_point(sK0,sK1),sK1) | apart_point_and_line(intersection_point(sK0,sK1),sK2) | ~spl4_7),
  inference(resolution,[],[f1445,f58])).
fof(f1445,plain,(
  ( ! [X0] : (apart_point_and_line(X0,sK3) | apart_point_and_line(X0,sK1) | apart_point_and_line(X0,sK2)) ) | ~spl4_7),
  inference(avatar_component_clause,[],[f1444])).
fof(f1444,plain,(
  spl4_7 <=> ! [X0] : (apart_point_and_line(X0,sK2) | apart_point_and_line(X0,sK1) | apart_point_and_line(X0,sK3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_7])])).
fof(f1598,plain,(
  ~spl4_9),
  inference(avatar_contradiction_clause,[],[f1597])).
fof(f1597,plain,(
  $false | ~spl4_9),
  inference(subsumption_resolution,[],[f1594,f56])).
fof(f56,plain,(
  convergent_lines(sK2,sK3)),
  inference(cnf_transformation,[],[f40])).
fof(f1594,plain,(
  ~convergent_lines(sK2,sK3) | ~spl4_9),
  inference(resolution,[],[f1453,f49])).
fof(f1453,plain,(
  apart_point_and_line(intersection_point(sK2,sK3),sK2) | ~spl4_9),
  inference(avatar_component_clause,[],[f1451])).
fof(f1451,plain,(
  spl4_9 <=> apart_point_and_line(intersection_point(sK2,sK3),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_9])])).
fof(f1561,plain,(
  ~spl4_8),
  inference(avatar_contradiction_clause,[],[f1560])).
fof(f1560,plain,(
  $false | ~spl4_8),
  inference(subsumption_resolution,[],[f1557,f56])).
fof(f1557,plain,(
  ~convergent_lines(sK2,sK3) | ~spl4_8),
  inference(resolution,[],[f1449,f50])).
fof(f1449,plain,(
  apart_point_and_line(intersection_point(sK2,sK3),sK3) | ~spl4_8),
  inference(avatar_component_clause,[],[f1447])).
fof(f1447,plain,(
  spl4_8 <=> apart_point_and_line(intersection_point(sK2,sK3),sK3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_8])])).
fof(f1552,plain,(
  spl4_8 | spl4_9 | spl4_28 | ~spl4_1),
  inference(avatar_split_clause,[],[f1545,f61,f1550,f1451,f1447])).
fof(f61,plain,(
  spl4_1 <=> apart_point_and_line(intersection_point(sK2,sK3),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_1])])).
fof(f1545,plain,(
  ( ! [X0] : (apart_point_and_line(X0,sK0) | apart_point_and_line(intersection_point(sK2,sK3),sK2) | apart_point_and_line(intersection_point(sK2,sK3),sK3) | apart_point_and_line(X0,sK2) | apart_point_and_line(X0,sK3)) ) | ~spl4_1),
  inference(resolution,[],[f1534,f89])).
fof(f89,plain,(
  ( ! [X0,X1] : (~distinct_points(X1,X0) | apart_point_and_line(X1,sK2) | apart_point_and_line(X1,sK3) | apart_point_and_line(X0,sK2) | apart_point_and_line(X0,sK3)) )),
  inference(resolution,[],[f86,f51])).
fof(f51,plain,(
  ( ! [X2,X0,X3,X1] : (~distinct_lines(X2,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2) | apart_point_and_line(X1,X3) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2,X3] : (apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2) | ~distinct_lines(X2,X3) | ~distinct_points(X0,X1))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : ((apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2)) | (~distinct_lines(X2,X3) | ~distinct_points(X0,X1)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2,X3] : ((distinct_lines(X2,X3) & distinct_points(X0,X1)) => (apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X3,X4] : ((distinct_lines(X3,X4) & distinct_points(X0,X1)) => (apart_point_and_line(X1,X4) | apart_point_and_line(X1,X3) | apart_point_and_line(X0,X4) | apart_point_and_line(X0,X3)))),
  file('GEO196+1.tptp',unknown)).
fof(f86,plain,(
  distinct_lines(sK3,sK2)),
  inference(resolution,[],[f79,f43])).
fof(f43,plain,(
  ( ! [X0] : (~convergent_lines(X0,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : ~convergent_lines(X0,X0)),
  file('GEO196+1.tptp',unknown)).
fof(f79,plain,(
  ( ! [X3] : (convergent_lines(sK2,X3) | distinct_lines(sK3,X3)) )),
  inference(resolution,[],[f54,f56])).
fof(f54,plain,(
  ( ! [X2,X0,X1] : (~convergent_lines(X0,X1) | distinct_lines(X1,X2) | convergent_lines(X0,X2)) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2] : (convergent_lines(X0,X2) | distinct_lines(X1,X2) | ~convergent_lines(X0,X1))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : ((convergent_lines(X0,X2) | distinct_lines(X1,X2)) | ~convergent_lines(X0,X1))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0,X1,X2] : (convergent_lines(X0,X1) => (convergent_lines(X0,X2) | distinct_lines(X1,X2)))),
  file('GEO196+1.tptp',unknown)).
fof(f1534,plain,(
  ( ! [X1] : (distinct_points(intersection_point(sK2,sK3),X1) | apart_point_and_line(X1,sK0)) ) | ~spl4_1),
  inference(resolution,[],[f63,f52])).
fof(f52,plain,(
  ( ! [X2,X0,X1] : (~apart_point_and_line(X0,X1) | distinct_points(X0,X2) | apart_point_and_line(X2,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : (apart_point_and_line(X2,X1) | distinct_points(X0,X2) | ~apart_point_and_line(X0,X1))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : ((apart_point_and_line(X2,X1) | distinct_points(X0,X2)) | ~apart_point_and_line(X0,X1))),
  inference(ennf_transformation,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2] : (apart_point_and_line(X0,X1) => (apart_point_and_line(X2,X1) | distinct_points(X0,X2)))),
  file('GEO196+1.tptp',unknown)).
fof(f63,plain,(
  apart_point_and_line(intersection_point(sK2,sK3),sK0) | ~spl4_1),
  inference(avatar_component_clause,[],[f61])).
fof(f1532,plain,(
  spl4_8 | spl4_9 | spl4_7 | ~spl4_2),
  inference(avatar_split_clause,[],[f1442,f65,f1444,f1451,f1447])).
fof(f65,plain,(
  spl4_2 <=> apart_point_and_line(intersection_point(sK2,sK3),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_2])])).
fof(f1442,plain,(
  ( ! [X61] : (apart_point_and_line(X61,sK2) | apart_point_and_line(X61,sK3) | apart_point_and_line(intersection_point(sK2,sK3),sK2) | apart_point_and_line(intersection_point(sK2,sK3),sK3) | apart_point_and_line(X61,sK1)) ) | ~spl4_2),
  inference(resolution,[],[f89,f217])).
fof(f217,plain,(
  ( ! [X0] : (distinct_points(X0,intersection_point(sK2,sK3)) | apart_point_and_line(X0,sK1)) ) | ~spl4_2),
  inference(resolution,[],[f92,f41])).
fof(f41,plain,(
  ( ! [X0] : (~distinct_points(X0,X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0] : ~distinct_points(X0,X0)),
  file('GEO196+1.tptp',unknown)).
fof(f92,plain,(
  ( ! [X0,X1] : (apart_point_and_line(X0,sK1) | distinct_points(intersection_point(sK2,sK3),X1) | distinct_points(X0,X1)) ) | ~spl4_2),
  inference(resolution,[],[f71,f44])).
fof(f44,plain,(
  ( ! [X2,X0,X1] : (~distinct_points(X0,X1) | distinct_points(X0,X2) | distinct_points(X1,X2)) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2] : (distinct_points(X1,X2) | distinct_points(X0,X2) | ~distinct_points(X0,X1))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : ((distinct_points(X1,X2) | distinct_points(X0,X2)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2] : (distinct_points(X0,X1) => (distinct_points(X1,X2) | distinct_points(X0,X2)))),
  file('GEO196+1.tptp',unknown)).
fof(f71,plain,(
  ( ! [X0] : (distinct_points(intersection_point(sK2,sK3),X0) | apart_point_and_line(X0,sK1)) ) | ~spl4_2),
  inference(resolution,[],[f52,f67])).
fof(f67,plain,(
  apart_point_and_line(intersection_point(sK2,sK3),sK1) | ~spl4_2),
  inference(avatar_component_clause,[],[f65])).
fof(f68,plain,(
  spl4_1 | spl4_2),
  inference(avatar_split_clause,[],[f59,f65,f61])).
fof(f59,plain,(
  apart_point_and_line(intersection_point(sK2,sK3),sK1) | apart_point_and_line(intersection_point(sK2,sK3),sK0)),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for GEO196+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.085 s
% ------------------------------
% ------------------------------
