% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP011-4
% SZS output start Proof for GRP011-4
fof(f168,plain,(
  $false),
  inference(subsumption_resolution,[],[f167,f5])).
fof(f5,axiom,(
  b != d),
  file('GRP011-4.tptp',unknown)).
fof(f167,plain,(
  b = d),
  inference(forward_demodulation,[],[f166,f53])).
fof(f53,plain,(
  ( ! [X5] : (multiply(X5,identity) = X5) )),
  inference(superposition,[],[f17,f15])).
fof(f15,plain,(
  ( ! [X1] : (multiply(inverse(inverse(X1)),identity) = X1) )),
  inference(superposition,[],[f10,f3])).
fof(f3,axiom,(
  ( ! [X0] : (identity = multiply(inverse(X0),X0)) )),
  file('GRP011-4.tptp',unknown)).
fof(f10,plain,(
  ( ! [X2,X3] : (multiply(inverse(X2),multiply(X2,X3)) = X3) )),
  inference(forward_demodulation,[],[f7,f2])).
fof(f2,axiom,(
  ( ! [X0] : (multiply(identity,X0) = X0) )),
  file('GRP011-4.tptp',unknown)).
fof(f7,plain,(
  ( ! [X2,X3] : (multiply(inverse(X2),multiply(X2,X3)) = multiply(identity,X3)) )),
  inference(superposition,[],[f1,f3])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (multiply(multiply(X0,X1),X2) = multiply(X0,multiply(X1,X2))) )),
  file('GRP011-4.tptp',unknown)).
fof(f17,plain,(
  ( ! [X6,X5] : (multiply(X5,X6) = multiply(inverse(inverse(X5)),X6)) )),
  inference(superposition,[],[f10,f10])).
fof(f166,plain,(
  d = multiply(b,identity)),
  inference(forward_demodulation,[],[f155,f53])).
fof(f155,plain,(
  multiply(b,identity) = multiply(d,identity)),
  inference(superposition,[],[f11,f51])).
fof(f51,plain,(
  ( ! [X0] : (identity = multiply(X0,inverse(X0))) )),
  inference(superposition,[],[f17,f3])).
fof(f11,plain,(
  ( ! [X4] : (multiply(d,multiply(c,X4)) = multiply(b,multiply(c,X4))) )),
  inference(forward_demodulation,[],[f8,f1])).
fof(f8,plain,(
  ( ! [X4] : (multiply(d,multiply(c,X4)) = multiply(multiply(b,c),X4)) )),
  inference(superposition,[],[f1,f4])).
fof(f4,axiom,(
  multiply(b,c) = multiply(d,c)),
  file('GRP011-4.tptp',unknown)).
% SZS output end Proof for GRP011-4
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
