% Refutation found. Thanks to Tanya!
% SZS status Theorem for GEO178+3
% SZS output start Proof for GEO178+3
fof(f157,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f140,f154,f156])).
fof(f156,plain,(
  spl3_2),
  inference(avatar_split_clause,[],[f155,f137])).
fof(f137,plain,(
  spl3_2 <=> distinct_points(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl3_2])])).
fof(f155,plain,(
  distinct_points(sK2,sK1)),
  inference(subsumption_resolution,[],[f148,f129])).
fof(f129,plain,(
  distinct_points(sK0,sK1)),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  (~distinct_points(sK2,sK1) | ~distinct_points(sK2,sK0)) & apart_point_and_line(sK2,line_connecting(sK0,sK1)) & distinct_points(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f90,f91])).
fof(f91,plain,(
  ? [X0,X1,X2] : ((~distinct_points(X2,X1) | ~distinct_points(X2,X0)) & apart_point_and_line(X2,line_connecting(X0,X1)) & distinct_points(X0,X1)) => ((~distinct_points(sK2,sK1) | ~distinct_points(sK2,sK0)) & apart_point_and_line(sK2,line_connecting(sK0,sK1)) & distinct_points(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f90,plain,(
  ? [X0,X1,X2] : ((~distinct_points(X2,X1) | ~distinct_points(X2,X0)) & apart_point_and_line(X2,line_connecting(X0,X1)) & distinct_points(X0,X1))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ? [X0,X1,X2] : ((~distinct_points(X2,X1) | ~distinct_points(X2,X0)) & (apart_point_and_line(X2,line_connecting(X0,X1)) & distinct_points(X0,X1)))),
  inference(ennf_transformation,[],[f37])).
fof(f37,negated_conjecture,(
  ~! [X0,X1,X2] : ((apart_point_and_line(X2,line_connecting(X0,X1)) & distinct_points(X0,X1)) => (distinct_points(X2,X1) & distinct_points(X2,X0)))),
  inference(negated_conjecture,[],[f36])).
fof(f36,conjecture,(
  ! [X0,X1,X2] : ((apart_point_and_line(X2,line_connecting(X0,X1)) & distinct_points(X0,X1)) => (distinct_points(X2,X1) & distinct_points(X2,X0)))),
  file('GEO178+3.tptp',unknown)).
fof(f148,plain,(
  distinct_points(sK2,sK1) | ~distinct_points(sK0,sK1)),
  inference(resolution,[],[f146,f100])).
fof(f100,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X1,line_connecting(X0,X1)) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : (~apart_point_and_line(X1,line_connecting(X0,X1)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1] : (distinct_points(X0,X1) => ~apart_point_and_line(X1,line_connecting(X0,X1)))),
  file('GEO178+3.tptp',unknown)).
fof(f146,plain,(
  ( ! [X0] : (apart_point_and_line(X0,line_connecting(sK0,sK1)) | distinct_points(sK2,X0)) )),
  inference(resolution,[],[f104,f130])).
fof(f130,plain,(
  apart_point_and_line(sK2,line_connecting(sK0,sK1))),
  inference(cnf_transformation,[],[f92])).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (~apart_point_and_line(X0,X1) | distinct_points(X0,X2) | apart_point_and_line(X2,X1)) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : (apart_point_and_line(X2,X1) | distinct_points(X0,X2) | ~apart_point_and_line(X0,X1))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : ((apart_point_and_line(X2,X1) | distinct_points(X0,X2)) | ~apart_point_and_line(X0,X1))),
  inference(ennf_transformation,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2] : (apart_point_and_line(X0,X1) => (apart_point_and_line(X2,X1) | distinct_points(X0,X2)))),
  file('GEO178+3.tptp',unknown)).
fof(f154,plain,(
  spl3_1),
  inference(avatar_contradiction_clause,[],[f153])).
fof(f153,plain,(
  $false | spl3_1),
  inference(subsumption_resolution,[],[f152,f129])).
fof(f152,plain,(
  ~distinct_points(sK0,sK1) | spl3_1),
  inference(subsumption_resolution,[],[f147,f135])).
fof(f135,plain,(
  ~distinct_points(sK2,sK0) | spl3_1),
  inference(avatar_component_clause,[],[f133])).
fof(f133,plain,(
  spl3_1 <=> distinct_points(sK2,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl3_1])])).
fof(f147,plain,(
  distinct_points(sK2,sK0) | ~distinct_points(sK0,sK1)),
  inference(resolution,[],[f146,f99])).
fof(f99,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X0,line_connecting(X0,X1)) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : (~apart_point_and_line(X0,line_connecting(X0,X1)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1] : (distinct_points(X0,X1) => ~apart_point_and_line(X0,line_connecting(X0,X1)))),
  file('GEO178+3.tptp',unknown)).
fof(f140,plain,(
  ~spl3_1 | ~spl3_2),
  inference(avatar_split_clause,[],[f131,f137,f133])).
fof(f131,plain,(
  ~distinct_points(sK2,sK1) | ~distinct_points(sK2,sK0)),
  inference(cnf_transformation,[],[f92])).
% SZS output end Proof for GEO178+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
