% Refutation found. Thanks to Tanya!
% SZS status Theorem for GEO223+2
% SZS output start Proof for GEO223+2
fof(f432,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f169,f173,f200,f232,f428])).
fof(f428,plain,(
  ~spl2_6),
  inference(avatar_contradiction_clause,[],[f427])).
fof(f427,plain,(
  $false | ~spl2_6),
  inference(subsumption_resolution,[],[f423,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (~convergent_lines(parallel_through_point(X1,X0),X1)) )),
  inference(cnf_transformation,[],[f13])).
fof(f13,axiom,(
  ! [X0,X1] : ~convergent_lines(parallel_through_point(X1,X0),X1)),
  file('GEO223+2.tptp',unknown)).
fof(f423,plain,(
  convergent_lines(parallel_through_point(sK1,sK0),sK1) | ~spl2_6),
  inference(resolution,[],[f199,f78])).
fof(f78,plain,(
  ( ! [X0,X1] : (~unorthogonal_lines(orthogonal_through_point(X1,X0),X1)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : ~unorthogonal_lines(orthogonal_through_point(X1,X0),X1)),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8,X5] : ~unorthogonal_lines(orthogonal_through_point(X5,X8),X5)),
  file('GEO223+2.tptp',unknown)).
fof(f199,plain,(
  ( ! [X1] : (unorthogonal_lines(orthogonal_through_point(sK1,sK0),X1) | convergent_lines(parallel_through_point(sK1,sK0),X1)) ) | ~spl2_6),
  inference(avatar_component_clause,[],[f198])).
fof(f198,plain,(
  spl2_6 <=> ! [X1] : (convergent_lines(parallel_through_point(sK1,sK0),X1) | unorthogonal_lines(orthogonal_through_point(sK1,sK0),X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_6])])).
fof(f232,plain,(
  spl2_4),
  inference(avatar_contradiction_clause,[],[f231])).
fof(f231,plain,(
  $false | spl2_4),
  inference(subsumption_resolution,[],[f226,f58])).
fof(f58,plain,(
  ( ! [X0] : (~convergent_lines(X0,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : ~convergent_lines(X0,X0)),
  file('GEO223+2.tptp',unknown)).
fof(f226,plain,(
  convergent_lines(sK1,sK1) | spl2_4),
  inference(resolution,[],[f208,f70])).
fof(f208,plain,(
  ( ! [X0] : (convergent_lines(parallel_through_point(sK1,sK0),X0) | convergent_lines(sK1,X0)) ) | spl2_4),
  inference(resolution,[],[f205,f61])).
fof(f61,plain,(
  ( ! [X2,X0,X1] : (~convergent_lines(X0,X1) | convergent_lines(X0,X2) | convergent_lines(X1,X2)) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : (convergent_lines(X1,X2) | convergent_lines(X0,X2) | ~convergent_lines(X0,X1))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : ((convergent_lines(X1,X2) | convergent_lines(X0,X2)) | ~convergent_lines(X0,X1))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2] : (convergent_lines(X0,X1) => (convergent_lines(X1,X2) | convergent_lines(X0,X2)))),
  file('GEO223+2.tptp',unknown)).
fof(f205,plain,(
  convergent_lines(sK1,parallel_through_point(sK1,sK0)) | spl2_4),
  inference(resolution,[],[f192,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X1] : (convergent_lines(X1,orthogonal_through_point(X0,X2)) | convergent_lines(X0,X1)) )),
  inference(resolution,[],[f91,f58])).
fof(f91,plain,(
  ( ! [X4,X2,X5,X3] : (convergent_lines(X2,X3) | convergent_lines(orthogonal_through_point(X2,X4),X5) | convergent_lines(X3,X5)) )),
  inference(resolution,[],[f88,f61])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (convergent_lines(orthogonal_through_point(X0,X1),X2) | convergent_lines(X0,X2)) )),
  inference(resolution,[],[f61,f82])).
fof(f82,plain,(
  ( ! [X0,X1] : (convergent_lines(orthogonal_through_point(X0,X1),X0)) )),
  inference(resolution,[],[f73,f78])).
fof(f73,plain,(
  ( ! [X0,X1] : (unorthogonal_lines(X0,X1) | convergent_lines(X0,X1)) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : (unorthogonal_lines(X0,X1) | convergent_lines(X0,X1))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X5,X6] : (unorthogonal_lines(X5,X6) | convergent_lines(X5,X6))),
  file('GEO223+2.tptp',unknown)).
fof(f192,plain,(
  ~convergent_lines(parallel_through_point(sK1,sK0),orthogonal_through_point(sK1,sK0)) | spl2_4),
  inference(avatar_component_clause,[],[f190])).
fof(f190,plain,(
  spl2_4 <=> convergent_lines(parallel_through_point(sK1,sK0),orthogonal_through_point(sK1,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_4])])).
fof(f200,plain,(
  ~spl2_4 | spl2_6 | ~spl2_3),
  inference(avatar_split_clause,[],[f187,f171,f198,f190])).
fof(f171,plain,(
  spl2_3 <=> ! [X0] : (unorthogonal_lines(parallel_through_point(sK1,sK0),X0) | unorthogonal_lines(orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_3])])).
fof(f187,plain,(
  ( ! [X1] : (convergent_lines(parallel_through_point(sK1,sK0),X1) | unorthogonal_lines(orthogonal_through_point(sK1,sK0),X1) | ~convergent_lines(parallel_through_point(sK1,sK0),orthogonal_through_point(sK1,sK0))) ) | ~spl2_3),
  inference(resolution,[],[f182,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X1] : (~unorthogonal_lines(X0,X1) | convergent_lines(X0,X2) | unorthogonal_lines(X1,X2) | ~convergent_lines(X0,X1)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((unorthogonal_lines(X1,X2) & convergent_lines(X1,X2)) | (unorthogonal_lines(X0,X2) & convergent_lines(X0,X2)) | ~unorthogonal_lines(X0,X1) | ~convergent_lines(X0,X1))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : (((unorthogonal_lines(X1,X2) & convergent_lines(X1,X2)) | (unorthogonal_lines(X0,X2) & convergent_lines(X0,X2))) | (~unorthogonal_lines(X0,X1) | ~convergent_lines(X0,X1)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : ((unorthogonal_lines(X0,X1) & convergent_lines(X0,X1)) => ((unorthogonal_lines(X1,X2) & convergent_lines(X1,X2)) | (unorthogonal_lines(X0,X2) & convergent_lines(X0,X2))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X5,X6,X7] : ((unorthogonal_lines(X5,X6) & convergent_lines(X5,X6)) => ((unorthogonal_lines(X6,X7) & convergent_lines(X6,X7)) | (unorthogonal_lines(X5,X7) & convergent_lines(X5,X7))))),
  file('GEO223+2.tptp',unknown)).
fof(f182,plain,(
  unorthogonal_lines(parallel_through_point(sK1,sK0),orthogonal_through_point(sK1,sK0)) | ~spl2_3),
  inference(resolution,[],[f172,f78])).
fof(f172,plain,(
  ( ! [X0] : (unorthogonal_lines(orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0),X0) | unorthogonal_lines(parallel_through_point(sK1,sK0),X0)) ) | ~spl2_3),
  inference(avatar_component_clause,[],[f171])).
fof(f173,plain,(
  spl2_2 | spl2_3),
  inference(avatar_split_clause,[],[f136,f171,f111])).
fof(f111,plain,(
  spl2_2 <=> ! [X0] : (apart_point_and_line(X0,orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0)) | apart_point_and_line(X0,parallel_through_point(sK1,sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_2])])).
fof(f136,plain,(
  ( ! [X0,X1] : (unorthogonal_lines(parallel_through_point(sK1,sK0),X0) | apart_point_and_line(X1,orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0)) | apart_point_and_line(X1,parallel_through_point(sK1,sK0)) | unorthogonal_lines(orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0),X0)) )),
  inference(resolution,[],[f80,f81])).
fof(f81,plain,(
  distinct_lines(parallel_through_point(sK1,sK0),orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0))),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  distinct_lines(parallel_through_point(sK1,sK0),orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f53,f54])).
fof(f54,plain,(
  ? [X0,X1] : distinct_lines(parallel_through_point(X1,X0),orthogonal_through_point(orthogonal_through_point(X1,X0),X0)) => distinct_lines(parallel_through_point(sK1,sK0),orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0))),
  introduced(choice_axiom,[])).
fof(f53,plain,(
  ? [X0,X1] : distinct_lines(parallel_through_point(X1,X0),orthogonal_through_point(orthogonal_through_point(X1,X0),X0))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1] : ~distinct_lines(parallel_through_point(X1,X0),orthogonal_through_point(orthogonal_through_point(X1,X0),X0))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X8,X5] : ~distinct_lines(parallel_through_point(X5,X8),orthogonal_through_point(orthogonal_through_point(X5,X8),X8))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X8,X5] : ~distinct_lines(parallel_through_point(X5,X8),orthogonal_through_point(orthogonal_through_point(X5,X8),X8))),
  file('GEO223+2.tptp',unknown)).
fof(f80,plain,(
  ( ! [X2,X0,X3,X1] : (~distinct_lines(X1,X2) | unorthogonal_lines(X1,X3) | apart_point_and_line(X0,X2) | apart_point_and_line(X0,X1) | unorthogonal_lines(X2,X3)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3] : (unorthogonal_lines(X2,X3) | unorthogonal_lines(X1,X3) | apart_point_and_line(X0,X2) | apart_point_and_line(X0,X1) | ~distinct_lines(X1,X2))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : ((unorthogonal_lines(X2,X3) | unorthogonal_lines(X1,X3) | apart_point_and_line(X0,X2) | apart_point_and_line(X0,X1)) | ~distinct_lines(X1,X2))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : (distinct_lines(X1,X2) => (unorthogonal_lines(X2,X3) | unorthogonal_lines(X1,X3) | apart_point_and_line(X0,X2) | apart_point_and_line(X0,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8,X5,X6,X7] : (distinct_lines(X5,X6) => (unorthogonal_lines(X6,X7) | unorthogonal_lines(X5,X7) | apart_point_and_line(X8,X6) | apart_point_and_line(X8,X5)))),
  file('GEO223+2.tptp',unknown)).
fof(f169,plain,(
  ~spl2_2),
  inference(avatar_contradiction_clause,[],[f168])).
fof(f168,plain,(
  $false | ~spl2_2),
  inference(subsumption_resolution,[],[f163,f71])).
fof(f71,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X0,parallel_through_point(X1,X0))) )),
  inference(cnf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0,X1] : ~apart_point_and_line(X0,parallel_through_point(X1,X0))),
  file('GEO223+2.tptp',unknown)).
fof(f163,plain,(
  apart_point_and_line(sK0,parallel_through_point(sK1,sK0)) | ~spl2_2),
  inference(resolution,[],[f112,f79])).
fof(f79,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X0,orthogonal_through_point(X1,X0))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : ~apart_point_and_line(X0,orthogonal_through_point(X1,X0))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X8,X5] : ~apart_point_and_line(X8,orthogonal_through_point(X5,X8))),
  file('GEO223+2.tptp',unknown)).
fof(f112,plain,(
  ( ! [X0] : (apart_point_and_line(X0,orthogonal_through_point(orthogonal_through_point(sK1,sK0),sK0)) | apart_point_and_line(X0,parallel_through_point(sK1,sK0))) ) | ~spl2_2),
  inference(avatar_component_clause,[],[f111])).
% SZS output end Proof for GEO223+2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
