% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GEO067-3
% SZS output start Proof for GEO067-3
fof(f163,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f125,f138,f155,f157,f159,f162])).
fof(f162,plain,(
  ~spl0_2 | spl0_5),
  inference(avatar_contradiction_clause,[],[f161])).
fof(f161,plain,(
  $false | (~spl0_2 | spl0_5)),
  inference(subsumption_resolution,[],[f160,f148])).
fof(f148,plain,(
  ( ! [X0,X1] : (colinear(X0,X1,X0)) )),
  inference(resolution,[],[f20,f49])).
fof(f49,axiom,(
  ( ! [X10,X3] : (between(X10,X3,X3)) )),
  file('GEO067-3.tptp',unknown)).
fof(f20,axiom,(
  ( ! [X2,X0,X1] : (~between(X1,X2,X0) | colinear(X0,X1,X2)) )),
  file('GEO067-3.tptp',unknown)).
fof(f160,plain,(
  ~colinear(x,z,x) | (~spl0_2 | spl0_5)),
  inference(forward_demodulation,[],[f137,f124])).
fof(f124,plain,(
  x = y | ~spl0_2),
  inference(avatar_component_clause,[],[f122])).
fof(f122,plain,(
  spl0_2 <=> x = y),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f137,plain,(
  ~colinear(x,z,y) | spl0_5),
  inference(avatar_component_clause,[],[f135])).
fof(f135,plain,(
  spl0_5 <=> colinear(x,z,y)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f159,plain,(
  spl0_4),
  inference(avatar_contradiction_clause,[],[f158])).
fof(f158,plain,(
  $false | spl0_4),
  inference(subsumption_resolution,[],[f133,f145])).
fof(f145,plain,(
  ( ! [X0,X1] : (colinear(X0,X1,X1)) )),
  inference(resolution,[],[f19,f49])).
fof(f19,axiom,(
  ( ! [X2,X0,X1] : (~between(X0,X1,X2) | colinear(X0,X1,X2)) )),
  file('GEO067-3.tptp',unknown)).
fof(f133,plain,(
  ~colinear(y,x,x) | spl0_4),
  inference(avatar_component_clause,[],[f131])).
fof(f131,plain,(
  spl0_4 <=> colinear(y,x,x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f157,plain,(
  spl0_3),
  inference(avatar_contradiction_clause,[],[f156])).
fof(f156,plain,(
  $false | spl0_3),
  inference(subsumption_resolution,[],[f129,f148])).
fof(f129,plain,(
  ~colinear(x,y,x) | spl0_3),
  inference(avatar_component_clause,[],[f127])).
fof(f127,plain,(
  spl0_3 <=> colinear(x,y,x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f155,plain,(
  spl0_1),
  inference(avatar_contradiction_clause,[],[f154])).
fof(f154,plain,(
  $false | spl0_1),
  inference(resolution,[],[f146,f120])).
fof(f120,plain,(
  ~colinear(x,x,y) | spl0_1),
  inference(avatar_component_clause,[],[f118])).
fof(f118,plain,(
  spl0_1 <=> colinear(x,x,y)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f146,plain,(
  ( ! [X2,X3] : (colinear(X2,X2,X3)) )),
  inference(resolution,[],[f19,f52])).
fof(f52,axiom,(
  ( ! [X10,X3] : (between(X10,X10,X3)) )),
  file('GEO067-3.tptp',unknown)).
fof(f138,plain,(
  ~spl0_1 | ~spl0_3 | ~spl0_4 | ~spl0_5),
  inference(avatar_split_clause,[],[f98,f135,f131,f127,f118])).
fof(f98,axiom,(
  ~colinear(x,z,y) | ~colinear(y,x,x) | ~colinear(x,y,x) | ~colinear(x,x,y)),
  file('GEO067-3.tptp',unknown)).
fof(f125,plain,(
  ~spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f116,f122,f118])).
fof(f116,plain,(
  x = y | ~colinear(x,x,y)),
  inference(subsumption_resolution,[],[f115,f92])).
fof(f92,axiom,(
  ( ! [X10,X5,X3] : (colinear(X3,X5,X10) | ~colinear(X10,X3,X5)) )),
  file('GEO067-3.tptp',unknown)).
fof(f115,plain,(
  x = y | ~colinear(x,y,x) | ~colinear(x,x,y)),
  inference(subsumption_resolution,[],[f97,f91])).
fof(f91,axiom,(
  ( ! [X10,X5,X3] : (colinear(X5,X3,X10) | ~colinear(X10,X3,X5)) )),
  file('GEO067-3.tptp',unknown)).
fof(f97,axiom,(
  x = y | ~colinear(y,x,x) | ~colinear(x,y,x) | ~colinear(x,x,y)),
  file('GEO067-3.tptp',unknown)).
% SZS output end Proof for GEO067-3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
