% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP039-3
% SZS output start Proof for GRP039-3
fof(f2892,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f52,f86,f447,f578,f2816,f2865,f2891])).
fof(f2891,plain,(
  spl0_3 | spl0_8 | ~spl0_7),
  inference(avatar_split_clause,[],[f2890,f530,f534,f437])).
fof(f437,plain,(
  spl0_3 <=> subgroup_member(inverse(a))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f534,plain,(
  spl0_8 <=> ! [X21] : (~product(X21,b,identity) | ~subgroup_member(X21))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f530,plain,(
  spl0_7 <=> subgroup_member(inverse(c))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f2890,plain,(
  ( ! [X21] : (~product(X21,b,identity) | subgroup_member(inverse(a)) | ~subgroup_member(X21)) ) | ~spl0_7),
  inference(subsumption_resolution,[],[f2428,f531])).
fof(f531,plain,(
  subgroup_member(inverse(c)) | ~spl0_7),
  inference(avatar_component_clause,[],[f530])).
fof(f2428,plain,(
  ( ! [X21] : (~product(X21,b,identity) | subgroup_member(inverse(a)) | ~subgroup_member(inverse(c)) | ~subgroup_member(X21)) )),
  inference(resolution,[],[f258,f41])).
fof(f41,plain,(
  ( ! [X2,X0,X1] : (~product(X1,X0,X2) | subgroup_member(X2) | ~subgroup_member(inverse(X0)) | ~subgroup_member(X1)) )),
  inference(superposition,[],[f9,f12])).
fof(f12,axiom,(
  ( ! [X6] : (inverse(inverse(X6)) = X6) )),
  file('GRP039-3.tptp',unknown)).
fof(f9,axiom,(
  ( ! [X6,X8,X7] : (~product(X6,inverse(X7),X8) | subgroup_member(X8) | ~subgroup_member(X7) | ~subgroup_member(X6)) )),
  file('GRP039-3.tptp',unknown)).
fof(f258,plain,(
  ( ! [X14] : (product(X14,c,inverse(a)) | ~product(X14,b,identity)) )),
  inference(resolution,[],[f78,f18])).
fof(f18,axiom,(
  product(b,inverse(a),c)),
  file('GRP039-3.tptp',unknown)).
fof(f78,plain,(
  ( ! [X21,X19,X20,X18] : (~product(X21,X20,X19) | product(X18,X19,X20) | ~product(X18,X21,identity)) )),
  inference(resolution,[],[f7,f1])).
fof(f1,axiom,(
  ( ! [X0] : (product(identity,X0,X0)) )),
  file('GRP039-3.tptp',unknown)).
fof(f7,axiom,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~product(X4,X2,X3) | product(X0,X5,X3) | ~product(X1,X2,X5) | ~product(X0,X1,X4)) )),
  file('GRP039-3.tptp',unknown)).
fof(f2865,plain,(
  ~spl0_1 | ~spl0_7),
  inference(avatar_split_clause,[],[f2864,f530,f45])).
fof(f45,plain,(
  spl0_1 <=> subgroup_member(a)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f2864,plain,(
  ~subgroup_member(a) | ~spl0_7),
  inference(subsumption_resolution,[],[f2863,f531])).
fof(f2863,plain,(
  ~subgroup_member(inverse(c)) | ~subgroup_member(a)),
  inference(subsumption_resolution,[],[f224,f20])).
fof(f20,axiom,(
  ~subgroup_member(d)),
  file('GRP039-3.tptp',unknown)).
fof(f224,plain,(
  subgroup_member(d) | ~subgroup_member(inverse(c)) | ~subgroup_member(a)),
  inference(resolution,[],[f41,f19])).
fof(f19,axiom,(
  product(a,c,d)),
  file('GRP039-3.tptp',unknown)).
fof(f2816,plain,(
  ~spl0_8),
  inference(avatar_contradiction_clause,[],[f2815])).
fof(f2815,plain,(
  $false | ~spl0_8),
  inference(subsumption_resolution,[],[f2813,f17])).
fof(f17,axiom,(
  subgroup_member(b)),
  file('GRP039-3.tptp',unknown)).
fof(f2813,plain,(
  ~subgroup_member(b) | ~spl0_8),
  inference(resolution,[],[f2749,f14])).
fof(f14,axiom,(
  ( ! [X6] : (subgroup_member(inverse(X6)) | ~subgroup_member(X6)) )),
  file('GRP039-3.tptp',unknown)).
fof(f2749,plain,(
  ~subgroup_member(inverse(b)) | ~spl0_8),
  inference(resolution,[],[f535,f3])).
fof(f3,axiom,(
  ( ! [X0] : (product(inverse(X0),X0,identity)) )),
  file('GRP039-3.tptp',unknown)).
fof(f535,plain,(
  ( ! [X21] : (~product(X21,b,identity) | ~subgroup_member(X21)) ) | ~spl0_8),
  inference(avatar_component_clause,[],[f534])).
fof(f578,plain,(
  ~spl0_2 | spl0_7),
  inference(avatar_contradiction_clause,[],[f577])).
fof(f577,plain,(
  $false | (~spl0_2 | spl0_7)),
  inference(subsumption_resolution,[],[f575,f51])).
fof(f51,plain,(
  subgroup_member(c) | ~spl0_2),
  inference(avatar_component_clause,[],[f49])).
fof(f49,plain,(
  spl0_2 <=> subgroup_member(c)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f575,plain,(
  ~subgroup_member(c) | spl0_7),
  inference(resolution,[],[f532,f14])).
fof(f532,plain,(
  ~subgroup_member(inverse(c)) | spl0_7),
  inference(avatar_component_clause,[],[f530])).
fof(f447,plain,(
  spl0_1 | ~spl0_3),
  inference(avatar_contradiction_clause,[],[f446])).
fof(f446,plain,(
  $false | (spl0_1 | ~spl0_3)),
  inference(subsumption_resolution,[],[f445,f47])).
fof(f47,plain,(
  ~subgroup_member(a) | spl0_1),
  inference(avatar_component_clause,[],[f45])).
fof(f445,plain,(
  subgroup_member(a) | ~spl0_3),
  inference(resolution,[],[f439,f23])).
fof(f23,plain,(
  ( ! [X0] : (~subgroup_member(inverse(X0)) | subgroup_member(X0)) )),
  inference(superposition,[],[f14,f12])).
fof(f439,plain,(
  subgroup_member(inverse(a)) | ~spl0_3),
  inference(avatar_component_clause,[],[f437])).
fof(f86,plain,(
  spl0_2 | spl0_1),
  inference(avatar_split_clause,[],[f85,f45,f49])).
fof(f85,plain,(
  subgroup_member(c) | spl0_1),
  inference(subsumption_resolution,[],[f84,f47])).
fof(f84,plain,(
  subgroup_member(c) | subgroup_member(a) | spl0_1),
  inference(subsumption_resolution,[],[f83,f20])).
fof(f83,plain,(
  subgroup_member(c) | subgroup_member(d) | subgroup_member(a) | spl0_1),
  inference(superposition,[],[f15,f73])).
fof(f73,plain,(
  c = element_in_O2(a,d) | spl0_1),
  inference(subsumption_resolution,[],[f72,f47])).
fof(f72,plain,(
  c = element_in_O2(a,d) | subgroup_member(a)),
  inference(subsumption_resolution,[],[f71,f20])).
fof(f71,plain,(
  c = element_in_O2(a,d) | subgroup_member(d) | subgroup_member(a)),
  inference(resolution,[],[f60,f16])).
fof(f16,axiom,(
  ( ! [X6,X7] : (product(X6,element_in_O2(X6,X7),X7) | subgroup_member(X7) | subgroup_member(X6)) )),
  file('GRP039-3.tptp',unknown)).
fof(f60,plain,(
  ( ! [X15] : (~product(a,X15,d) | c = X15) )),
  inference(resolution,[],[f10,f19])).
fof(f10,axiom,(
  ( ! [X6,X8,X7,X9] : (~product(X6,X9,X8) | X7 = X9 | ~product(X6,X7,X8)) )),
  file('GRP039-3.tptp',unknown)).
fof(f15,axiom,(
  ( ! [X6,X7] : (subgroup_member(element_in_O2(X6,X7)) | subgroup_member(X7) | subgroup_member(X6)) )),
  file('GRP039-3.tptp',unknown)).
fof(f52,plain,(
  ~spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f43,f49,f45])).
fof(f43,plain,(
  subgroup_member(c) | ~subgroup_member(a)),
  inference(subsumption_resolution,[],[f40,f17])).
fof(f40,plain,(
  subgroup_member(c) | ~subgroup_member(a) | ~subgroup_member(b)),
  inference(resolution,[],[f9,f18])).
% SZS output end Proof for GRP039-3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5884
% Time elapsed: 0.056 s
% ------------------------------
% ------------------------------
