% Refutation found. Thanks to Tanya!
% SZS status Theorem for KRS136+1
% SZS output start Proof for KRS136+1
fof(f26,plain,(
  $false),
  inference(subsumption_resolution,[],[f24,f20])).
fof(f20,plain,(
  ~xsd_string(sK0)),
  inference(subsumption_resolution,[],[f19,f14])).
fof(f14,plain,(
  ( ! [X0] : (~xsd_integer(X0) | ~xsd_string(X0)) )),
  inference(cnf_transformation,[],[f7])).
fof(f7,plain,(
  ! [X0] : ((xsd_string(X0) | xsd_integer(X0)) & (~xsd_integer(X0) | ~xsd_string(X0)))),
  inference(nnf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (xsd_string(X0) <=> ~xsd_integer(X0))),
  file('KRS136+1.tptp',unknown)).
fof(f19,plain,(
  xsd_integer(sK0) | ~xsd_string(sK0)),
  inference(subsumption_resolution,[],[f18,f12])).
fof(f12,plain,(
  ( ! [X0] : (cowlThing(X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0] : (~cowlNothing(X0) & cowlThing(X0))),
  file('KRS136+1.tptp',unknown)).
fof(f18,plain,(
  xsd_integer(sK0) | ~xsd_string(sK0) | ~cowlThing(sK1)),
  inference(subsumption_resolution,[],[f17,f13])).
fof(f13,plain,(
  ( ! [X0] : (~cowlNothing(X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f17,plain,(
  xsd_integer(sK0) | ~xsd_string(sK0) | cowlNothing(sK1) | ~cowlThing(sK1)),
  inference(cnf_transformation,[],[f11])).
fof(f11,plain,(
  ((xsd_integer(sK0) | ~xsd_string(sK0)) & (~xsd_integer(sK0) | xsd_string(sK0))) | (cowlNothing(sK1) | ~cowlThing(sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f8,f10,f9])).
fof(f9,plain,(
  ? [X0] : ((xsd_integer(X0) | ~xsd_string(X0)) & (~xsd_integer(X0) | xsd_string(X0))) => ((xsd_integer(sK0) | ~xsd_string(sK0)) & (~xsd_integer(sK0) | xsd_string(sK0)))),
  introduced(choice_axiom,[])).
fof(f10,plain,(
  ? [X1] : (cowlNothing(X1) | ~cowlThing(X1)) => (cowlNothing(sK1) | ~cowlThing(sK1))),
  introduced(choice_axiom,[])).
fof(f8,plain,(
  ? [X0] : ((xsd_integer(X0) | ~xsd_string(X0)) & (~xsd_integer(X0) | xsd_string(X0))) | ? [X1] : (cowlNothing(X1) | ~cowlThing(X1))),
  inference(nnf_transformation,[],[f6])).
fof(f6,plain,(
  ? [X0] : (xsd_string(X0) <~> ~xsd_integer(X0)) | ? [X1] : (cowlNothing(X1) | ~cowlThing(X1))),
  inference(ennf_transformation,[],[f5])).
fof(f5,plain,(
  ~(! [X0] : (xsd_string(X0) <=> ~xsd_integer(X0)) & ! [X1] : (~cowlNothing(X1) & cowlThing(X1)))),
  inference(rectify,[],[f4])).
fof(f4,negated_conjecture,(
  ~(! [X0] : (xsd_string(X0) <=> ~xsd_integer(X0)) & ! [X0] : (~cowlNothing(X0) & cowlThing(X0)))),
  inference(negated_conjecture,[],[f3])).
fof(f3,conjecture,(
  ! [X0] : (xsd_string(X0) <=> ~xsd_integer(X0)) & ! [X0] : (~cowlNothing(X0) & cowlThing(X0))),
  file('KRS136+1.tptp',unknown)).
fof(f24,plain,(
  xsd_string(sK0)),
  inference(resolution,[],[f15,f23])).
fof(f23,plain,(
  ~xsd_integer(sK0)),
  inference(subsumption_resolution,[],[f22,f12])).
fof(f22,plain,(
  ~xsd_integer(sK0) | ~cowlThing(sK1)),
  inference(subsumption_resolution,[],[f21,f13])).
fof(f21,plain,(
  ~xsd_integer(sK0) | cowlNothing(sK1) | ~cowlThing(sK1)),
  inference(subsumption_resolution,[],[f16,f20])).
fof(f16,plain,(
  ~xsd_integer(sK0) | xsd_string(sK0) | cowlNothing(sK1) | ~cowlThing(sK1)),
  inference(cnf_transformation,[],[f11])).
fof(f15,plain,(
  ( ! [X0] : (xsd_integer(X0) | xsd_string(X0)) )),
  inference(cnf_transformation,[],[f7])).
% SZS output end Proof for KRS136+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
