% Refutation found. Thanks to Tanya!
% SZS status Theorem for LAT386+4
% SZS output start Proof for LAT386+4
fof(f48965,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f277,f282,f292,f297,f1592,f4546,f4907,f23204,f23342,f48943])).
fof(f48943,plain,(
  spl23_6 | ~spl23_7 | ~spl23_80 | ~spl23_403),
  inference(avatar_contradiction_clause,[],[f48942])).
fof(f48942,plain,(
  $false | (spl23_6 | ~spl23_7 | ~spl23_80 | ~spl23_403)),
  inference(subsumption_resolution,[],[f48941,f1442])).
fof(f1442,plain,(
  sdtlseqdt0(sK22,xp) | ~spl23_80),
  inference(avatar_component_clause,[],[f1440])).
fof(f1440,plain,(
  spl23_80 <=> sdtlseqdt0(sK22,xp)),
  introduced(avatar_definition,[new_symbols(naming,[spl23_80])])).
fof(f48941,plain,(
  ~sdtlseqdt0(sK22,xp) | (spl23_6 | ~spl23_7 | ~spl23_403)),
  inference(subsumption_resolution,[],[f48917,f240])).
fof(f240,plain,(
  aElementOf0(xp,xU)),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  aInfimumOfIn0(xp,xP,xU) & ! [X0] : (sdtlseqdt0(X0,xp) | (~aLowerBoundOfIn0(X0,xP,xU) & ((~sdtlseqdt0(X0,sK20(X0)) & aElementOf0(sK20(X0),xP)) | ~aElementOf0(X0,xU)))) & aLowerBoundOfIn0(xp,xP,xU) & ! [X2] : (sdtlseqdt0(xp,X2) | ~aElementOf0(X2,xP)) & aElementOf0(xp,xU) & aElementOf0(xp,xU)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f71,f136])).
fof(f136,plain,(
  ! [X0] : (? [X1] : (~sdtlseqdt0(X0,X1) & aElementOf0(X1,xP)) => (~sdtlseqdt0(X0,sK20(X0)) & aElementOf0(sK20(X0),xP)))),
  introduced(choice_axiom,[])).
fof(f71,plain,(
  aInfimumOfIn0(xp,xP,xU) & ! [X0] : (sdtlseqdt0(X0,xp) | (~aLowerBoundOfIn0(X0,xP,xU) & (? [X1] : (~sdtlseqdt0(X0,X1) & aElementOf0(X1,xP)) | ~aElementOf0(X0,xU)))) & aLowerBoundOfIn0(xp,xP,xU) & ! [X2] : (sdtlseqdt0(xp,X2) | ~aElementOf0(X2,xP)) & aElementOf0(xp,xU) & aElementOf0(xp,xU)),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  aInfimumOfIn0(xp,xP,xU) & ! [X0] : ((aLowerBoundOfIn0(X0,xP,xU) | (! [X1] : (aElementOf0(X1,xP) => sdtlseqdt0(X0,X1)) & aElementOf0(X0,xU))) => sdtlseqdt0(X0,xp)) & aLowerBoundOfIn0(xp,xP,xU) & ! [X2] : (aElementOf0(X2,xP) => sdtlseqdt0(xp,X2)) & aElementOf0(xp,xU) & aElementOf0(xp,xU)),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  aInfimumOfIn0(xp,xP,xU) & ! [X0] : ((aLowerBoundOfIn0(X0,xP,xU) | (! [X1] : (aElementOf0(X1,xP) => sdtlseqdt0(X0,X1)) & aElementOf0(X0,xU))) => sdtlseqdt0(X0,xp)) & aLowerBoundOfIn0(xp,xP,xU) & ! [X0] : (aElementOf0(X0,xP) => sdtlseqdt0(xp,X0)) & aElementOf0(xp,xU) & aElementOf0(xp,xU)),
  file('LAT386+4.tptp',unknown)).
fof(f48917,plain,(
  ~aElementOf0(xp,xU) | ~sdtlseqdt0(sK22,xp) | (spl23_6 | ~spl23_7 | ~spl23_403)),
  inference(resolution,[],[f23898,f291])).
fof(f291,plain,(
  ~sdtlseqdt0(sK22,sdtlpdtrp0(xf,xp)) | spl23_6),
  inference(avatar_component_clause,[],[f289])).
fof(f289,plain,(
  spl23_6 <=> sdtlseqdt0(sK22,sdtlpdtrp0(xf,xp))),
  introduced(avatar_definition,[new_symbols(naming,[spl23_6])])).
fof(f23898,plain,(
  ( ! [X0] : (sdtlseqdt0(sK22,sdtlpdtrp0(xf,X0)) | ~aElementOf0(X0,xU) | ~sdtlseqdt0(sK22,X0)) ) | (~spl23_7 | ~spl23_403)),
  inference(subsumption_resolution,[],[f23853,f12270])).
fof(f12270,plain,(
  aElementOf0(sK22,xU) | ~spl23_403),
  inference(avatar_component_clause,[],[f12269])).
fof(f12269,plain,(
  spl23_403 <=> aElementOf0(sK22,xU)),
  introduced(avatar_definition,[new_symbols(naming,[spl23_403])])).
fof(f23853,plain,(
  ( ! [X0] : (sdtlseqdt0(sK22,sdtlpdtrp0(xf,X0)) | ~aElementOf0(X0,xU) | ~aElementOf0(sK22,xU) | ~sdtlseqdt0(sK22,X0)) ) | ~spl23_7),
  inference(superposition,[],[f261,f23690])).
fof(f23690,plain,(
  sK22 = sdtlpdtrp0(xf,sK22) | ~spl23_7),
  inference(resolution,[],[f23317,f222])).
fof(f222,plain,(
  ( ! [X0] : (~aElementOf0(X0,xS) | sdtlpdtrp0(xf,X0) = X0) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  xS = cS1142(xf) & ! [X0] : ((aElementOf0(X0,xS) | (~aFixedPointOf0(X0,xf) & (sdtlpdtrp0(xf,X0) != X0 | ~aElementOf0(X0,szDzozmdt0(xf))))) & ((aFixedPointOf0(X0,xf) & sdtlpdtrp0(xf,X0) = X0 & aElementOf0(X0,szDzozmdt0(xf))) | ~aElementOf0(X0,xS))) & aSet0(xS)),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  xS = cS1142(xf) & ! [X0] : (((aFixedPointOf0(X0,xf) | (sdtlpdtrp0(xf,X0) = X0 & aElementOf0(X0,szDzozmdt0(xf)))) => aElementOf0(X0,xS)) & (aElementOf0(X0,xS) => (aFixedPointOf0(X0,xf) & sdtlpdtrp0(xf,X0) = X0 & aElementOf0(X0,szDzozmdt0(xf))))) & aSet0(xS)),
  file('LAT386+4.tptp',unknown)).
fof(f23317,plain,(
  aElementOf0(sK22,xS) | ~spl23_7),
  inference(resolution,[],[f296,f228])).
fof(f228,plain,(
  ( ! [X0] : (~aElementOf0(X0,xT) | aElementOf0(X0,xS)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  aSubsetOf0(xT,xS) & ! [X0] : (aElementOf0(X0,xS) | ~aElementOf0(X0,xT)) & aSet0(xT)),
  inference(ennf_transformation,[],[f26])).
fof(f26,axiom,(
  aSubsetOf0(xT,xS) & ! [X0] : (aElementOf0(X0,xT) => aElementOf0(X0,xS)) & aSet0(xT)),
  file('LAT386+4.tptp',unknown)).
fof(f296,plain,(
  aElementOf0(sK22,xT) | ~spl23_7),
  inference(avatar_component_clause,[],[f294])).
fof(f294,plain,(
  spl23_7 <=> aElementOf0(sK22,xT)),
  introduced(avatar_definition,[new_symbols(naming,[spl23_7])])).
fof(f261,plain,(
  ( ! [X0,X1] : (sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)) | ~aElementOf0(X1,xU) | ~aElementOf0(X0,xU) | ~sdtlseqdt0(X0,X1)) )),
  inference(forward_demodulation,[],[f260,f259])).
fof(f259,plain,(
  xU = szDzozmdt0(xf)),
  inference(forward_demodulation,[],[f217,f218])).
fof(f218,plain,(
  xU = szRzazndt0(xf)),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  isOn0(xf,xU) & xU = szRzazndt0(xf) & szDzozmdt0(xf) = szRzazndt0(xf) & isMonotone0(xf) & ! [X0,X1] : (sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)) | ~sdtlseqdt0(X0,X1) | ~aElementOf0(X1,szDzozmdt0(xf)) | ~aElementOf0(X0,szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (sP2(X2) | (~aSubsetOf0(X2,xU) & ((~aElementOf0(sK18(X2),xU) & aElementOf0(sK18(X2),X2)) | ~aSet0(X2)))) & aSet0(xU)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f76,f132])).
fof(f132,plain,(
  ! [X2] : (? [X3] : (~aElementOf0(X3,xU) & aElementOf0(X3,X2)) => (~aElementOf0(sK18(X2),xU) & aElementOf0(sK18(X2),X2)))),
  introduced(choice_axiom,[])).
fof(f76,plain,(
  isOn0(xf,xU) & xU = szRzazndt0(xf) & szDzozmdt0(xf) = szRzazndt0(xf) & isMonotone0(xf) & ! [X0,X1] : (sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)) | ~sdtlseqdt0(X0,X1) | ~aElementOf0(X1,szDzozmdt0(xf)) | ~aElementOf0(X0,szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (sP2(X2) | (~aSubsetOf0(X2,xU) & (? [X3] : (~aElementOf0(X3,xU) & aElementOf0(X3,X2)) | ~aSet0(X2)))) & aSet0(xU)),
  inference(definition_folding,[],[f66,f75,f74,f73])).
fof(f73,plain,(
  ! [X2] : (? [X5] : (aSupremumOfIn0(X5,X2,xU) & ! [X6] : (sdtlseqdt0(X5,X6) | (~aUpperBoundOfIn0(X6,X2,xU) & (? [X7] : (~sdtlseqdt0(X7,X6) & aElementOf0(X7,X2)) | ~aElementOf0(X6,xU)))) & aUpperBoundOfIn0(X5,X2,xU) & ! [X8] : (sdtlseqdt0(X8,X5) | ~aElementOf0(X8,X2)) & aElementOf0(X5,xU) & aElementOf0(X5,xU)) | ~sP0(X2))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f74,plain,(
  ! [X4,X2] : (! [X9] : (sdtlseqdt0(X9,X4) | (~aLowerBoundOfIn0(X9,X2,xU) & (? [X10] : (~sdtlseqdt0(X9,X10) & aElementOf0(X10,X2)) | ~aElementOf0(X9,xU)))) | ~sP1(X4,X2))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f75,plain,(
  ! [X2] : (? [X4] : (sP0(X2) & aInfimumOfIn0(X4,X2,xU) & sP1(X4,X2) & aLowerBoundOfIn0(X4,X2,xU) & ! [X11] : (sdtlseqdt0(X4,X11) | ~aElementOf0(X11,X2)) & aElementOf0(X4,xU) & aElementOf0(X4,xU)) | ~sP2(X2))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f66,plain,(
  isOn0(xf,xU) & xU = szRzazndt0(xf) & szDzozmdt0(xf) = szRzazndt0(xf) & isMonotone0(xf) & ! [X0,X1] : (sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)) | ~sdtlseqdt0(X0,X1) | ~aElementOf0(X1,szDzozmdt0(xf)) | ~aElementOf0(X0,szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (? [X4] : (? [X5] : (aSupremumOfIn0(X5,X2,xU) & ! [X6] : (sdtlseqdt0(X5,X6) | (~aUpperBoundOfIn0(X6,X2,xU) & (? [X7] : (~sdtlseqdt0(X7,X6) & aElementOf0(X7,X2)) | ~aElementOf0(X6,xU)))) & aUpperBoundOfIn0(X5,X2,xU) & ! [X8] : (sdtlseqdt0(X8,X5) | ~aElementOf0(X8,X2)) & aElementOf0(X5,xU) & aElementOf0(X5,xU)) & aInfimumOfIn0(X4,X2,xU) & ! [X9] : (sdtlseqdt0(X9,X4) | (~aLowerBoundOfIn0(X9,X2,xU) & (? [X10] : (~sdtlseqdt0(X9,X10) & aElementOf0(X10,X2)) | ~aElementOf0(X9,xU)))) & aLowerBoundOfIn0(X4,X2,xU) & ! [X11] : (sdtlseqdt0(X4,X11) | ~aElementOf0(X11,X2)) & aElementOf0(X4,xU) & aElementOf0(X4,xU)) | (~aSubsetOf0(X2,xU) & (? [X3] : (~aElementOf0(X3,xU) & aElementOf0(X3,X2)) | ~aSet0(X2)))) & aSet0(xU)),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  isOn0(xf,xU) & xU = szRzazndt0(xf) & szDzozmdt0(xf) = szRzazndt0(xf) & isMonotone0(xf) & ! [X0,X1] : ((sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)) | ~sdtlseqdt0(X0,X1)) | (~aElementOf0(X1,szDzozmdt0(xf)) | ~aElementOf0(X0,szDzozmdt0(xf)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (? [X4] : (? [X5] : (aSupremumOfIn0(X5,X2,xU) & ! [X6] : (sdtlseqdt0(X5,X6) | (~aUpperBoundOfIn0(X6,X2,xU) & (? [X7] : (~sdtlseqdt0(X7,X6) & aElementOf0(X7,X2)) | ~aElementOf0(X6,xU)))) & aUpperBoundOfIn0(X5,X2,xU) & ! [X8] : (sdtlseqdt0(X8,X5) | ~aElementOf0(X8,X2)) & aElementOf0(X5,xU) & aElementOf0(X5,xU)) & aInfimumOfIn0(X4,X2,xU) & ! [X9] : (sdtlseqdt0(X9,X4) | (~aLowerBoundOfIn0(X9,X2,xU) & (? [X10] : (~sdtlseqdt0(X9,X10) & aElementOf0(X10,X2)) | ~aElementOf0(X9,xU)))) & aLowerBoundOfIn0(X4,X2,xU) & ! [X11] : (sdtlseqdt0(X4,X11) | ~aElementOf0(X11,X2)) & aElementOf0(X4,xU) & aElementOf0(X4,xU)) | (~aSubsetOf0(X2,xU) & (? [X3] : (~aElementOf0(X3,xU) & aElementOf0(X3,X2)) | ~aSet0(X2)))) & aSet0(xU)),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  isOn0(xf,xU) & xU = szRzazndt0(xf) & szDzozmdt0(xf) = szRzazndt0(xf) & isMonotone0(xf) & ! [X0,X1] : ((aElementOf0(X1,szDzozmdt0(xf)) & aElementOf0(X0,szDzozmdt0(xf))) => (sdtlseqdt0(X0,X1) => sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : ((aSubsetOf0(X2,xU) | (! [X3] : (aElementOf0(X3,X2) => aElementOf0(X3,xU)) & aSet0(X2))) => ? [X4] : (? [X5] : (aSupremumOfIn0(X5,X2,xU) & ! [X6] : ((aUpperBoundOfIn0(X6,X2,xU) | (! [X7] : (aElementOf0(X7,X2) => sdtlseqdt0(X7,X6)) & aElementOf0(X6,xU))) => sdtlseqdt0(X5,X6)) & aUpperBoundOfIn0(X5,X2,xU) & ! [X8] : (aElementOf0(X8,X2) => sdtlseqdt0(X8,X5)) & aElementOf0(X5,xU) & aElementOf0(X5,xU)) & aInfimumOfIn0(X4,X2,xU) & ! [X9] : ((aLowerBoundOfIn0(X9,X2,xU) | (! [X10] : (aElementOf0(X10,X2) => sdtlseqdt0(X9,X10)) & aElementOf0(X9,xU))) => sdtlseqdt0(X9,X4)) & aLowerBoundOfIn0(X4,X2,xU) & ! [X11] : (aElementOf0(X11,X2) => sdtlseqdt0(X4,X11)) & aElementOf0(X4,xU) & aElementOf0(X4,xU))) & aSet0(xU)),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  isOn0(xf,xU) & xU = szRzazndt0(xf) & szDzozmdt0(xf) = szRzazndt0(xf) & isMonotone0(xf) & ! [X0,X1] : ((aElementOf0(X1,szDzozmdt0(xf)) & aElementOf0(X0,szDzozmdt0(xf))) => (sdtlseqdt0(X0,X1) => sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X0] : ((aSubsetOf0(X0,xU) | (! [X1] : (aElementOf0(X1,X0) => aElementOf0(X1,xU)) & aSet0(X0))) => ? [X1] : (? [X2] : (aSupremumOfIn0(X2,X0,xU) & ! [X3] : ((aUpperBoundOfIn0(X3,X0,xU) | (! [X4] : (aElementOf0(X4,X0) => sdtlseqdt0(X4,X3)) & aElementOf0(X3,xU))) => sdtlseqdt0(X2,X3)) & aUpperBoundOfIn0(X2,X0,xU) & ! [X3] : (aElementOf0(X3,X0) => sdtlseqdt0(X3,X2)) & aElementOf0(X2,xU) & aElementOf0(X2,xU)) & aInfimumOfIn0(X1,X0,xU) & ! [X2] : ((aLowerBoundOfIn0(X2,X0,xU) | (! [X3] : (aElementOf0(X3,X0) => sdtlseqdt0(X2,X3)) & aElementOf0(X2,xU))) => sdtlseqdt0(X2,X1)) & aLowerBoundOfIn0(X1,X0,xU) & ! [X2] : (aElementOf0(X2,X0) => sdtlseqdt0(X1,X2)) & aElementOf0(X1,xU) & aElementOf0(X1,xU))) & aSet0(xU)),
  file('LAT386+4.tptp',unknown)).
fof(f217,plain,(
  szDzozmdt0(xf) = szRzazndt0(xf)),
  inference(cnf_transformation,[],[f133])).
fof(f260,plain,(
  ( ! [X0,X1] : (~aElementOf0(X1,xU) | sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)) | ~sdtlseqdt0(X0,X1) | ~aElementOf0(X0,szDzozmdt0(xf))) )),
  inference(forward_demodulation,[],[f215,f259])).
fof(f215,plain,(
  ( ! [X0,X1] : (sdtlseqdt0(sdtlpdtrp0(xf,X0),sdtlpdtrp0(xf,X1)) | ~sdtlseqdt0(X0,X1) | ~aElementOf0(X1,szDzozmdt0(xf)) | ~aElementOf0(X0,szDzozmdt0(xf))) )),
  inference(cnf_transformation,[],[f133])).
fof(f23342,plain,(
  spl23_403 | ~spl23_7 | ~spl23_39),
  inference(avatar_split_clause,[],[f23323,f916,f294,f12269])).
fof(f916,plain,(
  spl23_39 <=> aSubsetOf0(xT,xU)),
  introduced(avatar_definition,[new_symbols(naming,[spl23_39])])).
fof(f23323,plain,(
  aElementOf0(sK22,xU) | (~spl23_7 | ~spl23_39)),
  inference(resolution,[],[f296,f5160])).
fof(f5160,plain,(
  ( ! [X4] : (~aElementOf0(X4,xT) | aElementOf0(X4,xU)) ) | ~spl23_39),
  inference(subsumption_resolution,[],[f5156,f209])).
fof(f209,plain,(
  aSet0(xU)),
  inference(cnf_transformation,[],[f133])).
fof(f5156,plain,(
  ( ! [X4] : (~aElementOf0(X4,xT) | aElementOf0(X4,xU) | ~aSet0(xU)) ) | ~spl23_39),
  inference(resolution,[],[f917,f148])).
fof(f148,plain,(
  ( ! [X0,X3,X1] : (~aSubsetOf0(X1,X0) | ~aElementOf0(X3,X1) | aElementOf0(X3,X0) | ~aSet0(X0)) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (! [X1] : ((aSubsetOf0(X1,X0) | (~aElementOf0(sK5(X0,X1),X0) & aElementOf0(sK5(X0,X1),X1)) | ~aSet0(X1)) & ((! [X3] : (aElementOf0(X3,X0) | ~aElementOf0(X3,X1)) & aSet0(X1)) | ~aSubsetOf0(X1,X0))) | ~aSet0(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f85,f86])).
fof(f86,plain,(
  ! [X1,X0] : (? [X2] : (~aElementOf0(X2,X0) & aElementOf0(X2,X1)) => (~aElementOf0(sK5(X0,X1),X0) & aElementOf0(sK5(X0,X1),X1)))),
  introduced(choice_axiom,[])).
fof(f85,plain,(
  ! [X0] : (! [X1] : ((aSubsetOf0(X1,X0) | ? [X2] : (~aElementOf0(X2,X0) & aElementOf0(X2,X1)) | ~aSet0(X1)) & ((! [X3] : (aElementOf0(X3,X0) | ~aElementOf0(X3,X1)) & aSet0(X1)) | ~aSubsetOf0(X1,X0))) | ~aSet0(X0))),
  inference(rectify,[],[f84])).
fof(f84,plain,(
  ! [X0] : (! [X1] : ((aSubsetOf0(X1,X0) | ? [X2] : (~aElementOf0(X2,X0) & aElementOf0(X2,X1)) | ~aSet0(X1)) & ((! [X2] : (aElementOf0(X2,X0) | ~aElementOf0(X2,X1)) & aSet0(X1)) | ~aSubsetOf0(X1,X0))) | ~aSet0(X0))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : (! [X1] : ((aSubsetOf0(X1,X0) | (? [X2] : (~aElementOf0(X2,X0) & aElementOf0(X2,X1)) | ~aSet0(X1))) & ((! [X2] : (aElementOf0(X2,X0) | ~aElementOf0(X2,X1)) & aSet0(X1)) | ~aSubsetOf0(X1,X0))) | ~aSet0(X0))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (! [X1] : (aSubsetOf0(X1,X0) <=> (! [X2] : (aElementOf0(X2,X0) | ~aElementOf0(X2,X1)) & aSet0(X1))) | ~aSet0(X0))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0] : (aSet0(X0) => ! [X1] : (aSubsetOf0(X1,X0) <=> (! [X2] : (aElementOf0(X2,X1) => aElementOf0(X2,X0)) & aSet0(X1))))),
  file('LAT386+4.tptp',unknown)).
fof(f917,plain,(
  aSubsetOf0(xT,xU) | ~spl23_39),
  inference(avatar_component_clause,[],[f916])).
fof(f23204,plain,(
  spl23_3 | ~spl23_4 | ~spl23_40),
  inference(avatar_contradiction_clause,[],[f23203])).
fof(f23203,plain,(
  $false | (spl23_3 | ~spl23_4 | ~spl23_40)),
  inference(subsumption_resolution,[],[f23201,f1467])).
fof(f1467,plain,(
  aElementOf0(sK21,xU) | ~spl23_4),
  inference(resolution,[],[f281,f231])).
fof(f231,plain,(
  ( ! [X0] : (~aElementOf0(X0,xP) | aElementOf0(X0,xU)) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  xP = cS1241(xU,xf,xT) & ! [X0] : ((aElementOf0(X0,xP) | (~aUpperBoundOfIn0(X0,xT,xU) & (~sdtlseqdt0(sK19(X0),X0) & aElementOf0(sK19(X0),xT))) | ~sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) | ~aElementOf0(X0,xU)) & ((aUpperBoundOfIn0(X0,xT,xU) & ! [X2] : (sdtlseqdt0(X2,X0) | ~aElementOf0(X2,xT)) & sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) & aElementOf0(X0,xU)) | ~aElementOf0(X0,xP))) & aSet0(xP)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f70,f134])).
fof(f134,plain,(
  ! [X0] : (? [X1] : (~sdtlseqdt0(X1,X0) & aElementOf0(X1,xT)) => (~sdtlseqdt0(sK19(X0),X0) & aElementOf0(sK19(X0),xT)))),
  introduced(choice_axiom,[])).
fof(f70,plain,(
  xP = cS1241(xU,xf,xT) & ! [X0] : ((aElementOf0(X0,xP) | (~aUpperBoundOfIn0(X0,xT,xU) & ? [X1] : (~sdtlseqdt0(X1,X0) & aElementOf0(X1,xT))) | ~sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) | ~aElementOf0(X0,xU)) & ((aUpperBoundOfIn0(X0,xT,xU) & ! [X2] : (sdtlseqdt0(X2,X0) | ~aElementOf0(X2,xT)) & sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) & aElementOf0(X0,xU)) | ~aElementOf0(X0,xP))) & aSet0(xP)),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  xP = cS1241(xU,xf,xT) & ! [X0] : ((aElementOf0(X0,xP) | ((~aUpperBoundOfIn0(X0,xT,xU) & ? [X1] : (~sdtlseqdt0(X1,X0) & aElementOf0(X1,xT))) | ~sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) | ~aElementOf0(X0,xU))) & ((aUpperBoundOfIn0(X0,xT,xU) & ! [X2] : (sdtlseqdt0(X2,X0) | ~aElementOf0(X2,xT)) & sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) & aElementOf0(X0,xU)) | ~aElementOf0(X0,xP))) & aSet0(xP)),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  xP = cS1241(xU,xf,xT) & ! [X0] : ((((aUpperBoundOfIn0(X0,xT,xU) | ! [X1] : (aElementOf0(X1,xT) => sdtlseqdt0(X1,X0))) & sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) & aElementOf0(X0,xU)) => aElementOf0(X0,xP)) & (aElementOf0(X0,xP) => (aUpperBoundOfIn0(X0,xT,xU) & ! [X2] : (aElementOf0(X2,xT) => sdtlseqdt0(X2,X0)) & sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) & aElementOf0(X0,xU)))) & aSet0(xP)),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  xP = cS1241(xU,xf,xT) & ! [X0] : ((((aUpperBoundOfIn0(X0,xT,xU) | ! [X1] : (aElementOf0(X1,xT) => sdtlseqdt0(X1,X0))) & sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) & aElementOf0(X0,xU)) => aElementOf0(X0,xP)) & (aElementOf0(X0,xP) => (aUpperBoundOfIn0(X0,xT,xU) & ! [X1] : (aElementOf0(X1,xT) => sdtlseqdt0(X1,X0)) & sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) & aElementOf0(X0,xU)))) & aSet0(xP)),
  file('LAT386+4.tptp',unknown)).
fof(f281,plain,(
  aElementOf0(sK21,xP) | ~spl23_4),
  inference(avatar_component_clause,[],[f279])).
fof(f279,plain,(
  spl23_4 <=> aElementOf0(sK21,xP)),
  introduced(avatar_definition,[new_symbols(naming,[spl23_4])])).
fof(f23201,plain,(
  ~aElementOf0(sK21,xU) | (spl23_3 | ~spl23_4 | ~spl23_40)),
  inference(resolution,[],[f23171,f1360])).
fof(f1360,plain,(
  ( ! [X2] : (aElement0(sdtlpdtrp0(xf,X2)) | ~aElementOf0(X2,xU)) )),
  inference(subsumption_resolution,[],[f1358,f209])).
fof(f1358,plain,(
  ( ! [X2] : (~aElementOf0(X2,xU) | aElement0(sdtlpdtrp0(xf,X2)) | ~aSet0(xU)) )),
  inference(resolution,[],[f576,f144])).
fof(f144,plain,(
  ( ! [X0,X1] : (~aElementOf0(X1,X0) | aElement0(X1) | ~aSet0(X0)) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (! [X1] : (aElement0(X1) | ~aElementOf0(X1,X0)) | ~aSet0(X0))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1,X0) => aElement0(X1)))),
  file('LAT386+4.tptp',unknown)).
fof(f576,plain,(
  ( ! [X0] : (aElementOf0(sdtlpdtrp0(xf,X0),xU) | ~aElementOf0(X0,xU)) )),
  inference(forward_demodulation,[],[f575,f259])).
fof(f575,plain,(
  ( ! [X0] : (aElementOf0(sdtlpdtrp0(xf,X0),xU) | ~aElementOf0(X0,szDzozmdt0(xf))) )),
  inference(subsumption_resolution,[],[f570,f214])).
fof(f214,plain,(
  aFunction0(xf)),
  inference(cnf_transformation,[],[f133])).
fof(f570,plain,(
  ( ! [X0] : (aElementOf0(sdtlpdtrp0(xf,X0),xU) | ~aElementOf0(X0,szDzozmdt0(xf)) | ~aFunction0(xf)) )),
  inference(superposition,[],[f182,f218])).
fof(f182,plain,(
  ( ! [X0,X1] : (aElementOf0(sdtlpdtrp0(X0,X1),szRzazndt0(X0)) | ~aElementOf0(X1,szDzozmdt0(X0)) | ~aFunction0(X0)) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (! [X1] : (aElementOf0(sdtlpdtrp0(X0,X1),szRzazndt0(X0)) | ~aElementOf0(X1,szDzozmdt0(X0))) | ~aFunction0(X0))),
  inference(ennf_transformation,[],[f21])).
fof(f21,axiom,(
  ! [X0] : (aFunction0(X0) => ! [X1] : (aElementOf0(X1,szDzozmdt0(X0)) => aElementOf0(sdtlpdtrp0(X0,X1),szRzazndt0(X0))))),
  file('LAT386+4.tptp',unknown)).
fof(f23171,plain,(
  ~aElement0(sdtlpdtrp0(xf,sK21)) | (spl23_3 | ~spl23_4 | ~spl23_40)),
  inference(subsumption_resolution,[],[f23170,f1466])).
fof(f1466,plain,(
  sdtlseqdt0(xp,sK21) | ~spl23_4),
  inference(resolution,[],[f281,f241])).
fof(f241,plain,(
  ( ! [X2] : (~aElementOf0(X2,xP) | sdtlseqdt0(xp,X2)) )),
  inference(cnf_transformation,[],[f137])).
fof(f23170,plain,(
  ~aElement0(sdtlpdtrp0(xf,sK21)) | ~sdtlseqdt0(xp,sK21) | (spl23_3 | ~spl23_4 | ~spl23_40)),
  inference(subsumption_resolution,[],[f23169,f240])).
fof(f23169,plain,(
  ~aElement0(sdtlpdtrp0(xf,sK21)) | ~aElementOf0(xp,xU) | ~sdtlseqdt0(xp,sK21) | (spl23_3 | ~spl23_4 | ~spl23_40)),
  inference(subsumption_resolution,[],[f23168,f281])).
fof(f23168,plain,(
  ~aElement0(sdtlpdtrp0(xf,sK21)) | ~aElementOf0(sK21,xP) | ~aElementOf0(xp,xU) | ~sdtlseqdt0(xp,sK21) | (spl23_3 | ~spl23_4 | ~spl23_40)),
  inference(subsumption_resolution,[],[f23167,f1610])).
fof(f1610,plain,(
  aElement0(sdtlpdtrp0(xf,xp)) | ~spl23_40),
  inference(subsumption_resolution,[],[f1600,f209])).
fof(f1600,plain,(
  aElement0(sdtlpdtrp0(xf,xp)) | ~aSet0(xU) | ~spl23_40),
  inference(resolution,[],[f921,f144])).
fof(f921,plain,(
  aElementOf0(sdtlpdtrp0(xf,xp),xU) | ~spl23_40),
  inference(avatar_component_clause,[],[f920])).
fof(f920,plain,(
  spl23_40 <=> aElementOf0(sdtlpdtrp0(xf,xp),xU)),
  introduced(avatar_definition,[new_symbols(naming,[spl23_40])])).
fof(f23167,plain,(
  ~aElement0(sdtlpdtrp0(xf,sK21)) | ~aElement0(sdtlpdtrp0(xf,xp)) | ~aElementOf0(sK21,xP) | ~aElementOf0(xp,xU) | ~sdtlseqdt0(xp,sK21) | (spl23_3 | ~spl23_4)),
  inference(subsumption_resolution,[],[f23076,f1470])).
fof(f1470,plain,(
  aElement0(sK21) | ~spl23_4),
  inference(subsumption_resolution,[],[f1469,f230])).
fof(f230,plain,(
  aSet0(xP)),
  inference(cnf_transformation,[],[f135])).
fof(f1469,plain,(
  aElement0(sK21) | ~aSet0(xP) | ~spl23_4),
  inference(resolution,[],[f281,f144])).
fof(f23076,plain,(
  ~aElement0(sK21) | ~aElement0(sdtlpdtrp0(xf,sK21)) | ~aElement0(sdtlpdtrp0(xf,xp)) | ~aElementOf0(sK21,xP) | ~aElementOf0(xp,xU) | ~sdtlseqdt0(xp,sK21) | spl23_3),
  inference(resolution,[],[f3099,f276])).
fof(f276,plain,(
  ~sdtlseqdt0(sdtlpdtrp0(xf,xp),sK21) | spl23_3),
  inference(avatar_component_clause,[],[f274])).
fof(f274,plain,(
  spl23_3 <=> sdtlseqdt0(sdtlpdtrp0(xf,xp),sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl23_3])])).
fof(f3099,plain,(
  ( ! [X2,X1] : (sdtlseqdt0(sdtlpdtrp0(xf,X1),X2) | ~aElement0(X2) | ~aElement0(sdtlpdtrp0(xf,X2)) | ~aElement0(sdtlpdtrp0(xf,X1)) | ~aElementOf0(X2,xP) | ~aElementOf0(X1,xU) | ~sdtlseqdt0(X1,X2)) )),
  inference(subsumption_resolution,[],[f3083,f231])).
fof(f3083,plain,(
  ( ! [X2,X1] : (sdtlseqdt0(sdtlpdtrp0(xf,X1),X2) | ~aElement0(X2) | ~aElement0(sdtlpdtrp0(xf,X2)) | ~aElement0(sdtlpdtrp0(xf,X1)) | ~aElementOf0(X2,xP) | ~aElementOf0(X2,xU) | ~aElementOf0(X1,xU) | ~sdtlseqdt0(X1,X2)) )),
  inference(resolution,[],[f789,f261])).
fof(f789,plain,(
  ( ! [X6,X5] : (~sdtlseqdt0(X5,sdtlpdtrp0(xf,X6)) | sdtlseqdt0(X5,X6) | ~aElement0(X6) | ~aElement0(sdtlpdtrp0(xf,X6)) | ~aElement0(X5) | ~aElementOf0(X6,xP)) )),
  inference(resolution,[],[f153,f232])).
fof(f232,plain,(
  ( ! [X0] : (sdtlseqdt0(sdtlpdtrp0(xf,X0),X0) | ~aElementOf0(X0,xP)) )),
  inference(cnf_transformation,[],[f135])).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (~sdtlseqdt0(X1,X2) | sdtlseqdt0(X0,X2) | ~sdtlseqdt0(X0,X1) | ~aElement0(X2) | ~aElement0(X1) | ~aElement0(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : (sdtlseqdt0(X0,X2) | ~sdtlseqdt0(X1,X2) | ~sdtlseqdt0(X0,X1) | ~aElement0(X2) | ~aElement0(X1) | ~aElement0(X0))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : ((sdtlseqdt0(X0,X2) | (~sdtlseqdt0(X1,X2) | ~sdtlseqdt0(X0,X1))) | (~aElement0(X2) | ~aElement0(X1) | ~aElement0(X0)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2] : ((aElement0(X2) & aElement0(X1) & aElement0(X0)) => ((sdtlseqdt0(X1,X2) & sdtlseqdt0(X0,X1)) => sdtlseqdt0(X0,X2)))),
  file('LAT386+4.tptp',unknown)).
fof(f4907,plain,(
  spl23_39),
  inference(avatar_split_clause,[],[f4906,f916])).
fof(f4906,plain,(
  aSubsetOf0(xT,xU)),
  inference(subsumption_resolution,[],[f4905,f227])).
fof(f227,plain,(
  aSet0(xT)),
  inference(cnf_transformation,[],[f68])).
fof(f4905,plain,(
  aSubsetOf0(xT,xU) | ~aSet0(xT)),
  inference(subsumption_resolution,[],[f4900,f209])).
fof(f4900,plain,(
  aSubsetOf0(xT,xU) | ~aSet0(xU) | ~aSet0(xT)),
  inference(duplicate_literal_removal,[],[f4896])).
fof(f4896,plain,(
  aSubsetOf0(xT,xU) | ~aSet0(xU) | aSubsetOf0(xT,xU) | ~aSet0(xT) | ~aSet0(xU)),
  inference(resolution,[],[f1554,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (~aElementOf0(sK5(X0,X1),X0) | aSubsetOf0(X1,X0) | ~aSet0(X1) | ~aSet0(X0)) )),
  inference(cnf_transformation,[],[f87])).
fof(f1554,plain,(
  ( ! [X1] : (aElementOf0(sK5(X1,xT),xU) | aSubsetOf0(xT,X1) | ~aSet0(X1)) )),
  inference(resolution,[],[f527,f263])).
fof(f263,plain,(
  ( ! [X0] : (~aElementOf0(X0,xS) | aElementOf0(X0,xU)) )),
  inference(forward_demodulation,[],[f221,f259])).
fof(f221,plain,(
  ( ! [X0] : (aElementOf0(X0,szDzozmdt0(xf)) | ~aElementOf0(X0,xS)) )),
  inference(cnf_transformation,[],[f67])).
fof(f527,plain,(
  ( ! [X9] : (aElementOf0(sK5(X9,xT),xS) | ~aSet0(X9) | aSubsetOf0(xT,X9)) )),
  inference(subsumption_resolution,[],[f518,f227])).
fof(f518,plain,(
  ( ! [X9] : (aSubsetOf0(xT,X9) | ~aSet0(xT) | ~aSet0(X9) | aElementOf0(sK5(X9,xT),xS)) )),
  inference(resolution,[],[f149,f228])).
fof(f149,plain,(
  ( ! [X0,X1] : (aElementOf0(sK5(X0,X1),X1) | aSubsetOf0(X1,X0) | ~aSet0(X1) | ~aSet0(X0)) )),
  inference(cnf_transformation,[],[f87])).
fof(f4546,plain,(
  spl23_80 | ~spl23_7),
  inference(avatar_split_clause,[],[f4545,f294,f1440])).
fof(f4545,plain,(
  sdtlseqdt0(sK22,xp) | ~spl23_7),
  inference(subsumption_resolution,[],[f4535,f2234])).
fof(f2234,plain,(
  aElementOf0(sK22,xU) | ~spl23_7),
  inference(resolution,[],[f2225,f263])).
fof(f2225,plain,(
  aElementOf0(sK22,xS) | ~spl23_7),
  inference(resolution,[],[f296,f228])).
fof(f4535,plain,(
  sdtlseqdt0(sK22,xp) | ~aElementOf0(sK22,xU) | ~spl23_7),
  inference(duplicate_literal_removal,[],[f4531])).
fof(f4531,plain,(
  sdtlseqdt0(sK22,xp) | ~aElementOf0(sK22,xU) | sdtlseqdt0(sK22,xp) | ~aElementOf0(sK22,xU) | ~spl23_7),
  inference(resolution,[],[f2250,f244])).
fof(f244,plain,(
  ( ! [X0] : (~sdtlseqdt0(X0,sK20(X0)) | sdtlseqdt0(X0,xp) | ~aElementOf0(X0,xU)) )),
  inference(cnf_transformation,[],[f137])).
fof(f2250,plain,(
  ( ! [X10] : (sdtlseqdt0(sK22,sK20(X10)) | sdtlseqdt0(X10,xp) | ~aElementOf0(X10,xU)) ) | ~spl23_7),
  inference(resolution,[],[f2224,f243])).
fof(f243,plain,(
  ( ! [X0] : (aElementOf0(sK20(X0),xP) | sdtlseqdt0(X0,xp) | ~aElementOf0(X0,xU)) )),
  inference(cnf_transformation,[],[f137])).
fof(f2224,plain,(
  ( ! [X0] : (~aElementOf0(X0,xP) | sdtlseqdt0(sK22,X0)) ) | ~spl23_7),
  inference(resolution,[],[f296,f233])).
fof(f233,plain,(
  ( ! [X2,X0] : (~aElementOf0(X2,xT) | sdtlseqdt0(X2,X0) | ~aElementOf0(X0,xP)) )),
  inference(cnf_transformation,[],[f135])).
fof(f1592,plain,(
  spl23_40),
  inference(avatar_contradiction_clause,[],[f1591])).
fof(f1591,plain,(
  $false | spl23_40),
  inference(subsumption_resolution,[],[f1590,f240])).
fof(f1590,plain,(
  ~aElementOf0(xp,xU) | spl23_40),
  inference(resolution,[],[f922,f576])).
fof(f922,plain,(
  ~aElementOf0(sdtlpdtrp0(xf,xp),xU) | spl23_40),
  inference(avatar_component_clause,[],[f920])).
fof(f297,plain,(
  spl23_1 | spl23_7),
  inference(avatar_split_clause,[],[f250,f294,f265])).
fof(f265,plain,(
  spl23_1 <=> sP3),
  introduced(avatar_definition,[new_symbols(naming,[spl23_1])])).
fof(f250,plain,(
  aElementOf0(sK22,xT) | sP3),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  (~aUpperBoundOfIn0(sdtlpdtrp0(xf,xp),xT,xU) & (~sdtlseqdt0(sK22,sdtlpdtrp0(xf,xp)) & aElementOf0(sK22,xT))) | sP3),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f78,f142])).
fof(f142,plain,(
  ? [X0] : (~sdtlseqdt0(X0,sdtlpdtrp0(xf,xp)) & aElementOf0(X0,xT)) => (~sdtlseqdt0(sK22,sdtlpdtrp0(xf,xp)) & aElementOf0(sK22,xT))),
  introduced(choice_axiom,[])).
fof(f78,plain,(
  (~aUpperBoundOfIn0(sdtlpdtrp0(xf,xp),xT,xU) & ? [X0] : (~sdtlseqdt0(X0,sdtlpdtrp0(xf,xp)) & aElementOf0(X0,xT))) | sP3),
  inference(definition_folding,[],[f72,f77])).
fof(f77,plain,(
  (~aLowerBoundOfIn0(sdtlpdtrp0(xf,xp),xP,xU) & ? [X1] : (~sdtlseqdt0(sdtlpdtrp0(xf,xp),X1) & aElementOf0(X1,xP))) | ~sP3),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f72,plain,(
  (~aUpperBoundOfIn0(sdtlpdtrp0(xf,xp),xT,xU) & ? [X0] : (~sdtlseqdt0(X0,sdtlpdtrp0(xf,xp)) & aElementOf0(X0,xT))) | (~aLowerBoundOfIn0(sdtlpdtrp0(xf,xp),xP,xU) & ? [X1] : (~sdtlseqdt0(sdtlpdtrp0(xf,xp),X1) & aElementOf0(X1,xP)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~((aUpperBoundOfIn0(sdtlpdtrp0(xf,xp),xT,xU) | ! [X0] : (aElementOf0(X0,xT) => sdtlseqdt0(X0,sdtlpdtrp0(xf,xp)))) & (aLowerBoundOfIn0(sdtlpdtrp0(xf,xp),xP,xU) | ! [X1] : (aElementOf0(X1,xP) => sdtlseqdt0(sdtlpdtrp0(xf,xp),X1))))),
  inference(rectify,[],[f30])).
fof(f30,negated_conjecture,(
  ~((aUpperBoundOfIn0(sdtlpdtrp0(xf,xp),xT,xU) | ! [X0] : (aElementOf0(X0,xT) => sdtlseqdt0(X0,sdtlpdtrp0(xf,xp)))) & (aLowerBoundOfIn0(sdtlpdtrp0(xf,xp),xP,xU) | ! [X0] : (aElementOf0(X0,xP) => sdtlseqdt0(sdtlpdtrp0(xf,xp),X0))))),
  inference(negated_conjecture,[],[f29])).
fof(f29,conjecture,(
  (aUpperBoundOfIn0(sdtlpdtrp0(xf,xp),xT,xU) | ! [X0] : (aElementOf0(X0,xT) => sdtlseqdt0(X0,sdtlpdtrp0(xf,xp)))) & (aLowerBoundOfIn0(sdtlpdtrp0(xf,xp),xP,xU) | ! [X0] : (aElementOf0(X0,xP) => sdtlseqdt0(sdtlpdtrp0(xf,xp),X0)))),
  file('LAT386+4.tptp',unknown)).
fof(f292,plain,(
  spl23_1 | ~spl23_6),
  inference(avatar_split_clause,[],[f251,f289,f265])).
fof(f251,plain,(
  ~sdtlseqdt0(sK22,sdtlpdtrp0(xf,xp)) | sP3),
  inference(cnf_transformation,[],[f143])).
fof(f282,plain,(
  ~spl23_1 | spl23_4),
  inference(avatar_split_clause,[],[f247,f279,f265])).
fof(f247,plain,(
  aElementOf0(sK21,xP) | ~sP3),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  (~aLowerBoundOfIn0(sdtlpdtrp0(xf,xp),xP,xU) & (~sdtlseqdt0(sdtlpdtrp0(xf,xp),sK21) & aElementOf0(sK21,xP))) | ~sP3),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21])],[f139,f140])).
fof(f140,plain,(
  ? [X0] : (~sdtlseqdt0(sdtlpdtrp0(xf,xp),X0) & aElementOf0(X0,xP)) => (~sdtlseqdt0(sdtlpdtrp0(xf,xp),sK21) & aElementOf0(sK21,xP))),
  introduced(choice_axiom,[])).
fof(f139,plain,(
  (~aLowerBoundOfIn0(sdtlpdtrp0(xf,xp),xP,xU) & ? [X0] : (~sdtlseqdt0(sdtlpdtrp0(xf,xp),X0) & aElementOf0(X0,xP))) | ~sP3),
  inference(rectify,[],[f138])).
fof(f138,plain,(
  (~aLowerBoundOfIn0(sdtlpdtrp0(xf,xp),xP,xU) & ? [X1] : (~sdtlseqdt0(sdtlpdtrp0(xf,xp),X1) & aElementOf0(X1,xP))) | ~sP3),
  inference(nnf_transformation,[],[f77])).
fof(f277,plain,(
  ~spl23_1 | ~spl23_3),
  inference(avatar_split_clause,[],[f248,f274,f265])).
fof(f248,plain,(
  ~sdtlseqdt0(sdtlpdtrp0(xf,xp),sK21) | ~sP3),
  inference(cnf_transformation,[],[f141])).
% SZS output end Proof for LAT386+4
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 34413
% Time elapsed: 1.398 s
% ------------------------------
% ------------------------------
