% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for LAT081-1
% SZS output start Proof for LAT081-1
fof(f15132,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f15131])).
fof(f15131,plain,(
  meet(a,b) != meet(a,b)),
  inference(superposition,[],[f2,f14563])).
fof(f14563,plain,(
  ( ! [X2,X3] : (meet(X2,X3) = meet(X3,X2)) )),
  inference(superposition,[],[f7934,f7932])).
fof(f7932,plain,(
  ( ! [X70,X69] : (meet(X69,X70) = join(meet(X69,X70),meet(X70,X69))) )),
  inference(superposition,[],[f3811,f7143])).
fof(f7143,plain,(
  ( ! [X8,X9] : (meet(X8,X9) = meet(meet(X9,X8),meet(X8,X9))) )),
  inference(superposition,[],[f2914,f3811])).
fof(f2914,plain,(
  ( ! [X30,X31,X29] : (meet(X30,X29) = meet(meet(X29,join(X31,meet(X30,X29))),meet(X30,X29))) )),
  inference(superposition,[],[f2771,f2241])).
fof(f2241,plain,(
  ( ! [X165,X169] : (join(X165,meet(X169,X165)) = X165) )),
  inference(backward_demodulation,[],[f2233,f2028])).
fof(f2028,plain,(
  ( ! [X165,X169] : (join(meet(X165,X165),meet(X169,X165)) = X165) )),
  inference(backward_demodulation,[],[f1991,f1948])).
fof(f1948,plain,(
  ( ! [X165,X169] : (join(meet(X165,X165),meet(X169,join(X165,X165))) = X165) )),
  inference(backward_demodulation,[],[f1935,f666])).
fof(f666,plain,(
  ( ! [X167,X165,X169,X166,X168] : (join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))) = X165) )),
  inference(backward_demodulation,[],[f665,f437])).
fof(f437,plain,(
  ( ! [X167,X165,X169,X166,X170,X168] : (join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))) = join(meet(X165,X170),meet(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),join(X165,X170)))) )),
  inference(backward_demodulation,[],[f363,f325])).
fof(f325,plain,(
  ( ! [X167,X165,X173,X171,X169,X166,X174,X172,X170,X168] : (join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))) = join(meet(X165,X170),meet(join(meet(join(meet(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),X165),meet(X165,join(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),X165))),join(join(meet(X171,join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165))))),meet(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),X172)),join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))))),meet(join(meet(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),meet(meet(join(X171,join(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),X172)),join(X173,join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))))),join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))))),meet(X174,join(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),meet(meet(join(X171,join(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),X172)),join(X173,join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))))),join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))))))),join(join(meet(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),X165),meet(X165,join(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),X165))),join(join(meet(X171,join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165))))),meet(join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))),X172)),join(meet(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)),meet(X169,join(X165,meet(meet(join(X166,join(X165,X167)),join(X168,X165)),X165)))))))),join(X165,X170)))) )),
  inference(superposition,[],[f8,f172])).
fof(f172,plain,(
  ( ! [X2,X1] : (join(meet(X1,X2),meet(X1,join(X1,X2))) = X1) )),
  inference(superposition,[],[f71,f71])).
fof(f71,plain,(
  ( ! [X0,X8,X1] : (join(meet(join(meet(join(meet(X0,X1),meet(X1,join(X0,X1))),X1),meet(X1,join(join(meet(X0,X1),meet(X1,join(X0,X1))),X1))),X8),meet(X1,join(join(meet(join(meet(X0,X1),meet(X1,join(X0,X1))),X1),meet(X1,join(join(meet(X0,X1),meet(X1,join(X0,X1))),X1))),X8))) = X1) )),
  inference(superposition,[],[f13,f1])).
fof(f1,axiom,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (join(meet(join(meet(X0,X1),meet(X1,join(X0,X1))),X2),meet(join(meet(X0,join(join(meet(X3,X1),meet(X1,X4)),X1)),meet(join(meet(X1,meet(meet(join(X3,join(X1,X4)),join(X5,X1)),X1)),meet(X6,join(X1,meet(meet(join(X3,join(X1,X4)),join(X5,X1)),X1)))),join(X0,join(join(meet(X3,X1),meet(X1,X4)),X1)))),join(join(meet(X0,X1),meet(X1,join(X0,X1))),X2))) = X1) )),
  file('LAT081-1.tptp',unknown)).
fof(f13,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (join(meet(join(meet(X6,X1),meet(X1,join(X6,X1))),X7),meet(join(meet(X6,join(join(meet(X8,X1),meet(X1,X9)),X1)),meet(join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))),join(X6,join(join(meet(X8,X1),meet(X1,X9)),X1)))),join(join(meet(X6,X1),meet(X1,join(X6,X1))),X7))) = X1) )),
  inference(superposition,[],[f5,f1])).
fof(f5,plain,(
  ( ! [X2,X0,X10,X3,X1,X9] : (join(meet(X0,X1),meet(X1,join(X0,X1))) = join(meet(join(meet(X9,join(meet(X0,X1),meet(X1,join(X0,X1)))),meet(join(meet(X0,X1),meet(X1,join(X0,X1))),join(X9,join(meet(X0,X1),meet(X1,join(X0,X1)))))),X10),meet(join(meet(X9,join(join(meet(X2,join(meet(X0,X1),meet(X1,join(X0,X1)))),meet(join(meet(X0,X1),meet(X1,join(X0,X1))),X3)),join(meet(X0,X1),meet(X1,join(X0,X1))))),meet(X1,join(X9,join(join(meet(X2,join(meet(X0,X1),meet(X1,join(X0,X1)))),meet(join(meet(X0,X1),meet(X1,join(X0,X1))),X3)),join(meet(X0,X1),meet(X1,join(X0,X1))))))),join(join(meet(X9,join(meet(X0,X1),meet(X1,join(X0,X1)))),meet(join(meet(X0,X1),meet(X1,join(X0,X1))),join(X9,join(meet(X0,X1),meet(X1,join(X0,X1)))))),X10)))) )),
  inference(superposition,[],[f1,f1])).
fof(f8,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))) = join(meet(X1,X6),meet(join(meet(join(meet(X0,X1),meet(X1,join(X0,X1))),join(join(meet(X7,join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1))))),meet(join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))),X8)),join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))))),meet(join(meet(join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))),meet(meet(join(X7,join(join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))),X8)),join(X9,join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))))),join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))))),meet(X10,join(join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))),meet(meet(join(X7,join(join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))),X8)),join(X9,join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))))),join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))))))),join(join(meet(X0,X1),meet(X1,join(X0,X1))),join(join(meet(X7,join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1))))),meet(join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))),X8)),join(meet(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)),meet(join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)))),join(X0,join(join(meet(X2,X1),meet(X1,X3)),X1)))))))),join(X1,X6)))) )),
  inference(superposition,[],[f1,f1])).
fof(f363,plain,(
  ( ! [X271,X273,X275,X276,X272,X274] : (join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))) = X273) )),
  inference(forward_demodulation,[],[f362,f172])).
fof(f362,plain,(
  ( ! [X280,X271,X273,X275,X276,X272,X274] : (join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))) = join(meet(X273,X280),meet(X273,join(X273,X280)))) )),
  inference(forward_demodulation,[],[f338,f1])).
fof(f338,plain,(
  ( ! [X280,X271,X277,X273,X275,X276,X278,X272,X274] : (join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))) = join(meet(X273,X280),meet(join(meet(join(meet(X271,X273),meet(X273,join(X271,X273))),join(join(meet(X277,join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273))))),meet(join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))),X278)),join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))))),meet(join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))),join(join(meet(X271,X273),meet(X273,join(X271,X273))),join(join(meet(X277,join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273))))),meet(join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))),X278)),join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(join(meet(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)),meet(X276,join(X273,meet(meet(join(X272,join(X273,X274)),join(X275,X273)),X273)))),join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))))))),join(X273,X280)))) )),
  inference(superposition,[],[f8,f172])).
fof(f665,plain,(
  ( ! [X61,X59,X62,X60,X58,X53] : (join(meet(X53,X58),meet(join(meet(X53,meet(meet(join(X59,join(X53,X60)),join(X61,X53)),X53)),meet(X62,join(X53,meet(meet(join(X59,join(X53,X60)),join(X61,X53)),X53)))),join(X53,X58))) = X53) )),
  inference(forward_demodulation,[],[f441,f437])).
fof(f441,plain,(
  ( ! [X61,X59,X62,X60,X58,X53] : (join(meet(X53,X58),meet(join(meet(X53,join(join(meet(X59,X53),meet(X53,X60)),X53)),meet(join(meet(X53,meet(meet(join(X59,join(X53,X60)),join(X61,X53)),X53)),meet(X62,join(X53,meet(meet(join(X59,join(X53,X60)),join(X61,X53)),X53)))),join(X53,join(join(meet(X59,X53),meet(X53,X60)),X53)))),join(X53,X58))) = X53) )),
  inference(backward_demodulation,[],[f363,f178])).
fof(f178,plain,(
  ( ! [X61,X59,X57,X54,X52,X62,X60,X58,X56,X55,X53] : (join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))) = join(meet(X53,X58),meet(join(meet(X53,join(join(meet(X59,join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53))))),meet(join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))),X60)),join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))))),meet(join(meet(join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))),meet(meet(join(X59,join(join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))),X60)),join(X61,join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))))),join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))))),meet(X62,join(join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))),meet(meet(join(X59,join(join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))),X60)),join(X61,join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))))),join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))))))),join(X53,join(join(meet(X59,join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53))))),meet(join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))),X60)),join(meet(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)),meet(join(meet(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)),meet(X57,join(X53,meet(meet(join(X54,join(X53,X55)),join(X56,X53)),X53)))),join(join(meet(join(meet(X52,X53),meet(X53,join(X52,X53))),X53),meet(X53,join(join(meet(X52,X53),meet(X53,join(X52,X53))),X53))),join(join(meet(X54,X53),meet(X53,X55)),X53)))))))),join(X53,X58)))) )),
  inference(superposition,[],[f8,f71])).
fof(f1935,plain,(
  ( ! [X12,X10,X11,X9] : (meet(meet(join(X10,join(X9,X11)),join(X12,X9)),X9) = X9) )),
  inference(superposition,[],[f666,f750])).
fof(f750,plain,(
  ( ! [X222,X221] : (join(meet(X221,X222),meet(X222,join(X221,X222))) = X222) )),
  inference(forward_demodulation,[],[f747,f721])).
fof(f721,plain,(
  ( ! [X88,X87,X86] : (join(meet(join(meet(X87,X86),meet(X86,join(X87,X86))),X88),meet(X86,join(join(meet(X87,X86),meet(X86,join(X87,X86))),X88))) = X86) )),
  inference(backward_demodulation,[],[f709,f181])).
fof(f181,plain,(
  ( ! [X90,X88,X87,X89,X86] : (join(meet(join(meet(X87,X86),meet(X86,join(X87,X86))),X88),meet(join(meet(X87,join(join(meet(X89,X86),meet(X86,X90)),X86)),meet(X86,join(X87,join(join(meet(X89,X86),meet(X86,X90)),X86)))),join(join(meet(X87,X86),meet(X86,join(X87,X86))),X88))) = X86) )),
  inference(superposition,[],[f5,f71])).
fof(f709,plain,(
  ( ! [X271,X273,X272,X274] : (join(meet(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)),meet(X273,join(X271,join(join(meet(X272,X273),meet(X273,X274)),X273)))) = X273) )),
  inference(backward_demodulation,[],[f666,f363])).
fof(f747,plain,(
  ( ! [X222,X223,X221] : (join(meet(X221,X222),meet(X222,join(X221,X222))) = join(meet(join(meet(X221,X222),meet(X222,join(X221,X222))),X223),meet(X222,join(join(meet(X221,X222),meet(X222,join(X221,X222))),X223)))) )),
  inference(backward_demodulation,[],[f721,f331])).
fof(f331,plain,(
  ( ! [X222,X225,X223,X221,X224] : (join(meet(X221,X222),meet(X222,join(X221,X222))) = join(meet(join(meet(X221,X222),meet(X222,join(X221,X222))),X223),meet(join(meet(join(meet(X221,X222),meet(X222,join(X221,X222))),join(join(meet(X224,join(meet(X221,X222),meet(X222,join(X221,X222)))),meet(join(meet(X221,X222),meet(X222,join(X221,X222))),X225)),join(meet(X221,X222),meet(X222,join(X221,X222))))),meet(X222,join(join(meet(X221,X222),meet(X222,join(X221,X222))),join(join(meet(X224,join(meet(X221,X222),meet(X222,join(X221,X222)))),meet(join(meet(X221,X222),meet(X222,join(X221,X222))),X225)),join(meet(X221,X222),meet(X222,join(X221,X222))))))),join(join(meet(X221,X222),meet(X222,join(X221,X222))),X223)))) )),
  inference(superposition,[],[f5,f172])).
fof(f1991,plain,(
  ( ! [X2] : (join(X2,X2) = X2) )),
  inference(superposition,[],[f1948,f852])).
fof(f852,plain,(
  ( ! [X116,X117] : (join(meet(X117,join(X116,X116)),meet(X116,join(X117,join(X116,X116)))) = X116) )),
  inference(backward_demodulation,[],[f851,f691])).
fof(f691,plain,(
  ( ! [X116,X120,X117] : (join(meet(X117,join(X116,X116)),meet(X116,join(X117,join(X116,X116)))) = join(meet(X116,X120),meet(join(meet(X117,join(X116,X116)),meet(X116,join(X117,join(X116,X116)))),join(X116,X120)))) )),
  inference(backward_demodulation,[],[f666,f365])).
fof(f365,plain,(
  ( ! [X118,X116,X120,X119,X117] : (join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))) = join(meet(X116,X120),meet(join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))),join(X116,X120)))) )),
  inference(backward_demodulation,[],[f363,f319])).
fof(f319,plain,(
  ( ! [X123,X121,X118,X116,X124,X122,X120,X119,X117] : (join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))) = join(meet(X116,X120),meet(join(meet(join(meet(X117,X116),meet(X116,join(X117,X116))),join(join(meet(X121,join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116))))),meet(join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))),X122)),join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))))),meet(join(meet(join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))),meet(meet(join(X121,join(join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))),X122)),join(X123,join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))))),join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))))),meet(X124,join(join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))),meet(meet(join(X121,join(join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))),X122)),join(X123,join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))))),join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))))))),join(join(meet(X117,X116),meet(X116,join(X117,X116))),join(join(meet(X121,join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116))))),meet(join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))),X122)),join(meet(X117,join(X116,X116)),meet(join(meet(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)),meet(X119,join(X116,meet(meet(join(X116,join(X116,join(X116,X116))),join(X118,X116)),X116)))),join(X117,join(X116,X116)))))))),join(X116,X120)))) )),
  inference(superposition,[],[f8,f172])).
fof(f851,plain,(
  ( ! [X215,X213,X214] : (join(meet(X213,X215),meet(join(meet(X214,join(X213,X213)),meet(X213,join(X214,join(X213,X213)))),join(X213,X215))) = X213) )),
  inference(forward_demodulation,[],[f763,f750])).
fof(f763,plain,(
  ( ! [X215,X213,X214] : (join(meet(join(meet(X214,X213),meet(X213,join(X214,X213))),X215),meet(join(meet(X214,join(X213,X213)),meet(X213,join(X214,join(X213,X213)))),join(join(meet(X214,X213),meet(X213,join(X214,X213))),X215))) = X213) )),
  inference(backward_demodulation,[],[f750,f195])).
fof(f195,plain,(
  ( ! [X212,X215,X213,X214] : (join(meet(X212,X213),meet(X213,join(X212,X213))) = join(meet(join(meet(X214,join(meet(X212,X213),meet(X213,join(X212,X213)))),meet(join(meet(X212,X213),meet(X213,join(X212,X213))),join(X214,join(meet(X212,X213),meet(X213,join(X212,X213)))))),X215),meet(join(meet(X214,join(join(meet(X212,X213),meet(X213,join(X212,X213))),join(meet(X212,X213),meet(X213,join(X212,X213))))),meet(X213,join(X214,join(join(meet(X212,X213),meet(X213,join(X212,X213))),join(meet(X212,X213),meet(X213,join(X212,X213))))))),join(join(meet(X214,join(meet(X212,X213),meet(X213,join(X212,X213)))),meet(join(meet(X212,X213),meet(X213,join(X212,X213))),join(X214,join(meet(X212,X213),meet(X213,join(X212,X213)))))),X215)))) )),
  inference(superposition,[],[f5,f71])).
fof(f2233,plain,(
  ( ! [X0] : (meet(X0,X0) = X0) )),
  inference(superposition,[],[f1991,f2028])).
fof(f2771,plain,(
  ( ! [X37,X38,X36] : (meet(meet(join(X37,X36),join(X38,X36)),X36) = X36) )),
  inference(superposition,[],[f1935,f1991])).
fof(f3811,plain,(
  ( ! [X0,X1] : (join(X0,meet(X0,X1)) = X0) )),
  inference(forward_demodulation,[],[f3805,f1991])).
fof(f3805,plain,(
  ( ! [X0,X1] : (join(X0,X0) = join(X0,meet(X0,X1))) )),
  inference(backward_demodulation,[],[f3782,f2599])).
fof(f2599,plain,(
  ( ! [X0,X1] : (join(X0,meet(X0,X1)) = join(meet(join(X0,meet(X0,X1)),X0),meet(join(X0,meet(X0,X1)),X0))) )),
  inference(superposition,[],[f172,f2397])).
fof(f2397,plain,(
  ( ! [X0,X1] : (join(join(X0,meet(X0,X1)),X0) = X0) )),
  inference(superposition,[],[f1992,f2233])).
fof(f1992,plain,(
  ( ! [X4,X5,X3] : (join(join(meet(X3,X4),meet(X4,X5)),X4) = X4) )),
  inference(superposition,[],[f1948,f709])).
fof(f3782,plain,(
  ( ! [X0,X1] : (meet(join(X0,X1),X0) = X0) )),
  inference(superposition,[],[f3639,f3171])).
fof(f3171,plain,(
  ( ! [X2,X3] : (meet(X2,X3) = meet(meet(X2,X3),X3)) )),
  inference(forward_demodulation,[],[f3164,f3093])).
fof(f3093,plain,(
  ( ! [X38,X42] : (meet(X42,X38) = join(meet(meet(X42,X38),X38),meet(X42,X38))) )),
  inference(forward_demodulation,[],[f3092,f2995])).
fof(f2995,plain,(
  ( ! [X4,X5] : (meet(meet(X5,X4),X4) = meet(meet(X5,X4),meet(meet(X5,X4),X4))) )),
  inference(superposition,[],[f2930,f2244])).
fof(f2244,plain,(
  ( ! [X23,X27] : (meet(X27,X23) = join(meet(X23,meet(X27,X23)),meet(meet(X27,X23),X23))) )),
  inference(backward_demodulation,[],[f2233,f2031])).
fof(f2031,plain,(
  ( ! [X23,X27] : (meet(X27,X23) = join(meet(meet(X23,X23),meet(X27,X23)),meet(meet(X27,X23),X23))) )),
  inference(backward_demodulation,[],[f1991,f1952])).
fof(f1952,plain,(
  ( ! [X23,X27] : (meet(X27,join(X23,X23)) = join(meet(meet(X23,X23),meet(X27,join(X23,X23))),meet(meet(X27,join(X23,X23)),X23))) )),
  inference(forward_demodulation,[],[f1940,f1935])).
fof(f1940,plain,(
  ( ! [X26,X24,X23,X27,X25] : (meet(X27,join(X23,meet(meet(join(X24,join(X23,X25)),join(X26,X23)),X23))) = join(meet(meet(X23,meet(meet(join(X24,join(X23,X25)),join(X26,X23)),X23)),meet(X27,join(X23,meet(meet(join(X24,join(X23,X25)),join(X26,X23)),X23)))),meet(meet(X27,join(X23,meet(meet(join(X24,join(X23,X25)),join(X26,X23)),X23))),X23))) )),
  inference(superposition,[],[f750,f666])).
fof(f2930,plain,(
  ( ! [X0,X1] : (meet(join(X0,X1),X1) = X1) )),
  inference(superposition,[],[f2771,f2233])).
fof(f3092,plain,(
  ( ! [X38,X42] : (meet(X42,X38) = join(meet(meet(X42,X38),meet(meet(X42,X38),X38)),meet(X42,X38))) )),
  inference(forward_demodulation,[],[f3077,f2233])).
fof(f3077,plain,(
  ( ! [X38,X42] : (meet(X42,X38) = join(meet(meet(X42,X38),meet(meet(X42,X38),X38)),meet(meet(X42,X38),meet(X42,X38)))) )),
  inference(backward_demodulation,[],[f3003,f2247])).
fof(f2247,plain,(
  ( ! [X38,X42] : (meet(X38,meet(X42,X38)) = join(meet(meet(X38,meet(X42,X38)),meet(meet(X42,X38),X38)),meet(meet(X38,meet(X42,X38)),meet(X42,X38)))) )),
  inference(backward_demodulation,[],[f2233,f2034])).
fof(f2034,plain,(
  ( ! [X38,X42] : (meet(meet(X38,X38),meet(X42,X38)) = join(meet(meet(meet(X38,X38),meet(X42,X38)),meet(meet(X42,X38),X38)),meet(meet(meet(X38,X38),meet(X42,X38)),meet(X42,X38)))) )),
  inference(backward_demodulation,[],[f1991,f1955])).
fof(f1955,plain,(
  ( ! [X38,X42] : (meet(meet(X38,X38),meet(X42,join(X38,X38))) = join(meet(meet(meet(X38,X38),meet(X42,join(X38,X38))),meet(meet(X42,join(X38,X38)),X38)),meet(meet(meet(X38,X38),meet(X42,join(X38,X38))),meet(X42,join(X38,X38))))) )),
  inference(forward_demodulation,[],[f1943,f1935])).
fof(f1943,plain,(
  ( ! [X39,X41,X38,X42,X40] : (meet(meet(X38,meet(meet(join(X39,join(X38,X40)),join(X41,X38)),X38)),meet(X42,join(X38,meet(meet(join(X39,join(X38,X40)),join(X41,X38)),X38)))) = join(meet(meet(meet(X38,meet(meet(join(X39,join(X38,X40)),join(X41,X38)),X38)),meet(X42,join(X38,meet(meet(join(X39,join(X38,X40)),join(X41,X38)),X38)))),meet(meet(X42,join(X38,meet(meet(join(X39,join(X38,X40)),join(X41,X38)),X38))),X38)),meet(meet(meet(X38,meet(meet(join(X39,join(X38,X40)),join(X41,X38)),X38)),meet(X42,join(X38,meet(meet(join(X39,join(X38,X40)),join(X41,X38)),X38)))),meet(X42,join(X38,meet(meet(join(X39,join(X38,X40)),join(X41,X38)),X38)))))) )),
  inference(superposition,[],[f1613,f666])).
fof(f1613,plain,(
  ( ! [X4,X3] : (meet(X3,X4) = join(meet(meet(X3,X4),meet(X4,join(X3,X4))),meet(meet(X3,X4),X4))) )),
  inference(superposition,[],[f172,f750])).
fof(f3003,plain,(
  ( ! [X19,X20] : (meet(X20,X19) = meet(X19,meet(X20,X19))) )),
  inference(superposition,[],[f2930,f2241])).
fof(f3164,plain,(
  ( ! [X2,X3] : (meet(meet(X2,X3),X3) = join(meet(meet(X2,X3),X3),meet(X2,X3))) )),
  inference(backward_demodulation,[],[f3137,f2528])).
fof(f2528,plain,(
  ( ! [X2,X3] : (meet(meet(X2,X3),X3) = join(meet(meet(X2,X3),X3),meet(meet(meet(X2,X3),X3),meet(X2,X3)))) )),
  inference(forward_demodulation,[],[f2523,f2233])).
fof(f2523,plain,(
  ( ! [X2,X3] : (meet(meet(X2,X3),X3) = join(meet(meet(meet(X2,X3),X3),meet(meet(X2,X3),X3)),meet(meet(meet(X2,X3),X3),meet(X2,X3)))) )),
  inference(superposition,[],[f304,f2405])).
fof(f2405,plain,(
  ( ! [X15,X16] : (join(meet(X15,X16),X16) = X16) )),
  inference(superposition,[],[f1992,f2241])).
fof(f304,plain,(
  ( ! [X59,X60] : (meet(X59,X60) = join(meet(meet(X59,X60),meet(X59,join(X59,X60))),meet(meet(X59,X60),X59))) )),
  inference(superposition,[],[f172,f172])).
fof(f3137,plain,(
  ( ! [X15,X16] : (meet(X16,X15) = meet(meet(meet(X16,X15),X15),meet(X16,X15))) )),
  inference(superposition,[],[f2910,f2241])).
fof(f2910,plain,(
  ( ! [X20,X18] : (meet(meet(X18,join(X20,X18)),X18) = X18) )),
  inference(superposition,[],[f2771,f2281])).
fof(f2281,plain,(
  ( ! [X17,X13] : (join(meet(X13,meet(X17,X13)),X13) = X13) )),
  inference(forward_demodulation,[],[f2242,f2233])).
fof(f2242,plain,(
  ( ! [X17,X13] : (join(meet(X13,meet(X17,X13)),meet(X13,X13)) = X13) )),
  inference(backward_demodulation,[],[f2233,f2029])).
fof(f2029,plain,(
  ( ! [X17,X13] : (meet(X13,X13) = join(meet(meet(X13,X13),meet(X17,X13)),meet(meet(X13,X13),X13))) )),
  inference(backward_demodulation,[],[f1991,f1950])).
fof(f1950,plain,(
  ( ! [X17,X13] : (meet(X13,X13) = join(meet(meet(X13,X13),meet(X17,join(X13,X13))),meet(meet(X13,X13),X13))) )),
  inference(forward_demodulation,[],[f1938,f1935])).
fof(f1938,plain,(
  ( ! [X14,X17,X15,X13,X16] : (meet(X13,meet(meet(join(X14,join(X13,X15)),join(X16,X13)),X13)) = join(meet(meet(X13,meet(meet(join(X14,join(X13,X15)),join(X16,X13)),X13)),meet(X17,join(X13,meet(meet(join(X14,join(X13,X15)),join(X16,X13)),X13)))),meet(meet(X13,meet(meet(join(X14,join(X13,X15)),join(X16,X13)),X13)),X13))) )),
  inference(superposition,[],[f172,f666])).
fof(f3639,plain,(
  ( ! [X17,X16] : (meet(meet(join(X16,X17),X16),X16) = X16) )),
  inference(superposition,[],[f2775,f2405])).
fof(f2775,plain,(
  ( ! [X0,X3,X1] : (meet(meet(join(X0,X1),join(X3,X0)),X0) = X0) )),
  inference(superposition,[],[f1935,f2281])).
fof(f7934,plain,(
  ( ! [X74,X73] : (meet(X73,X74) = join(meet(X74,X73),meet(X73,X74))) )),
  inference(superposition,[],[f4156,f7143])).
fof(f4156,plain,(
  ( ! [X2,X1] : (join(meet(X1,X2),X1) = X1) )),
  inference(forward_demodulation,[],[f4155,f3999])).
fof(f3999,plain,(
  ( ! [X2,X3] : (meet(X2,X3) = meet(X2,meet(X2,X3))) )),
  inference(superposition,[],[f3782,f172])).
fof(f4155,plain,(
  ( ! [X2,X1] : (join(meet(X1,meet(X1,X2)),X1) = X1) )),
  inference(forward_demodulation,[],[f4108,f2233])).
fof(f4108,plain,(
  ( ! [X2,X1] : (join(meet(X1,meet(X1,X2)),meet(X1,X1)) = X1) )),
  inference(superposition,[],[f172,f3811])).
fof(f2,axiom,(
  meet(a,b) != meet(b,a)),
  file('LAT081-1.tptp',unknown)).
% SZS output end Proof for LAT081-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 2421453
% Time elapsed: 15.836 s
% ------------------------------
% ------------------------------
