% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for LCL202-1
% SZS output start Proof for LCL202-1
fof(f2593,plain,(
  $false),
  inference(subsumption_resolution,[],[f2582,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (axiom(or(not(X0),or(X1,X0)))) )),
  file('LCL202-1.tptp',unknown)).
fof(f2582,plain,(
  ~axiom(or(not(q),or(p,q)))),
  inference(resolution,[],[f429,f69])).
fof(f69,plain,(
  ~theorem(or(not(not(or(p,q))),not(q)))),
  inference(resolution,[],[f40,f9])).
fof(f9,axiom,(
  ~theorem(or(not(not(or(p,q))),or(p,not(q))))),
  file('LCL202-1.tptp',unknown)).
fof(f40,plain,(
  ( ! [X10,X11,X9] : (theorem(or(X9,or(X10,X11))) | ~theorem(or(X9,X11))) )),
  inference(resolution,[],[f18,f15])).
fof(f15,plain,(
  ( ! [X4,X3] : (theorem(or(X4,X3)) | ~theorem(X3)) )),
  inference(resolution,[],[f7,f2])).
fof(f7,axiom,(
  ( ! [X4,X3] : (~axiom(or(not(X4),X3)) | ~theorem(X4) | theorem(X3)) )),
  file('LCL202-1.tptp',unknown)).
fof(f18,plain,(
  ( ! [X2,X0,X1] : (~theorem(or(X0,or(X1,X2))) | theorem(or(X1,or(X0,X2)))) )),
  inference(resolution,[],[f4,f7])).
fof(f4,axiom,(
  ( ! [X2,X0,X1] : (axiom(or(not(or(X0,or(X1,X2))),or(X1,or(X0,X2))))) )),
  file('LCL202-1.tptp',unknown)).
fof(f429,plain,(
  ( ! [X10,X11] : (theorem(or(not(not(X11)),not(X10))) | ~axiom(or(not(X10),X11))) )),
  inference(resolution,[],[f65,f14])).
fof(f14,plain,(
  ( ! [X2,X1] : (~theorem(or(X1,X2)) | theorem(or(X2,X1))) )),
  inference(resolution,[],[f7,f3])).
fof(f3,axiom,(
  ( ! [X0,X1] : (axiom(or(not(or(X0,X1)),or(X1,X0)))) )),
  file('LCL202-1.tptp',unknown)).
fof(f65,plain,(
  ( ! [X2,X1] : (theorem(or(not(X1),not(not(X2)))) | ~axiom(or(not(X1),X2))) )),
  inference(resolution,[],[f59,f8])).
fof(f8,axiom,(
  ( ! [X4,X5,X3] : (~theorem(or(not(X4),X5)) | ~axiom(or(not(X3),X4)) | theorem(or(not(X3),X5))) )),
  file('LCL202-1.tptp',unknown)).
fof(f59,plain,(
  ( ! [X2] : (theorem(or(X2,not(X2)))) )),
  inference(resolution,[],[f54,f14])).
fof(f54,plain,(
  ( ! [X2] : (theorem(or(not(X2),X2))) )),
  inference(resolution,[],[f36,f13])).
fof(f13,plain,(
  ( ! [X0] : (~theorem(or(X0,X0)) | theorem(X0)) )),
  inference(resolution,[],[f7,f1])).
fof(f1,axiom,(
  ( ! [X0] : (axiom(or(not(or(X0,X0)),X0))) )),
  file('LCL202-1.tptp',unknown)).
fof(f36,plain,(
  ( ! [X0,X1] : (theorem(or(X0,or(not(X1),X1)))) )),
  inference(resolution,[],[f18,f11])).
fof(f11,plain,(
  ( ! [X0,X1] : (theorem(or(not(X0),or(X1,X0)))) )),
  inference(resolution,[],[f2,f6])).
fof(f6,axiom,(
  ( ! [X3] : (~axiom(X3) | theorem(X3)) )),
  file('LCL202-1.tptp',unknown)).
% SZS output end Proof for LCL202-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7419
% Time elapsed: 0.074 s
% ------------------------------
% ------------------------------
