% Refutation found. Thanks to Tanya!
% SZS status Theorem for MGT001+1
% SZS output start Proof for MGT001+1
fof(f355,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f176,f209,f217,f232,f237,f280,f288,f347,f354])).
fof(f354,plain,(
  ~spl11_30),
  inference(avatar_contradiction_clause,[],[f353])).
fof(f353,plain,(
  $false | ~spl11_30),
  inference(subsumption_resolution,[],[f352,f46])).
fof(f46,plain,(
  organization(sK3,sK5)),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ~greater(sK10,sK9) & greater(sK8,sK7) & survival_chance(sK4,sK10,sK6) & survival_chance(sK3,sK9,sK5) & inertia(sK4,sK8,sK6) & inertia(sK3,sK7,sK5) & reorganization_free(sK4,sK6,sK6) & reorganization_free(sK3,sK5,sK5) & organization(sK4,sK6) & organization(sK3,sK5)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f25,f35])).
fof(f35,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~greater(X7,X6) & greater(X5,X4) & survival_chance(X1,X7,X3) & survival_chance(X0,X6,X2) & inertia(X1,X5,X3) & inertia(X0,X4,X2) & reorganization_free(X1,X3,X3) & reorganization_free(X0,X2,X2) & organization(X1,X3) & organization(X0,X2)) => (~greater(sK10,sK9) & greater(sK8,sK7) & survival_chance(sK4,sK10,sK6) & survival_chance(sK3,sK9,sK5) & inertia(sK4,sK8,sK6) & inertia(sK3,sK7,sK5) & reorganization_free(sK4,sK6,sK6) & reorganization_free(sK3,sK5,sK5) & organization(sK4,sK6) & organization(sK3,sK5))),
  introduced(choice_axiom,[])).
fof(f25,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~greater(X7,X6) & greater(X5,X4) & survival_chance(X1,X7,X3) & survival_chance(X0,X6,X2) & inertia(X1,X5,X3) & inertia(X0,X4,X2) & reorganization_free(X1,X3,X3) & reorganization_free(X0,X2,X2) & organization(X1,X3) & organization(X0,X2))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~greater(X7,X6) & (greater(X5,X4) & survival_chance(X1,X7,X3) & survival_chance(X0,X6,X2) & inertia(X1,X5,X3) & inertia(X0,X4,X2) & reorganization_free(X1,X3,X3) & reorganization_free(X0,X2,X2) & organization(X1,X3) & organization(X0,X2)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : ((greater(X5,X4) & survival_chance(X1,X7,X3) & survival_chance(X0,X6,X2) & inertia(X1,X5,X3) & inertia(X0,X4,X2) & reorganization_free(X1,X3,X3) & reorganization_free(X0,X2,X2) & organization(X1,X3) & organization(X0,X2)) => greater(X7,X6))),
  inference(rectify,[],[f8])).
fof(f8,negated_conjecture,(
  ~! [X0,X5,X12,X13,X16,X17,X10,X11] : ((greater(X17,X16) & survival_chance(X5,X11,X13) & survival_chance(X0,X10,X12) & inertia(X5,X17,X13) & inertia(X0,X16,X12) & reorganization_free(X5,X13,X13) & reorganization_free(X0,X12,X12) & organization(X5,X13) & organization(X0,X12)) => greater(X11,X10))),
  inference(negated_conjecture,[],[f7])).
fof(f7,conjecture,(
  ! [X0,X5,X12,X13,X16,X17,X10,X11] : ((greater(X17,X16) & survival_chance(X5,X11,X13) & survival_chance(X0,X10,X12) & inertia(X5,X17,X13) & inertia(X0,X16,X12) & reorganization_free(X5,X13,X13) & reorganization_free(X0,X12,X12) & organization(X5,X13) & organization(X0,X12)) => greater(X11,X10))),
  file('MGT001+1.tptp',unknown)).
fof(f352,plain,(
  ~organization(sK3,sK5) | ~spl11_30),
  inference(resolution,[],[f276,f38])).
fof(f38,plain,(
  ( ! [X0,X1] : (accountability(X0,sK1(X0,X1),X1) | ~organization(X0,X1)) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : (accountability(X0,sK1(X0,X1),X1) | ~organization(X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f16,f28])).
fof(f28,plain,(
  ! [X1,X0] : (? [X2] : accountability(X0,X2,X1) => accountability(X0,sK1(X0,X1),X1))),
  introduced(choice_axiom,[])).
fof(f16,plain,(
  ! [X0,X1] : (? [X2] : accountability(X0,X2,X1) | ~organization(X0,X1))),
  inference(ennf_transformation,[],[f9])).
fof(f9,plain,(
  ! [X0,X1] : (organization(X0,X1) => ? [X2] : accountability(X0,X2,X1))),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1] : (organization(X0,X1) => ? [X3] : accountability(X0,X3,X1))),
  file('MGT001+1.tptp',unknown)).
fof(f276,plain,(
  ( ! [X1] : (~accountability(sK3,X1,sK5)) ) | ~spl11_30),
  inference(avatar_component_clause,[],[f275])).
fof(f275,plain,(
  spl11_30 <=> ! [X1] : ~accountability(sK3,X1,sK5)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_30])])).
fof(f347,plain,(
  spl11_20 | spl11_30 | spl11_31 | ~spl11_16 | ~spl11_23),
  inference(avatar_split_clause,[],[f346,f215,f174,f278,f275,f204])).
fof(f204,plain,(
  spl11_20 <=> ! [X5] : ~reliability(sK4,X5,sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_20])])).
fof(f278,plain,(
  spl11_31 <=> ! [X0] : ~reliability(sK3,X0,sK5)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_31])])).
fof(f174,plain,(
  spl11_16 <=> ! [X1,X0] : (~greater(X0,X1) | ~reliability(sK3,X1,sK5) | ~reliability(sK4,X0,sK6))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_16])])).
fof(f215,plain,(
  spl11_23 <=> ! [X9,X11,X13,X7,X8,X10] : (~reproducibility(sK3,X7,sK5) | ~organization(X10,X11) | ~reliability(X10,X9,X11) | ~accountability(X10,X13,X11) | ~reproducibility(X10,X7,X11) | greater(X8,X9) | ~reliability(sK4,X8,sK6))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_23])])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (~reliability(sK3,X0,sK5) | ~accountability(sK3,X1,sK5) | ~reliability(sK4,X2,sK6)) ) | (~spl11_16 | ~spl11_23)),
  inference(subsumption_resolution,[],[f345,f175])).
fof(f175,plain,(
  ( ! [X0,X1] : (~reliability(sK4,X0,sK6) | ~reliability(sK3,X1,sK5) | ~greater(X0,X1)) ) | ~spl11_16),
  inference(avatar_component_clause,[],[f174])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (~reliability(sK3,X0,sK5) | ~accountability(sK3,X1,sK5) | greater(X2,X0) | ~reliability(sK4,X2,sK6)) ) | ~spl11_23),
  inference(subsumption_resolution,[],[f344,f46])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (~reliability(sK3,X0,sK5) | ~accountability(sK3,X1,sK5) | ~organization(sK3,sK5) | greater(X2,X0) | ~reliability(sK4,X2,sK6)) ) | ~spl11_23),
  inference(duplicate_literal_removal,[],[f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (~reliability(sK3,X0,sK5) | ~accountability(sK3,X1,sK5) | ~organization(sK3,sK5) | greater(X2,X0) | ~reliability(sK4,X2,sK6) | ~organization(sK3,sK5)) ) | ~spl11_23),
  inference(resolution,[],[f301,f39])).
fof(f39,plain,(
  ( ! [X0,X1] : (reproducibility(X0,sK2(X0,X1),X1) | ~organization(X0,X1)) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (reproducibility(X0,sK2(X0,X1),X1) | ~organization(X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f17,f30])).
fof(f30,plain,(
  ! [X1,X0] : (? [X2] : reproducibility(X0,X2,X1) => reproducibility(X0,sK2(X0,X1),X1))),
  introduced(choice_axiom,[])).
fof(f17,plain,(
  ! [X0,X1] : (? [X2] : reproducibility(X0,X2,X1) | ~organization(X0,X1))),
  inference(ennf_transformation,[],[f10])).
fof(f10,plain,(
  ! [X0,X1] : (organization(X0,X1) => ? [X2] : reproducibility(X0,X2,X1))),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1] : (organization(X0,X1) => ? [X4] : reproducibility(X0,X4,X1))),
  file('MGT001+1.tptp',unknown)).
fof(f301,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~reproducibility(X0,sK2(sK3,sK5),X1) | ~reliability(X0,X2,X1) | ~accountability(X0,X3,X1) | ~organization(X0,X1) | greater(X4,X2) | ~reliability(sK4,X4,sK6)) ) | ~spl11_23),
  inference(subsumption_resolution,[],[f300,f46])).
fof(f300,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~organization(X0,X1) | ~reliability(X0,X2,X1) | ~accountability(X0,X3,X1) | ~reproducibility(X0,sK2(sK3,sK5),X1) | greater(X4,X2) | ~reliability(sK4,X4,sK6) | ~organization(sK3,sK5)) ) | ~spl11_23),
  inference(resolution,[],[f216,f39])).
fof(f216,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (~reproducibility(sK3,X7,sK5) | ~organization(X10,X11) | ~reliability(X10,X9,X11) | ~accountability(X10,X13,X11) | ~reproducibility(X10,X7,X11) | greater(X8,X9) | ~reliability(sK4,X8,sK6)) ) | ~spl11_23),
  inference(avatar_component_clause,[],[f215])).
fof(f288,plain,(
  ~spl11_31),
  inference(avatar_contradiction_clause,[],[f287])).
fof(f287,plain,(
  $false | ~spl11_31),
  inference(subsumption_resolution,[],[f286,f46])).
fof(f286,plain,(
  ~organization(sK3,sK5) | ~spl11_31),
  inference(resolution,[],[f279,f37])).
fof(f37,plain,(
  ( ! [X0,X1] : (reliability(X0,sK0(X0,X1),X1) | ~organization(X0,X1)) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : (reliability(X0,sK0(X0,X1),X1) | ~organization(X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f15,f26])).
fof(f26,plain,(
  ! [X1,X0] : (? [X2] : reliability(X0,X2,X1) => reliability(X0,sK0(X0,X1),X1))),
  introduced(choice_axiom,[])).
fof(f15,plain,(
  ! [X0,X1] : (? [X2] : reliability(X0,X2,X1) | ~organization(X0,X1))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : (organization(X0,X1) => ? [X2] : reliability(X0,X2,X1))),
  file('MGT001+1.tptp',unknown)).
fof(f279,plain,(
  ( ! [X0] : (~reliability(sK3,X0,sK5)) ) | ~spl11_31),
  inference(avatar_component_clause,[],[f278])).
fof(f280,plain,(
  spl11_22 | spl11_30 | spl11_31 | ~spl11_15 | ~spl11_21),
  inference(avatar_split_clause,[],[f273,f207,f171,f278,f275,f212])).
fof(f212,plain,(
  spl11_22 <=> ! [X12] : ~accountability(sK4,X12,sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_22])])).
fof(f171,plain,(
  spl11_15 <=> ! [X3,X2] : (~greater(X2,X3) | ~accountability(sK3,X3,sK5) | ~accountability(sK4,X2,sK6))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_15])])).
fof(f207,plain,(
  spl11_21 <=> ! [X1,X3,X0,X2,X4,X6] : (~reproducibility(sK3,X0,sK5) | ~organization(X3,X4) | ~reliability(X3,X6,X4) | ~accountability(X3,X2,X4) | ~reproducibility(X3,X0,X4) | greater(X1,X2) | ~accountability(sK4,X1,sK6))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_21])])).
fof(f273,plain,(
  ( ! [X2,X0,X1] : (~reliability(sK3,X0,sK5) | ~accountability(sK3,X1,sK5) | ~accountability(sK4,X2,sK6)) ) | (~spl11_15 | ~spl11_21)),
  inference(subsumption_resolution,[],[f272,f172])).
fof(f172,plain,(
  ( ! [X2,X3] : (~accountability(sK4,X2,sK6) | ~accountability(sK3,X3,sK5) | ~greater(X2,X3)) ) | ~spl11_15),
  inference(avatar_component_clause,[],[f171])).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (~reliability(sK3,X0,sK5) | ~accountability(sK3,X1,sK5) | greater(X2,X1) | ~accountability(sK4,X2,sK6)) ) | ~spl11_21),
  inference(subsumption_resolution,[],[f271,f46])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (~reliability(sK3,X0,sK5) | ~accountability(sK3,X1,sK5) | ~organization(sK3,sK5) | greater(X2,X1) | ~accountability(sK4,X2,sK6)) ) | ~spl11_21),
  inference(duplicate_literal_removal,[],[f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (~reliability(sK3,X0,sK5) | ~accountability(sK3,X1,sK5) | ~organization(sK3,sK5) | greater(X2,X1) | ~accountability(sK4,X2,sK6) | ~organization(sK3,sK5)) ) | ~spl11_21),
  inference(resolution,[],[f268,f39])).
fof(f268,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~reproducibility(X0,sK2(sK3,sK5),X1) | ~reliability(X0,X2,X1) | ~accountability(X0,X3,X1) | ~organization(X0,X1) | greater(X4,X3) | ~accountability(sK4,X4,sK6)) ) | ~spl11_21),
  inference(subsumption_resolution,[],[f267,f46])).
fof(f267,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~organization(X0,X1) | ~reliability(X0,X2,X1) | ~accountability(X0,X3,X1) | ~reproducibility(X0,sK2(sK3,sK5),X1) | greater(X4,X3) | ~accountability(sK4,X4,sK6) | ~organization(sK3,sK5)) ) | ~spl11_21),
  inference(resolution,[],[f208,f39])).
fof(f208,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (~reproducibility(sK3,X0,sK5) | ~organization(X3,X4) | ~reliability(X3,X6,X4) | ~accountability(X3,X2,X4) | ~reproducibility(X3,X0,X4) | greater(X1,X2) | ~accountability(sK4,X1,sK6)) ) | ~spl11_21),
  inference(avatar_component_clause,[],[f207])).
fof(f237,plain,(
  ~spl11_22),
  inference(avatar_contradiction_clause,[],[f236])).
fof(f236,plain,(
  $false | ~spl11_22),
  inference(subsumption_resolution,[],[f235,f47])).
fof(f47,plain,(
  organization(sK4,sK6)),
  inference(cnf_transformation,[],[f36])).
fof(f235,plain,(
  ~organization(sK4,sK6) | ~spl11_22),
  inference(resolution,[],[f213,f38])).
fof(f213,plain,(
  ( ! [X12] : (~accountability(sK4,X12,sK6)) ) | ~spl11_22),
  inference(avatar_component_clause,[],[f212])).
fof(f232,plain,(
  ~spl11_20),
  inference(avatar_contradiction_clause,[],[f231])).
fof(f231,plain,(
  $false | ~spl11_20),
  inference(subsumption_resolution,[],[f230,f47])).
fof(f230,plain,(
  ~organization(sK4,sK6) | ~spl11_20),
  inference(resolution,[],[f205,f37])).
fof(f205,plain,(
  ( ! [X5] : (~reliability(sK4,X5,sK6)) ) | ~spl11_20),
  inference(avatar_component_clause,[],[f204])).
fof(f217,plain,(
  spl11_22 | spl11_23),
  inference(avatar_split_clause,[],[f210,f215,f212])).
fof(f210,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~reproducibility(sK3,X7,sK5) | greater(X8,X9) | ~reproducibility(X10,X7,X11) | ~accountability(sK4,X12,sK6) | ~accountability(X10,X13,X11) | ~reliability(sK4,X8,sK6) | ~reliability(X10,X9,X11) | ~organization(X10,X11)) )),
  inference(subsumption_resolution,[],[f198,f47])).
fof(f198,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~reproducibility(sK3,X7,sK5) | greater(X8,X9) | ~reproducibility(X10,X7,X11) | ~accountability(sK4,X12,sK6) | ~accountability(X10,X13,X11) | ~reliability(sK4,X8,sK6) | ~reliability(X10,X9,X11) | ~organization(sK4,sK6) | ~organization(X10,X11)) )),
  inference(resolution,[],[f196,f57])).
fof(f57,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~greater(sK2(X0,X1),X2) | greater(X3,X4) | ~reproducibility(X5,X2,X6) | ~accountability(X0,X7,X1) | ~accountability(X5,X8,X6) | ~reliability(X0,X3,X1) | ~reliability(X5,X4,X6) | ~organization(X0,X1) | ~organization(X5,X6)) )),
  inference(duplicate_literal_removal,[],[f56])).
fof(f56,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~greater(sK2(X0,X1),X2) | greater(X3,X4) | ~reproducibility(X5,X2,X6) | ~accountability(X0,X7,X1) | ~accountability(X5,X8,X6) | ~reliability(X0,X3,X1) | ~reliability(X5,X4,X6) | ~organization(X0,X1) | ~organization(X5,X6) | ~organization(X0,X1)) )),
  inference(resolution,[],[f41,f39])).
fof(f41,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~reproducibility(X1,X9,X3) | ~greater(X9,X8) | greater(X5,X4) | ~reproducibility(X0,X8,X2) | ~accountability(X1,X7,X3) | ~accountability(X0,X6,X2) | ~reliability(X1,X5,X3) | ~reliability(X0,X4,X2) | ~organization(X1,X3) | ~organization(X0,X2)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (((greater(X9,X8) | ~greater(X7,X6) | ~greater(X5,X4)) & ((greater(X7,X6) & greater(X5,X4)) | ~greater(X9,X8))) | ~reproducibility(X1,X9,X3) | ~reproducibility(X0,X8,X2) | ~accountability(X1,X7,X3) | ~accountability(X0,X6,X2) | ~reliability(X1,X5,X3) | ~reliability(X0,X4,X2) | ~organization(X1,X3) | ~organization(X0,X2))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (((greater(X9,X8) | (~greater(X7,X6) | ~greater(X5,X4))) & ((greater(X7,X6) & greater(X5,X4)) | ~greater(X9,X8))) | ~reproducibility(X1,X9,X3) | ~reproducibility(X0,X8,X2) | ~accountability(X1,X7,X3) | ~accountability(X0,X6,X2) | ~reliability(X1,X5,X3) | ~reliability(X0,X4,X2) | ~organization(X1,X3) | ~organization(X0,X2))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((greater(X9,X8) <=> (greater(X7,X6) & greater(X5,X4))) | ~reproducibility(X1,X9,X3) | ~reproducibility(X0,X8,X2) | ~accountability(X1,X7,X3) | ~accountability(X0,X6,X2) | ~reliability(X1,X5,X3) | ~reliability(X0,X4,X2) | ~organization(X1,X3) | ~organization(X0,X2))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((greater(X9,X8) <=> (greater(X7,X6) & greater(X5,X4))) | (~reproducibility(X1,X9,X3) | ~reproducibility(X0,X8,X2) | ~accountability(X1,X7,X3) | ~accountability(X0,X6,X2) | ~reliability(X1,X5,X3) | ~reliability(X0,X4,X2) | ~organization(X1,X3) | ~organization(X0,X2)))),
  inference(ennf_transformation,[],[f12])).
fof(f12,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((reproducibility(X1,X9,X3) & reproducibility(X0,X8,X2) & accountability(X1,X7,X3) & accountability(X0,X6,X2) & reliability(X1,X5,X3) & reliability(X0,X4,X2) & organization(X1,X3) & organization(X0,X2)) => (greater(X9,X8) <=> (greater(X7,X6) & greater(X5,X4))))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X0,X5,X12,X13,X6,X7,X8,X9,X14,X15] : ((reproducibility(X5,X15,X13) & reproducibility(X0,X14,X12) & accountability(X5,X9,X13) & accountability(X0,X8,X12) & reliability(X5,X7,X13) & reliability(X0,X6,X12) & organization(X5,X13) & organization(X0,X12)) => (greater(X15,X14) <=> (greater(X9,X8) & greater(X7,X6))))),
  file('MGT001+1.tptp',unknown)).
fof(f196,plain,(
  ( ! [X0] : (greater(sK2(sK4,sK6),X0) | ~reproducibility(sK3,X0,sK5)) )),
  inference(subsumption_resolution,[],[f195,f47])).
fof(f195,plain,(
  ( ! [X0] : (greater(sK2(sK4,sK6),X0) | ~reproducibility(sK3,X0,sK5) | ~organization(sK4,sK6)) )),
  inference(resolution,[],[f135,f39])).
fof(f135,plain,(
  ( ! [X0,X1] : (~reproducibility(sK4,X0,sK6) | greater(X0,X1) | ~reproducibility(sK3,X1,sK5)) )),
  inference(subsumption_resolution,[],[f134,f46])).
fof(f134,plain,(
  ( ! [X0,X1] : (greater(X0,X1) | ~reproducibility(sK4,X0,sK6) | ~reproducibility(sK3,X1,sK5) | ~organization(sK3,sK5)) )),
  inference(subsumption_resolution,[],[f133,f48])).
fof(f48,plain,(
  reorganization_free(sK3,sK5,sK5)),
  inference(cnf_transformation,[],[f36])).
fof(f133,plain,(
  ( ! [X0,X1] : (greater(X0,X1) | ~reproducibility(sK4,X0,sK6) | ~reproducibility(sK3,X1,sK5) | ~reorganization_free(sK3,sK5,sK5) | ~organization(sK3,sK5)) )),
  inference(subsumption_resolution,[],[f131,f54])).
fof(f54,plain,(
  greater(sK8,sK7)),
  inference(cnf_transformation,[],[f36])).
fof(f131,plain,(
  ( ! [X0,X1] : (greater(X0,X1) | ~greater(sK8,sK7) | ~reproducibility(sK4,X0,sK6) | ~reproducibility(sK3,X1,sK5) | ~reorganization_free(sK3,sK5,sK5) | ~organization(sK3,sK5)) )),
  inference(resolution,[],[f93,f50])).
fof(f50,plain,(
  inertia(sK3,sK7,sK5)),
  inference(cnf_transformation,[],[f36])).
fof(f93,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~inertia(X8,X5,X9) | greater(X6,X7) | ~greater(sK8,X5) | ~reproducibility(sK4,X6,sK6) | ~reproducibility(X8,X7,X9) | ~reorganization_free(X8,X9,X9) | ~organization(X8,X9)) )),
  inference(subsumption_resolution,[],[f92,f47])).
fof(f92,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~greater(sK8,X5) | greater(X6,X7) | ~inertia(X8,X5,X9) | ~reproducibility(sK4,X6,sK6) | ~reproducibility(X8,X7,X9) | ~reorganization_free(X8,X9,X9) | ~organization(sK4,sK6) | ~organization(X8,X9)) )),
  inference(subsumption_resolution,[],[f89,f49])).
fof(f49,plain,(
  reorganization_free(sK4,sK6,sK6)),
  inference(cnf_transformation,[],[f36])).
fof(f89,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~greater(sK8,X5) | greater(X6,X7) | ~inertia(X8,X5,X9) | ~reproducibility(sK4,X6,sK6) | ~reproducibility(X8,X7,X9) | ~reorganization_free(sK4,sK6,sK6) | ~reorganization_free(X8,X9,X9) | ~organization(sK4,sK6) | ~organization(X8,X9)) )),
  inference(resolution,[],[f45,f51])).
fof(f51,plain,(
  inertia(sK4,sK8,sK6)),
  inference(cnf_transformation,[],[f36])).
fof(f45,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~inertia(X1,X7,X3) | ~greater(X7,X6) | greater(X5,X4) | ~inertia(X0,X6,X2) | ~reproducibility(X1,X5,X3) | ~reproducibility(X0,X4,X2) | ~reorganization_free(X1,X3,X3) | ~reorganization_free(X0,X2,X2) | ~organization(X1,X3) | ~organization(X0,X2)) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (((greater(X5,X4) | ~greater(X7,X6)) & (greater(X7,X6) | ~greater(X5,X4))) | ~inertia(X1,X7,X3) | ~inertia(X0,X6,X2) | ~reproducibility(X1,X5,X3) | ~reproducibility(X0,X4,X2) | ~reorganization_free(X1,X3,X3) | ~reorganization_free(X0,X2,X2) | ~organization(X1,X3) | ~organization(X0,X2))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((greater(X5,X4) <=> greater(X7,X6)) | ~inertia(X1,X7,X3) | ~inertia(X0,X6,X2) | ~reproducibility(X1,X5,X3) | ~reproducibility(X0,X4,X2) | ~reorganization_free(X1,X3,X3) | ~reorganization_free(X0,X2,X2) | ~organization(X1,X3) | ~organization(X0,X2))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((greater(X5,X4) <=> greater(X7,X6)) | (~inertia(X1,X7,X3) | ~inertia(X0,X6,X2) | ~reproducibility(X1,X5,X3) | ~reproducibility(X0,X4,X2) | ~reorganization_free(X1,X3,X3) | ~reorganization_free(X0,X2,X2) | ~organization(X1,X3) | ~organization(X0,X2)))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((inertia(X1,X7,X3) & inertia(X0,X6,X2) & reproducibility(X1,X5,X3) & reproducibility(X0,X4,X2) & reorganization_free(X1,X3,X3) & reorganization_free(X0,X2,X2) & organization(X1,X3) & organization(X0,X2)) => (greater(X5,X4) <=> greater(X7,X6)))),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X0,X5,X12,X13,X14,X15,X16,X17] : ((inertia(X5,X17,X13) & inertia(X0,X16,X12) & reproducibility(X5,X15,X13) & reproducibility(X0,X14,X12) & reorganization_free(X5,X13,X13) & reorganization_free(X0,X12,X12) & organization(X5,X13) & organization(X0,X12)) => (greater(X15,X14) <=> greater(X17,X16)))),
  file('MGT001+1.tptp',unknown)).
fof(f209,plain,(
  spl11_20 | spl11_21),
  inference(avatar_split_clause,[],[f202,f207,f204])).
fof(f202,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~reproducibility(sK3,X0,sK5) | greater(X1,X2) | ~reproducibility(X3,X0,X4) | ~accountability(sK4,X1,sK6) | ~accountability(X3,X2,X4) | ~reliability(sK4,X5,sK6) | ~reliability(X3,X6,X4) | ~organization(X3,X4)) )),
  inference(subsumption_resolution,[],[f197,f47])).
fof(f197,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~reproducibility(sK3,X0,sK5) | greater(X1,X2) | ~reproducibility(X3,X0,X4) | ~accountability(sK4,X1,sK6) | ~accountability(X3,X2,X4) | ~reliability(sK4,X5,sK6) | ~reliability(X3,X6,X4) | ~organization(sK4,sK6) | ~organization(X3,X4)) )),
  inference(resolution,[],[f196,f59])).
fof(f59,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~greater(sK2(X0,X1),X2) | greater(X3,X4) | ~reproducibility(X5,X2,X6) | ~accountability(X0,X3,X1) | ~accountability(X5,X4,X6) | ~reliability(X0,X7,X1) | ~reliability(X5,X8,X6) | ~organization(X0,X1) | ~organization(X5,X6)) )),
  inference(duplicate_literal_removal,[],[f58])).
fof(f58,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~greater(sK2(X0,X1),X2) | greater(X3,X4) | ~reproducibility(X5,X2,X6) | ~accountability(X0,X3,X1) | ~accountability(X5,X4,X6) | ~reliability(X0,X7,X1) | ~reliability(X5,X8,X6) | ~organization(X0,X1) | ~organization(X5,X6) | ~organization(X0,X1)) )),
  inference(resolution,[],[f42,f39])).
fof(f42,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~reproducibility(X1,X9,X3) | ~greater(X9,X8) | greater(X7,X6) | ~reproducibility(X0,X8,X2) | ~accountability(X1,X7,X3) | ~accountability(X0,X6,X2) | ~reliability(X1,X5,X3) | ~reliability(X0,X4,X2) | ~organization(X1,X3) | ~organization(X0,X2)) )),
  inference(cnf_transformation,[],[f33])).
fof(f176,plain,(
  spl11_15 | spl11_16),
  inference(avatar_split_clause,[],[f169,f174,f171])).
fof(f169,plain,(
  ( ! [X2,X0,X3,X1] : (~greater(X0,X1) | ~greater(X2,X3) | ~accountability(sK4,X2,sK6) | ~accountability(sK3,X3,sK5) | ~reliability(sK4,X0,sK6) | ~reliability(sK3,X1,sK5)) )),
  inference(subsumption_resolution,[],[f168,f46])).
fof(f168,plain,(
  ( ! [X2,X0,X3,X1] : (~greater(X0,X1) | ~greater(X2,X3) | ~accountability(sK4,X2,sK6) | ~accountability(sK3,X3,sK5) | ~reliability(sK4,X0,sK6) | ~reliability(sK3,X1,sK5) | ~organization(sK3,sK5)) )),
  inference(subsumption_resolution,[],[f166,f55])).
fof(f55,plain,(
  ~greater(sK10,sK9)),
  inference(cnf_transformation,[],[f36])).
fof(f166,plain,(
  ( ! [X2,X0,X3,X1] : (~greater(X0,X1) | greater(sK10,sK9) | ~greater(X2,X3) | ~accountability(sK4,X2,sK6) | ~accountability(sK3,X3,sK5) | ~reliability(sK4,X0,sK6) | ~reliability(sK3,X1,sK5) | ~organization(sK3,sK5)) )),
  inference(resolution,[],[f97,f52])).
fof(f52,plain,(
  survival_chance(sK3,sK9,sK5)),
  inference(cnf_transformation,[],[f36])).
fof(f97,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~survival_chance(X12,X11,X13) | ~greater(X9,X10) | greater(sK10,X11) | ~greater(X7,X8) | ~accountability(sK4,X7,sK6) | ~accountability(X12,X8,X13) | ~reliability(sK4,X9,sK6) | ~reliability(X12,X10,X13) | ~organization(X12,X13)) )),
  inference(subsumption_resolution,[],[f95,f47])).
fof(f95,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~greater(X7,X8) | ~greater(X9,X10) | greater(sK10,X11) | ~survival_chance(X12,X11,X13) | ~accountability(sK4,X7,sK6) | ~accountability(X12,X8,X13) | ~reliability(sK4,X9,sK6) | ~reliability(X12,X10,X13) | ~organization(sK4,sK6) | ~organization(X12,X13)) )),
  inference(resolution,[],[f40,f53])).
fof(f53,plain,(
  survival_chance(sK4,sK10,sK6)),
  inference(cnf_transformation,[],[f36])).
fof(f40,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~survival_chance(X1,X7,X9) | ~greater(X5,X4) | ~greater(X3,X2) | greater(X7,X6) | ~survival_chance(X0,X6,X8) | ~accountability(X1,X5,X9) | ~accountability(X0,X4,X8) | ~reliability(X1,X3,X9) | ~reliability(X0,X2,X8) | ~organization(X1,X9) | ~organization(X0,X8)) )),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (greater(X7,X6) | ~greater(X5,X4) | ~greater(X3,X2) | ~survival_chance(X1,X7,X9) | ~survival_chance(X0,X6,X8) | ~accountability(X1,X5,X9) | ~accountability(X0,X4,X8) | ~reliability(X1,X3,X9) | ~reliability(X0,X2,X8) | ~organization(X1,X9) | ~organization(X0,X8))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (greater(X7,X6) | (~greater(X5,X4) | ~greater(X3,X2) | ~survival_chance(X1,X7,X9) | ~survival_chance(X0,X6,X8) | ~accountability(X1,X5,X9) | ~accountability(X0,X4,X8) | ~reliability(X1,X3,X9) | ~reliability(X0,X2,X8) | ~organization(X1,X9) | ~organization(X0,X8)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((greater(X5,X4) & greater(X3,X2) & survival_chance(X1,X7,X9) & survival_chance(X0,X6,X8) & accountability(X1,X5,X9) & accountability(X0,X4,X8) & reliability(X1,X3,X9) & reliability(X0,X2,X8) & organization(X1,X9) & organization(X0,X8)) => greater(X7,X6))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X0,X5,X6,X7,X8,X9,X10,X11,X12,X13] : ((greater(X9,X8) & greater(X7,X6) & survival_chance(X5,X11,X13) & survival_chance(X0,X10,X12) & accountability(X5,X9,X13) & accountability(X0,X8,X12) & reliability(X5,X7,X13) & reliability(X0,X6,X12) & organization(X5,X13) & organization(X0,X12)) => greater(X11,X10))),
  file('MGT001+1.tptp',unknown)).
% SZS output end Proof for MGT001+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
