% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM500+3
% SZS output start Proof for NUM500+3
fof(f1270,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f693,f714,f1269])).
fof(f1269,plain,(
  ~spl16_11),
  inference(avatar_contradiction_clause,[],[f1268])).
fof(f1268,plain,(
  $false | ~spl16_11),
  inference(subsumption_resolution,[],[f1267,f282])).
fof(f282,plain,(
  aNaturalNumber0(xk)),
  inference(cnf_transformation,[],[f45])).
fof(f45,axiom,(
  xk = sdtsldt0(sdtasdt0(xn,xm),xp) & sdtasdt0(xn,xm) = sdtasdt0(xp,xk) & aNaturalNumber0(xk)),
  file('NUM500+3.tptp',unknown)).
fof(f1267,plain,(
  ~aNaturalNumber0(xk) | ~spl16_11),
  inference(subsumption_resolution,[],[f1266,f285])).
fof(f285,plain,(
  sz00 != xk),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  sz10 != xk & sz00 != xk),
  inference(ennf_transformation,[],[f46])).
fof(f46,axiom,(
  ~(sz10 = xk | sz00 = xk)),
  file('NUM500+3.tptp',unknown)).
fof(f1266,plain,(
  sz00 = xk | ~aNaturalNumber0(xk) | ~spl16_11),
  inference(subsumption_resolution,[],[f1265,f286])).
fof(f286,plain,(
  sz10 != xk),
  inference(cnf_transformation,[],[f126])).
fof(f1265,plain,(
  sz10 = xk | sz00 = xk | ~aNaturalNumber0(xk) | ~spl16_11),
  inference(resolution,[],[f760,f242])).
fof(f242,plain,(
  ( ! [X0] : (isPrime0(sK6(X0)) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : ((isPrime0(sK6(X0)) & doDivides0(sK6(X0),X0) & aNaturalNumber0(sK6(X0))) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f119,f151])).
fof(f151,plain,(
  ! [X0] : (? [X1] : (isPrime0(X1) & doDivides0(X1,X0) & aNaturalNumber0(X1)) => (isPrime0(sK6(X0)) & doDivides0(sK6(X0),X0) & aNaturalNumber0(sK6(X0))))),
  introduced(choice_axiom,[])).
fof(f119,plain,(
  ! [X0] : (? [X1] : (isPrime0(X1) & doDivides0(X1,X0) & aNaturalNumber0(X1)) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : (? [X1] : (isPrime0(X1) & doDivides0(X1,X0) & aNaturalNumber0(X1)) | (sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,axiom,(
  ! [X0] : ((sz10 != X0 & sz00 != X0 & aNaturalNumber0(X0)) => ? [X1] : (isPrime0(X1) & doDivides0(X1,X0) & aNaturalNumber0(X1)))),
  file('NUM500+3.tptp',unknown)).
fof(f760,plain,(
  ~isPrime0(sK6(xk)) | ~spl16_11),
  inference(resolution,[],[f692,f295])).
fof(f295,plain,(
  ( ! [X0] : (~sP2(X0) | ~isPrime0(X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0] : ((~isPrime0(X0) & ((sK14(X0) != X0 & sz10 != sK14(X0) & doDivides0(sK14(X0),X0) & (sdtasdt0(sK14(X0),sK15(X0)) = X0 & aNaturalNumber0(sK15(X0))) & aNaturalNumber0(sK14(X0))) | sz10 = X0 | sz00 = X0)) | ~sP2(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f170,f172,f171])).
fof(f171,plain,(
  ! [X0] : (? [X1] : (X0 != X1 & sz10 != X1 & doDivides0(X1,X0) & ? [X2] : (sdtasdt0(X1,X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) => (sK14(X0) != X0 & sz10 != sK14(X0) & doDivides0(sK14(X0),X0) & ? [X2] : (sdtasdt0(sK14(X0),X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(sK14(X0))))),
  introduced(choice_axiom,[])).
fof(f172,plain,(
  ! [X0] : (? [X2] : (sdtasdt0(sK14(X0),X2) = X0 & aNaturalNumber0(X2)) => (sdtasdt0(sK14(X0),sK15(X0)) = X0 & aNaturalNumber0(sK15(X0))))),
  introduced(choice_axiom,[])).
fof(f170,plain,(
  ! [X0] : ((~isPrime0(X0) & (? [X1] : (X0 != X1 & sz10 != X1 & doDivides0(X1,X0) & ? [X2] : (sdtasdt0(X1,X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) | sz10 = X0 | sz00 = X0)) | ~sP2(X0))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((~isPrime0(X0) & (? [X1] : (X0 != X1 & sz10 != X1 & doDivides0(X1,X0) & ? [X2] : (sdtasdt0(X1,X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) | sz10 = X0 | sz00 = X0)) | ~sP2(X0))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f692,plain,(
  sP2(sK6(xk)) | ~spl16_11),
  inference(avatar_component_clause,[],[f690])).
fof(f690,plain,(
  spl16_11 <=> sP2(sK6(xk))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_11])])).
fof(f714,plain,(
  spl16_10),
  inference(avatar_contradiction_clause,[],[f713])).
fof(f713,plain,(
  $false | spl16_10),
  inference(subsumption_resolution,[],[f712,f282])).
fof(f712,plain,(
  ~aNaturalNumber0(xk) | spl16_10),
  inference(subsumption_resolution,[],[f711,f285])).
fof(f711,plain,(
  sz00 = xk | ~aNaturalNumber0(xk) | spl16_10),
  inference(subsumption_resolution,[],[f710,f286])).
fof(f710,plain,(
  sz10 = xk | sz00 = xk | ~aNaturalNumber0(xk) | spl16_10),
  inference(resolution,[],[f688,f240])).
fof(f240,plain,(
  ( ! [X0] : (aNaturalNumber0(sK6(X0)) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f152])).
fof(f688,plain,(
  ~aNaturalNumber0(sK6(xk)) | spl16_10),
  inference(avatar_component_clause,[],[f686])).
fof(f686,plain,(
  spl16_10 <=> aNaturalNumber0(sK6(xk))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_10])])).
fof(f693,plain,(
  ~spl16_10 | spl16_11),
  inference(avatar_split_clause,[],[f684,f690,f686])).
fof(f684,plain,(
  sP2(sK6(xk)) | ~aNaturalNumber0(sK6(xk))),
  inference(subsumption_resolution,[],[f683,f282])).
fof(f683,plain,(
  ~aNaturalNumber0(xk) | sP2(sK6(xk)) | ~aNaturalNumber0(sK6(xk))),
  inference(subsumption_resolution,[],[f682,f285])).
fof(f682,plain,(
  sz00 = xk | ~aNaturalNumber0(xk) | sP2(sK6(xk)) | ~aNaturalNumber0(sK6(xk))),
  inference(subsumption_resolution,[],[f680,f286])).
fof(f680,plain,(
  sz10 = xk | sz00 = xk | ~aNaturalNumber0(xk) | sP2(sK6(xk)) | ~aNaturalNumber0(sK6(xk))),
  inference(resolution,[],[f241,f297])).
fof(f297,plain,(
  ( ! [X0] : (~doDivides0(X0,xk) | sP2(X0) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : (sP2(X0) | (~doDivides0(X0,xk) & ! [X1] : (sdtasdt0(X0,X1) != xk | ~aNaturalNumber0(X1))) | ~aNaturalNumber0(X0))),
  inference(rectify,[],[f133])).
fof(f133,plain,(
  ! [X0] : (sP2(X0) | (~doDivides0(X0,xk) & ! [X3] : (xk != sdtasdt0(X0,X3) | ~aNaturalNumber0(X3))) | ~aNaturalNumber0(X0))),
  inference(definition_folding,[],[f128,f132])).
fof(f128,plain,(
  ! [X0] : ((~isPrime0(X0) & (? [X1] : (X0 != X1 & sz10 != X1 & doDivides0(X1,X0) & ? [X2] : (sdtasdt0(X1,X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) | sz10 = X0 | sz00 = X0)) | (~doDivides0(X0,xk) & ! [X3] : (xk != sdtasdt0(X0,X3) | ~aNaturalNumber0(X3))) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((~isPrime0(X0) & (? [X1] : ((X0 != X1 & sz10 != X1) & (doDivides0(X1,X0) & ? [X2] : (sdtasdt0(X1,X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(X1))) | sz10 = X0 | sz00 = X0)) | (~doDivides0(X0,xk) & ! [X3] : (xk != sdtasdt0(X0,X3) | ~aNaturalNumber0(X3))) | ~aNaturalNumber0(X0))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ~? [X0] : ((isPrime0(X0) | (! [X1] : ((doDivides0(X1,X0) & ? [X2] : (sdtasdt0(X1,X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) => (X0 = X1 | sz10 = X1)) & sz10 != X0 & sz00 != X0)) & (doDivides0(X0,xk) | ? [X3] : (xk = sdtasdt0(X0,X3) & aNaturalNumber0(X3))) & aNaturalNumber0(X0))),
  inference(rectify,[],[f49])).
fof(f49,negated_conjecture,(
  ~? [X0] : ((isPrime0(X0) | (! [X1] : ((doDivides0(X1,X0) & ? [X2] : (sdtasdt0(X1,X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) => (X0 = X1 | sz10 = X1)) & sz10 != X0 & sz00 != X0)) & (doDivides0(X0,xk) | ? [X1] : (sdtasdt0(X0,X1) = xk & aNaturalNumber0(X1))) & aNaturalNumber0(X0))),
  inference(negated_conjecture,[],[f48])).
fof(f48,conjecture,(
  ? [X0] : ((isPrime0(X0) | (! [X1] : ((doDivides0(X1,X0) & ? [X2] : (sdtasdt0(X1,X2) = X0 & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) => (X0 = X1 | sz10 = X1)) & sz10 != X0 & sz00 != X0)) & (doDivides0(X0,xk) | ? [X1] : (sdtasdt0(X0,X1) = xk & aNaturalNumber0(X1))) & aNaturalNumber0(X0))),
  file('NUM500+3.tptp',unknown)).
fof(f241,plain,(
  ( ! [X0] : (doDivides0(sK6(X0),X0) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f152])).
% SZS output end Proof for NUM500+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
