% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for RNG039-1
% SZS output start Proof for RNG039-1
fof(f1011,plain,(
  $false),
  inference(subsumption_resolution,[],[f1006,f588])).
fof(f588,plain,(
  additive_identity != d),
  inference(backward_demodulation,[],[f561,f196])).
fof(f196,plain,(
  b != d),
  inference(backward_demodulation,[],[f168,f60])).
fof(f60,axiom,(
  c != d),
  file('RNG039-1.tptp',unknown)).
fof(f168,plain,(
  b = c),
  inference(resolution,[],[f148,f107])).
fof(f107,plain,(
  product(a,b,b)),
  inference(superposition,[],[f91,f23])).
fof(f23,axiom,(
  ( ! [X10] : (multiply(X10,X10) = X10) )),
  file('RNG039-1.tptp',unknown)).
fof(f91,plain,(
  ( ! [X1] : (product(a,b,multiply(X1,b))) )),
  inference(superposition,[],[f32,f23])).
fof(f32,axiom,(
  ( ! [X10,X11] : (product(a,multiply(b,X10),multiply(X11,X10))) )),
  file('RNG039-1.tptp',unknown)).
fof(f148,plain,(
  ( ! [X37] : (~product(a,b,X37) | c = X37) )),
  inference(resolution,[],[f17,f58])).
fof(f58,axiom,(
  product(a,b,c)),
  file('RNG039-1.tptp',unknown)).
fof(f17,axiom,(
  ( ! [X4,X2,X0,X1] : (~product(X0,X1,X4) | X2 = X4 | ~product(X0,X1,X2)) )),
  file('RNG039-1.tptp',unknown)).
fof(f561,plain,(
  additive_identity = b),
  inference(resolution,[],[f222,f241])).
fof(f241,plain,(
  ( ! [X2] : (product(X2,b,b)) )),
  inference(superposition,[],[f31,f208])).
fof(f208,plain,(
  ( ! [X0] : (b = multiply(X0,b)) )),
  inference(backward_demodulation,[],[f168,f165])).
fof(f165,plain,(
  ( ! [X0] : (c = multiply(X0,b)) )),
  inference(resolution,[],[f148,f91])).
fof(f31,axiom,(
  ( ! [X10,X11] : (product(X10,multiply(X10,X11),multiply(X10,X11))) )),
  file('RNG039-1.tptp',unknown)).
fof(f222,plain,(
  ( ! [X17] : (~product(add(a,b),b,X17) | additive_identity = X17) )),
  inference(forward_demodulation,[],[f199,f22])).
fof(f22,axiom,(
  ( ! [X10] : (additive_identity = add(X10,X10)) )),
  file('RNG039-1.tptp',unknown)).
fof(f199,plain,(
  ( ! [X17] : (add(b,b) = X17 | ~product(add(a,b),b,X17)) )),
  inference(backward_demodulation,[],[f168,f133])).
fof(f133,plain,(
  ( ! [X17] : (add(c,b) = X17 | ~product(add(a,b),b,X17)) )),
  inference(resolution,[],[f17,f53])).
fof(f53,axiom,(
  product(add(a,b),b,add(c,b))),
  file('RNG039-1.tptp',unknown)).
fof(f1006,plain,(
  additive_identity = d),
  inference(resolution,[],[f953,f120])).
fof(f120,plain,(
  ( ! [X17,X16] : (~sum(X17,X17,X16) | additive_identity = X16) )),
  inference(resolution,[],[f16,f20])).
fof(f20,axiom,(
  ( ! [X10] : (sum(X10,X10,additive_identity)) )),
  file('RNG039-1.tptp',unknown)).
fof(f16,axiom,(
  ( ! [X4,X2,X0,X1] : (~sum(X0,X1,X4) | X2 = X4 | ~sum(X0,X1,X2)) )),
  file('RNG039-1.tptp',unknown)).
fof(f953,plain,(
  sum(a,a,d)),
  inference(superposition,[],[f74,f924])).
fof(f924,plain,(
  a = add(d,a)),
  inference(resolution,[],[f684,f125])).
fof(f125,plain,(
  ( ! [X0,X1] : (~product(X1,X1,X0) | X0 = X1) )),
  inference(resolution,[],[f17,f57])).
fof(f57,axiom,(
  ( ! [X10] : (product(X10,X10,X10)) )),
  file('RNG039-1.tptp',unknown)).
fof(f684,plain,(
  product(a,a,add(d,a))),
  inference(forward_demodulation,[],[f569,f350])).
fof(f350,plain,(
  ( ! [X1] : (add(additive_identity,X1) = X1) )),
  inference(resolution,[],[f114,f18])).
fof(f18,axiom,(
  ( ! [X10,X11] : (sum(X10,add(X10,X11),X11)) )),
  file('RNG039-1.tptp',unknown)).
fof(f114,plain,(
  ( ! [X0,X1] : (~sum(additive_identity,X1,X0) | X0 = X1) )),
  inference(resolution,[],[f16,f1])).
fof(f1,axiom,(
  ( ! [X0] : (sum(additive_identity,X0,X0)) )),
  file('RNG039-1.tptp',unknown)).
fof(f569,plain,(
  product(add(additive_identity,a),a,add(d,a))),
  inference(backward_demodulation,[],[f561,f55])).
fof(f55,axiom,(
  product(add(b,a),a,add(d,a))),
  file('RNG039-1.tptp',unknown)).
fof(f74,plain,(
  ( ! [X2,X3] : (sum(X2,add(X3,X2),X3)) )),
  inference(resolution,[],[f9,f19])).
fof(f19,axiom,(
  ( ! [X10,X11] : (sum(add(X10,X11),X11,X10)) )),
  file('RNG039-1.tptp',unknown)).
fof(f9,axiom,(
  ( ! [X0,X3,X1] : (~sum(X0,X1,X3) | sum(X1,X0,X3)) )),
  file('RNG039-1.tptp',unknown)).
% SZS output end Proof for RNG039-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.037 s
% ------------------------------
% ------------------------------
