% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWV353-2
% SZS output start Proof for SWV353-2
fof(f15,plain,(
  $false),
  inference(subsumption_resolution,[],[f14,f8])).
fof(f8,plain,(
  v_NAa != v_NA),
  inference(equality_resolution,[],[f5])).
fof(f5,axiom,(
  ( ! [X0] : (v_NAa != v_NA | c_Message_Omsg_OKey(v_K) != X0) )),
  file('SWV353-2.tptp',unknown)).
fof(f14,plain,(
  v_NAa = v_NA),
  inference(resolution,[],[f13,f10])).
fof(f10,plain,(
  c_in(c_Event_Oevent_OGets(c_Message_Oagent_OServer,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_Ba),c_Message_Omsg_OCrypt(c_Public_OshrK(v_Ba),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_Aa),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(v_NAa),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evso,tc_Event_Oevent),tc_Event_Oevent)),
  inference(subsumption_resolution,[],[f9,f1])).
fof(f1,axiom,(
  c_in(v_evso,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))),
  file('SWV353-2.tptp',unknown)).
fof(f9,plain,(
  c_in(c_Event_Oevent_OGets(c_Message_Oagent_OServer,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_Ba),c_Message_Omsg_OCrypt(c_Public_OshrK(v_Ba),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_Aa),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(v_NAa),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evso,tc_Event_Oevent),tc_Event_Oevent) | ~c_in(v_evso,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))),
  inference(resolution,[],[f6,f2])).
fof(f2,axiom,(
  c_in(c_Event_Oevent_OSays(c_Message_Oagent_OServer,v_Aa,c_Message_Omsg_OMPair(c_Message_Omsg_OCrypt(c_Public_OshrK(v_Aa),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_Ba),c_Message_Omsg_OMPair(c_Message_Omsg_OKey(v_K),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(v_NAa),c_Message_Omsg_ONonce(v_NB))))),v_X)),c_List_Oset(v_evso,tc_Event_Oevent),tc_Event_Oevent)),
  file('SWV353-2.tptp',unknown)).
fof(f6,axiom,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~c_in(c_Event_Oevent_OSays(c_Message_Oagent_OServer,X2,c_Message_Omsg_OMPair(c_Message_Omsg_OCrypt(c_Public_OshrK(X2),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X3),c_Message_Omsg_OMPair(X4,c_Message_Omsg_OMPair(X5,X6)))),X7)),c_List_Oset(X1,tc_Event_Oevent),tc_Event_Oevent) | c_in(c_Event_Oevent_OGets(c_Message_Oagent_OServer,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X3),c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X2),c_Message_Omsg_OMPair(X5,X6))))),c_List_Oset(X1,tc_Event_Oevent),tc_Event_Oevent) | ~c_in(X1,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))) )),
  file('SWV353-2.tptp',unknown)).
fof(f13,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~c_in(c_Event_Oevent_OGets(X1,c_Message_Omsg_OMPair(X2,c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X4),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X0),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evso,tc_Event_Oevent),tc_Event_Oevent) | v_NA = X0) )),
  inference(subsumption_resolution,[],[f12,f1])).
fof(f12,plain,(
  ( ! [X4,X2,X0,X3,X1] : (v_NA = X0 | ~c_in(c_Event_Oevent_OGets(X1,c_Message_Omsg_OMPair(X2,c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X4),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X0),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evso,tc_Event_Oevent),tc_Event_Oevent) | ~c_in(v_evso,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))) )),
  inference(subsumption_resolution,[],[f11,f4])).
fof(f4,axiom,(
  ~c_in(c_Message_Omsg_ONonce(v_NB),c_Message_Oanalz(c_Event_Oknows(c_Message_Oagent_OSpy,v_evso)),tc_Message_Omsg)),
  file('SWV353-2.tptp',unknown)).
fof(f11,plain,(
  ( ! [X4,X2,X0,X3,X1] : (c_in(c_Message_Omsg_ONonce(v_NB),c_Message_Oanalz(c_Event_Oknows(c_Message_Oagent_OSpy,v_evso)),tc_Message_Omsg) | v_NA = X0 | ~c_in(c_Event_Oevent_OGets(X1,c_Message_Omsg_OMPair(X2,c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X4),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X0),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evso,tc_Event_Oevent),tc_Event_Oevent) | ~c_in(v_evso,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))) )),
  inference(resolution,[],[f7,f3])).
fof(f3,axiom,(
  c_in(c_Event_Oevent_OSays(v_B,c_Message_Oagent_OServer,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_B),c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(v_NA),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evso,tc_Event_Oevent),tc_Event_Oevent)),
  file('SWV353-2.tptp',unknown)).
fof(f7,axiom,(
  ( ! [X6,X2,X14,X12,X10,X8,X7,X3,X1,X15,X13,X11,X9] : (~c_in(c_Event_Oevent_OSays(X13,X14,c_Message_Omsg_OMPair(X7,c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X2),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X15),X6))))),c_List_Oset(X1,tc_Event_Oevent),tc_Event_Oevent) | c_in(X6,c_Message_Oanalz(c_Event_Oknows(c_Message_Oagent_OSpy,X1)),tc_Message_Omsg) | X12 = X15 | ~c_in(c_Event_Oevent_OGets(X8,c_Message_Omsg_OMPair(X9,c_Message_Omsg_OCrypt(c_Public_OshrK(X10),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X11),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X12),X6))))),c_List_Oset(X1,tc_Event_Oevent),tc_Event_Oevent) | ~c_in(X1,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))) )),
  file('SWV353-2.tptp',unknown)).
% SZS output end Proof for SWV353-2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
