% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET853-1
% SZS output start Proof for SET853-1
fof(f29692,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f3395,f4114,f4163,f4182,f29689])).
fof(f29689,plain,(
  spl0_5 | ~spl0_24),
  inference(avatar_contradiction_clause,[],[f29688])).
fof(f29688,plain,(
  $false | (spl0_5 | ~spl0_24)),
  inference(subsumption_resolution,[],[f29687,f3394])).
fof(f3394,plain,(
  ~c_lessequals(v_xa,v_x,tc_set(tc_set(t_a))) | spl0_5),
  inference(avatar_component_clause,[],[f3392])).
fof(f3392,plain,(
  spl0_5 <=> c_lessequals(v_xa,v_x,tc_set(tc_set(t_a)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f29687,plain,(
  c_lessequals(v_xa,v_x,tc_set(tc_set(t_a))) | ~spl0_24),
  inference(subsumption_resolution,[],[f29671,f1368])).
fof(f1368,axiom,(
  c_in(v_xa,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))),
  file('SET853-1.tptp',unknown)).
fof(f29671,plain,(
  ~c_in(v_xa,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | c_lessequals(v_xa,v_x,tc_set(tc_set(t_a))) | ~spl0_24),
  inference(resolution,[],[f4162,f2056])).
fof(f2056,plain,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),v_xa,tc_set(tc_set(t_a)))),
  inference(subsumption_resolution,[],[f2055,f1370])).
fof(f1370,axiom,(
  v_xa != c_Zorn_Osucc(v_S,v_x,t_a)),
  file('SET853-1.tptp',unknown)).
fof(f2055,plain,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),v_xa,tc_set(tc_set(t_a))) | v_xa = c_Zorn_Osucc(v_S,v_x,t_a)),
  inference(resolution,[],[f309,f1369])).
fof(f1369,axiom,(
  c_lessequals(v_xa,c_Zorn_Osucc(v_S,v_x,t_a),tc_set(tc_set(t_a)))),
  file('SET853-1.tptp',unknown)).
fof(f309,axiom,(
  ( ! [X2,X15,X13] : (~c_lessequals(X15,X13,tc_set(X2)) | ~c_lessequals(X13,X15,tc_set(X2)) | X13 = X15) )),
  file('SET853-1.tptp',unknown)).
fof(f4162,plain,(
  ( ! [X0] : (c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X0,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | c_lessequals(X0,v_x,tc_set(tc_set(t_a)))) ) | ~spl0_24),
  inference(avatar_component_clause,[],[f4161])).
fof(f4161,plain,(
  spl0_24 <=> ! [X0] : (c_lessequals(X0,v_x,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X0,tc_set(tc_set(t_a))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_24])])).
fof(f4182,plain,(
  spl0_24 | ~spl0_23),
  inference(avatar_split_clause,[],[f4181,f4157,f4161])).
fof(f4157,plain,(
  spl0_23 <=> v_x = c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_23])])).
fof(f4181,plain,(
  ( ! [X3] : (c_lessequals(X3,v_x,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X3,tc_set(tc_set(t_a))) | ~c_in(X3,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))) ) | ~spl0_23),
  inference(subsumption_resolution,[],[f4174,f1367])).
fof(f1367,axiom,(
  c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))),
  file('SET853-1.tptp',unknown)).
fof(f4174,plain,(
  ( ! [X3] : (c_lessequals(X3,v_x,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X3,tc_set(tc_set(t_a))) | ~c_in(X3,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))) ) | ~spl0_23),
  inference(trivial_inequality_removal,[],[f4173])).
fof(f4173,plain,(
  ( ! [X3] : (v_x != v_x | c_lessequals(X3,v_x,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X3,tc_set(tc_set(t_a))) | ~c_in(X3,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))) ) | ~spl0_23),
  inference(superposition,[],[f1361,f4159])).
fof(f4159,plain,(
  v_x = c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a) | ~spl0_23),
  inference(avatar_component_clause,[],[f4157])).
fof(f1361,axiom,(
  ( ! [X2,X12,X40,X18] : (c_Zorn_OTFin__linear__lemma1__1(X40,X12,X2) != X12 | c_lessequals(X18,X12,tc_set(tc_set(X2))) | c_lessequals(c_Zorn_Osucc(X40,X12,X2),X18,tc_set(tc_set(X2))) | ~c_in(X18,c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2))) | ~c_in(X12,c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2)))) )),
  file('SET853-1.tptp',unknown)).
fof(f4163,plain,(
  spl0_23 | spl0_24 | spl0_5),
  inference(avatar_split_clause,[],[f4155,f3392,f4161,f4157])).
fof(f4155,plain,(
  ( ! [X0] : (c_lessequals(X0,v_x,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X0,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | v_x = c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a)) ) | spl0_5),
  inference(subsumption_resolution,[],[f4154,f4085])).
fof(f4085,plain,(
  c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | spl0_5),
  inference(subsumption_resolution,[],[f4084,f1367])).
fof(f4084,plain,(
  c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | spl0_5),
  inference(subsumption_resolution,[],[f4083,f1368])).
fof(f4083,plain,(
  c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_xa,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | spl0_5),
  inference(subsumption_resolution,[],[f4073,f3394])).
fof(f4073,plain,(
  c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | c_lessequals(v_xa,v_x,tc_set(tc_set(t_a))) | ~c_in(v_xa,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))),
  inference(resolution,[],[f1359,f2056])).
fof(f1359,axiom,(
  ( ! [X2,X12,X40,X18] : (c_in(c_Zorn_OTFin__linear__lemma1__1(X40,X12,X2),c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2))) | c_lessequals(c_Zorn_Osucc(X40,X12,X2),X18,tc_set(tc_set(X2))) | c_lessequals(X18,X12,tc_set(tc_set(X2))) | ~c_in(X18,c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2))) | ~c_in(X12,c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2)))) )),
  file('SET853-1.tptp',unknown)).
fof(f4154,plain,(
  ( ! [X0] : (c_lessequals(X0,v_x,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X0,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | v_x = c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a) | ~c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))) ) | spl0_5),
  inference(subsumption_resolution,[],[f4153,f4044])).
fof(f4044,plain,(
  c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),v_x,tc_set(tc_set(t_a))) | spl0_5),
  inference(subsumption_resolution,[],[f4043,f1367])).
fof(f4043,plain,(
  c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),v_x,tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | spl0_5),
  inference(subsumption_resolution,[],[f4042,f1368])).
fof(f4042,plain,(
  c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),v_x,tc_set(tc_set(t_a))) | ~c_in(v_xa,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | spl0_5),
  inference(subsumption_resolution,[],[f4017,f3394])).
fof(f4017,plain,(
  c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),v_x,tc_set(tc_set(t_a))) | c_lessequals(v_xa,v_x,tc_set(tc_set(t_a))) | ~c_in(v_xa,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))),
  inference(resolution,[],[f1360,f2056])).
fof(f1360,axiom,(
  ( ! [X2,X12,X40,X18] : (c_lessequals(c_Zorn_OTFin__linear__lemma1__1(X40,X12,X2),X12,tc_set(tc_set(X2))) | c_lessequals(c_Zorn_Osucc(X40,X12,X2),X18,tc_set(tc_set(X2))) | c_lessequals(X18,X12,tc_set(tc_set(X2))) | ~c_in(X18,c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2))) | ~c_in(X12,c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2)))) )),
  file('SET853-1.tptp',unknown)).
fof(f4153,plain,(
  ( ! [X0] : (c_lessequals(X0,v_x,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X0,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),v_x,tc_set(tc_set(t_a))) | v_x = c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a) | ~c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))) )),
  inference(subsumption_resolution,[],[f4145,f1367])).
fof(f4145,plain,(
  ( ! [X0] : (c_lessequals(X0,v_x,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),X0,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),v_x,tc_set(tc_set(t_a))) | v_x = c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a) | ~c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,v_x,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))) )),
  inference(resolution,[],[f1362,f1372])).
fof(f1372,axiom,(
  ( ! [X41] : (c_lessequals(c_Zorn_Osucc(v_S,X41,t_a),v_x,tc_set(tc_set(t_a))) | ~c_lessequals(X41,v_x,tc_set(tc_set(t_a))) | v_x = X41 | ~c_in(X41,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))) )),
  file('SET853-1.tptp',unknown)).
fof(f1362,axiom,(
  ( ! [X2,X12,X40,X18] : (~c_lessequals(c_Zorn_Osucc(X40,c_Zorn_OTFin__linear__lemma1__1(X40,X12,X2),X2),X12,tc_set(tc_set(X2))) | c_lessequals(X18,X12,tc_set(tc_set(X2))) | c_lessequals(c_Zorn_Osucc(X40,X12,X2),X18,tc_set(tc_set(X2))) | ~c_in(X18,c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2))) | ~c_in(X12,c_Zorn_OTFin(X40,X2),tc_set(tc_set(X2)))) )),
  file('SET853-1.tptp',unknown)).
fof(f4114,plain,(
  ~spl0_4),
  inference(avatar_contradiction_clause,[],[f4113])).
fof(f4113,plain,(
  $false | ~spl0_4),
  inference(subsumption_resolution,[],[f4098,f310])).
fof(f310,axiom,(
  ( ! [X2,X13] : (c_lessequals(X13,X13,tc_set(X2))) )),
  file('SET853-1.tptp',unknown)).
fof(f4098,plain,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),c_Zorn_Osucc(v_S,v_x,t_a),tc_set(tc_set(t_a))) | ~spl0_4),
  inference(backward_demodulation,[],[f3390,f1371])).
fof(f1371,axiom,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_xa,t_a),c_Zorn_Osucc(v_S,v_x,t_a),tc_set(tc_set(t_a)))),
  file('SET853-1.tptp',unknown)).
fof(f3390,plain,(
  v_x = v_xa | ~spl0_4),
  inference(avatar_component_clause,[],[f3388])).
fof(f3388,plain,(
  spl0_4 <=> v_x = v_xa),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f3395,plain,(
  spl0_4 | ~spl0_5),
  inference(avatar_split_clause,[],[f3386,f3392,f3388])).
fof(f3386,plain,(
  ~c_lessequals(v_xa,v_x,tc_set(tc_set(t_a))) | v_x = v_xa),
  inference(subsumption_resolution,[],[f3380,f1368])).
fof(f3380,plain,(
  ~c_lessequals(v_xa,v_x,tc_set(tc_set(t_a))) | v_x = v_xa | ~c_in(v_xa,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))),
  inference(resolution,[],[f1372,f2543])).
fof(f2543,plain,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_xa,t_a),v_x,tc_set(tc_set(t_a)))),
  inference(resolution,[],[f1366,f1371])).
fof(f1366,axiom,(
  ( ! [X2,X0,X23,X40] : (c_lessequals(X0,c_Zorn_Osucc(X40,X23,X2),tc_set(tc_set(X2))) | ~c_lessequals(X0,X23,tc_set(tc_set(X2)))) )),
  file('SET853-1.tptp',unknown)).
% SZS output end Proof for SET853-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 27504
% Time elapsed: 0.629 s
% ------------------------------
% ------------------------------
