% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SYN350-1
% SZS output start Proof for SYN350-1
fof(f30,plain,(
  $false),
  inference(subsumption_resolution,[],[f29,f10])).
fof(f10,plain,(
  ( ! [X0] : (~f(X0,a)) )),
  inference(subsumption_resolution,[],[f9,f6])).
fof(f6,axiom,(
  ( ! [X0,X1] : (~f(a,z(X0,X1)) | ~f(X1,z(X0,X1)) | ~f(X0,X1)) )),
  file('SYN350-1.tptp',unknown)).
fof(f9,plain,(
  ( ! [X0] : (f(a,z(X0,a)) | ~f(X0,a)) )),
  inference(factoring,[],[f5])).
fof(f5,axiom,(
  ( ! [X0,X1] : (f(a,z(X0,X1)) | f(X1,z(X0,X1)) | ~f(X0,X1)) )),
  file('SYN350-1.tptp',unknown)).
fof(f29,plain,(
  f(z(a,a),a)),
  inference(resolution,[],[f26,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (~f(a,z(X0,X1)) | f(z(X0,X1),a)) )),
  file('SYN350-1.tptp',unknown)).
fof(f26,plain,(
  ( ! [X1] : (f(X1,z(X1,a))) )),
  inference(subsumption_resolution,[],[f25,f4])).
fof(f4,axiom,(
  ( ! [X0,X1] : (f(X0,z(X0,X1)) | ~f(X1,z(X0,X1)) | ~f(a,z(X0,X1))) )),
  file('SYN350-1.tptp',unknown)).
fof(f25,plain,(
  ( ! [X1] : (f(X1,z(X1,a)) | f(a,z(X1,a))) )),
  inference(subsumption_resolution,[],[f18,f10])).
fof(f18,plain,(
  ( ! [X1] : (f(X1,z(X1,a)) | f(a,z(X1,a)) | f(z(X1,a),a)) )),
  inference(resolution,[],[f3,f1])).
fof(f3,axiom,(
  ( ! [X0,X1] : (f(X0,z(X0,X1)) | f(X1,z(X0,X1)) | f(a,z(X0,X1))) )),
  file('SYN350-1.tptp',unknown)).
% SZS output end Proof for SYN350-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
