% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SYN454-1
% SZS output start Proof for SYN454-1
fof(f4153,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f264,f319,f329,f336,f339,f341,f342,f343,f351,f356,f361,f396,f401,f406,f441,f446,f451,f456,f461,f466,f471,f476,f481,f486,f491,f496,f501,f506,f511,f516,f521,f536,f541,f546,f551,f561,f571,f576,f621,f626,f631,f636,f646,f651,f656,f661,f666,f671,f676,f681,f686,f691,f696,f701,f711,f716,f721,f726,f731,f736,f741,f746,f761,f766,f781,f786,f791,f796,f800,f814,f818,f819,f820,f824,f828,f831,f832,f837,f858,f862,f873,f880,f887,f891,f895,f896,f900,f904,f908,f912,f919,f920,f922,f923,f924,f925,f929,f930,f931,f932,f936,f937,f973,f991,f994,f1000,f1045,f1058,f1351,f1353,f1515,f1586,f1587,f1606,f1608,f1700,f1891,f2121,f2122,f2132,f2200,f2258,f2285,f2314,f2322,f2323,f2407,f2411,f2436,f2464,f2478,f2527,f2530,f2619,f2622,f2625,f2627,f2652,f2657,f2702,f2709,f2719,f2747,f2802,f2823,f2870,f2873,f2875,f2877,f2945,f2996,f3007,f3018,f3040,f3044,f3125,f3171,f3256,f3287,f3291,f3323,f3324,f3386,f3404,f3430,f3434,f3436,f3509,f3550,f3586,f3606,f3607,f3745,f3791,f3830,f3894,f3897,f3899,f3927,f3933,f3953,f4139])).
fof(f4139,plain,(
  ~spl0_41 | ~spl0_42 | ~spl0_43 | ~spl0_152),
  inference(avatar_contradiction_clause,[],[f4138])).
fof(f4138,plain,(
  $false | (~spl0_41 | ~spl0_42 | ~spl0_43 | ~spl0_152)),
  inference(subsumption_resolution,[],[f4137,f400])).
fof(f400,plain,(
  c1_1(a94) | ~spl0_42),
  inference(avatar_component_clause,[],[f398])).
fof(f398,plain,(
  spl0_42 <=> c1_1(a94)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_42])])).
fof(f4137,plain,(
  ~c1_1(a94) | (~spl0_41 | ~spl0_43 | ~spl0_152)),
  inference(subsumption_resolution,[],[f4120,f405])).
fof(f405,plain,(
  c2_1(a94) | ~spl0_43),
  inference(avatar_component_clause,[],[f403])).
fof(f403,plain,(
  spl0_43 <=> c2_1(a94)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_43])])).
fof(f4120,plain,(
  ~c2_1(a94) | ~c1_1(a94) | (~spl0_41 | ~spl0_152)),
  inference(resolution,[],[f935,f395])).
fof(f395,plain,(
  c0_1(a94) | ~spl0_41),
  inference(avatar_component_clause,[],[f393])).
fof(f393,plain,(
  spl0_41 <=> c0_1(a94)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_41])])).
fof(f935,plain,(
  ( ! [X0] : (~c0_1(X0) | ~c2_1(X0) | ~c1_1(X0)) ) | ~spl0_152),
  inference(avatar_component_clause,[],[f934])).
fof(f934,plain,(
  spl0_152 <=> ! [X0] : (~c2_1(X0) | ~c0_1(X0) | ~c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_152])])).
fof(f3953,plain,(
  spl0_95 | ~spl0_134 | ~spl0_150 | spl0_154),
  inference(avatar_contradiction_clause,[],[f3952])).
fof(f3952,plain,(
  $false | (spl0_95 | ~spl0_134 | ~spl0_150 | spl0_154)),
  inference(subsumption_resolution,[],[f3942,f2266])).
fof(f2266,plain,(
  ~c2_1(a116) | spl0_154),
  inference(avatar_component_clause,[],[f2264])).
fof(f2264,plain,(
  spl0_154 <=> c2_1(a116)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_154])])).
fof(f3942,plain,(
  c2_1(a116) | (spl0_95 | ~spl0_134 | ~spl0_150)),
  inference(resolution,[],[f3938,f665])).
fof(f665,plain,(
  ~c0_1(a116) | spl0_95),
  inference(avatar_component_clause,[],[f663])).
fof(f663,plain,(
  spl0_95 <=> c0_1(a116)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_95])])).
fof(f3938,plain,(
  ( ! [X1] : (c0_1(X1) | c2_1(X1)) ) | (~spl0_134 | ~spl0_150)),
  inference(subsumption_resolution,[],[f857,f918])).
fof(f918,plain,(
  ( ! [X1] : (~c3_1(X1) | c2_1(X1) | c0_1(X1)) ) | ~spl0_150),
  inference(avatar_component_clause,[],[f917])).
fof(f917,plain,(
  spl0_150 <=> ! [X1] : (c2_1(X1) | ~c3_1(X1) | c0_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_150])])).
fof(f857,plain,(
  ( ! [X1] : (c3_1(X1) | c0_1(X1) | c2_1(X1)) ) | ~spl0_134),
  inference(avatar_component_clause,[],[f856])).
fof(f856,plain,(
  spl0_134 <=> ! [X1] : (c3_1(X1) | c0_1(X1) | c2_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_134])])).
fof(f3933,plain,(
  ~spl0_70 | spl0_109 | ~spl0_136 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f3932])).
fof(f3932,plain,(
  $false | (~spl0_70 | spl0_109 | ~spl0_136 | ~spl0_146)),
  inference(subsumption_resolution,[],[f3915,f735])).
fof(f735,plain,(
  ~c3_1(a102) | spl0_109),
  inference(avatar_component_clause,[],[f733])).
fof(f733,plain,(
  spl0_109 <=> c3_1(a102)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_109])])).
fof(f3915,plain,(
  c3_1(a102) | (~spl0_70 | ~spl0_136 | ~spl0_146)),
  inference(resolution,[],[f3900,f540])).
fof(f540,plain,(
  c1_1(a102) | ~spl0_70),
  inference(avatar_component_clause,[],[f538])).
fof(f538,plain,(
  spl0_70 <=> c1_1(a102)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_70])])).
fof(f3900,plain,(
  ( ! [X1] : (~c1_1(X1) | c3_1(X1)) ) | (~spl0_136 | ~spl0_146)),
  inference(subsumption_resolution,[],[f865,f903])).
fof(f903,plain,(
  ( ! [X0] : (c3_1(X0) | ~c0_1(X0) | ~c1_1(X0)) ) | ~spl0_146),
  inference(avatar_component_clause,[],[f902])).
fof(f902,plain,(
  spl0_146 <=> ! [X0] : (c3_1(X0) | ~c0_1(X0) | ~c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_146])])).
fof(f865,plain,(
  ( ! [X1] : (~c1_1(X1) | c3_1(X1) | c0_1(X1)) ) | ~spl0_136),
  inference(avatar_component_clause,[],[f864])).
fof(f864,plain,(
  spl0_136 <=> ! [X1] : (c3_1(X1) | ~c1_1(X1) | c0_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_136])])).
fof(f3927,plain,(
  ~spl0_55 | spl0_89 | ~spl0_136 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f3926])).
fof(f3926,plain,(
  $false | (~spl0_55 | spl0_89 | ~spl0_136 | ~spl0_146)),
  inference(subsumption_resolution,[],[f3910,f635])).
fof(f635,plain,(
  ~c3_1(a125) | spl0_89),
  inference(avatar_component_clause,[],[f633])).
fof(f633,plain,(
  spl0_89 <=> c3_1(a125)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_89])])).
fof(f3910,plain,(
  c3_1(a125) | (~spl0_55 | ~spl0_136 | ~spl0_146)),
  inference(resolution,[],[f3900,f465])).
fof(f465,plain,(
  c1_1(a125) | ~spl0_55),
  inference(avatar_component_clause,[],[f463])).
fof(f463,plain,(
  spl0_55 <=> c1_1(a125)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_55])])).
fof(f3899,plain,(
  spl0_164 | ~spl0_65 | spl0_102 | ~spl0_133),
  inference(avatar_split_clause,[],[f3898,f853,f698,f513,f3003])).
fof(f3003,plain,(
  spl0_164 <=> c0_1(a106)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_164])])).
fof(f513,plain,(
  spl0_65 <=> c1_1(a106)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_65])])).
fof(f698,plain,(
  spl0_102 <=> c2_1(a106)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_102])])).
fof(f853,plain,(
  spl0_133 <=> ! [X0] : (c2_1(X0) | ~c1_1(X0) | c0_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_133])])).
fof(f3898,plain,(
  c0_1(a106) | (~spl0_65 | spl0_102 | ~spl0_133)),
  inference(subsumption_resolution,[],[f3771,f700])).
fof(f700,plain,(
  ~c2_1(a106) | spl0_102),
  inference(avatar_component_clause,[],[f698])).
fof(f3771,plain,(
  c2_1(a106) | c0_1(a106) | (~spl0_65 | ~spl0_133)),
  inference(resolution,[],[f854,f515])).
fof(f515,plain,(
  c1_1(a106) | ~spl0_65),
  inference(avatar_component_clause,[],[f513])).
fof(f854,plain,(
  ( ! [X0] : (~c1_1(X0) | c2_1(X0) | c0_1(X0)) ) | ~spl0_133),
  inference(avatar_component_clause,[],[f853])).
fof(f3897,plain,(
  ~spl0_66 | spl0_102 | ~spl0_145 | ~spl0_164),
  inference(avatar_contradiction_clause,[],[f3896])).
fof(f3896,plain,(
  $false | (~spl0_66 | spl0_102 | ~spl0_145 | ~spl0_164)),
  inference(subsumption_resolution,[],[f3895,f520])).
fof(f520,plain,(
  c3_1(a106) | ~spl0_66),
  inference(avatar_component_clause,[],[f518])).
fof(f518,plain,(
  spl0_66 <=> c3_1(a106)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_66])])).
fof(f3895,plain,(
  ~c3_1(a106) | (spl0_102 | ~spl0_145 | ~spl0_164)),
  inference(subsumption_resolution,[],[f3889,f700])).
fof(f3889,plain,(
  c2_1(a106) | ~c3_1(a106) | (~spl0_145 | ~spl0_164)),
  inference(resolution,[],[f3005,f899])).
fof(f899,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0) | ~c3_1(X0)) ) | ~spl0_145),
  inference(avatar_component_clause,[],[f898])).
fof(f898,plain,(
  spl0_145 <=> ! [X0] : (c2_1(X0) | ~c0_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_145])])).
fof(f3005,plain,(
  c0_1(a106) | ~spl0_164),
  inference(avatar_component_clause,[],[f3003])).
fof(f3894,plain,(
  ~spl0_65 | ~spl0_66 | ~spl0_149 | ~spl0_164),
  inference(avatar_contradiction_clause,[],[f3893])).
fof(f3893,plain,(
  $false | (~spl0_65 | ~spl0_66 | ~spl0_149 | ~spl0_164)),
  inference(subsumption_resolution,[],[f3892,f515])).
fof(f3892,plain,(
  ~c1_1(a106) | (~spl0_66 | ~spl0_149 | ~spl0_164)),
  inference(subsumption_resolution,[],[f3888,f520])).
fof(f3888,plain,(
  ~c3_1(a106) | ~c1_1(a106) | (~spl0_149 | ~spl0_164)),
  inference(resolution,[],[f3005,f915])).
fof(f915,plain,(
  ( ! [X0] : (~c0_1(X0) | ~c3_1(X0) | ~c1_1(X0)) ) | ~spl0_149),
  inference(avatar_component_clause,[],[f914])).
fof(f914,plain,(
  spl0_149 <=> ! [X0] : (~c3_1(X0) | ~c0_1(X0) | ~c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_149])])).
fof(f3830,plain,(
  spl0_100 | ~spl0_64 | spl0_101 | ~spl0_138),
  inference(avatar_split_clause,[],[f3826,f871,f693,f508,f688])).
fof(f688,plain,(
  spl0_100 <=> c0_1(a109)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_100])])).
fof(f508,plain,(
  spl0_64 <=> c2_1(a109)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_64])])).
fof(f693,plain,(
  spl0_101 <=> c3_1(a109)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_101])])).
fof(f871,plain,(
  spl0_138 <=> ! [X1] : (c3_1(X1) | ~c2_1(X1) | c0_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_138])])).
fof(f3826,plain,(
  c0_1(a109) | (~spl0_64 | spl0_101 | ~spl0_138)),
  inference(subsumption_resolution,[],[f3815,f695])).
fof(f695,plain,(
  ~c3_1(a109) | spl0_101),
  inference(avatar_component_clause,[],[f693])).
fof(f3815,plain,(
  c3_1(a109) | c0_1(a109) | (~spl0_64 | ~spl0_138)),
  inference(resolution,[],[f872,f510])).
fof(f510,plain,(
  c2_1(a109) | ~spl0_64),
  inference(avatar_component_clause,[],[f508])).
fof(f872,plain,(
  ( ! [X1] : (~c2_1(X1) | c3_1(X1) | c0_1(X1)) ) | ~spl0_138),
  inference(avatar_component_clause,[],[f871])).
fof(f3791,plain,(
  spl0_154 | ~spl0_57 | spl0_95 | ~spl0_133),
  inference(avatar_split_clause,[],[f3790,f853,f663,f473,f2264])).
fof(f473,plain,(
  spl0_57 <=> c1_1(a116)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_57])])).
fof(f3790,plain,(
  c2_1(a116) | (~spl0_57 | spl0_95 | ~spl0_133)),
  inference(subsumption_resolution,[],[f3770,f665])).
fof(f3770,plain,(
  c2_1(a116) | c0_1(a116) | (~spl0_57 | ~spl0_133)),
  inference(resolution,[],[f854,f475])).
fof(f475,plain,(
  c1_1(a116) | ~spl0_57),
  inference(avatar_component_clause,[],[f473])).
fof(f3745,plain,(
  ~spl0_61 | ~spl0_60 | ~spl0_131 | ~spl0_145),
  inference(avatar_split_clause,[],[f3690,f898,f844,f488,f493])).
fof(f493,plain,(
  spl0_61 <=> c3_1(a112)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_61])])).
fof(f488,plain,(
  spl0_60 <=> c0_1(a112)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_60])])).
fof(f844,plain,(
  spl0_131 <=> ! [X0] : (~c3_1(X0) | ~c0_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_131])])).
fof(f3690,plain,(
  ~c3_1(a112) | (~spl0_60 | ~spl0_131 | ~spl0_145)),
  inference(resolution,[],[f3642,f490])).
fof(f490,plain,(
  c0_1(a112) | ~spl0_60),
  inference(avatar_component_clause,[],[f488])).
fof(f3642,plain,(
  ( ! [X0] : (~c0_1(X0) | ~c3_1(X0)) ) | (~spl0_131 | ~spl0_145)),
  inference(subsumption_resolution,[],[f899,f845])).
fof(f845,plain,(
  ( ! [X0] : (~c0_1(X0) | ~c3_1(X0) | ~c2_1(X0)) ) | ~spl0_131),
  inference(avatar_component_clause,[],[f844])).
fof(f3607,plain,(
  spl0_86 | ~spl0_122 | ~spl0_125 | ~spl0_139 | ~spl0_145 | ~spl0_147 | ~spl0_150),
  inference(avatar_contradiction_clause,[],[f3596])).
fof(f3596,plain,(
  $false | (spl0_86 | ~spl0_122 | ~spl0_125 | ~spl0_139 | ~spl0_145 | ~spl0_147 | ~spl0_150)),
  inference(resolution,[],[f3593,f620])).
fof(f620,plain,(
  ~c2_1(a145) | spl0_86),
  inference(avatar_component_clause,[],[f618])).
fof(f618,plain,(
  spl0_86 <=> c2_1(a145)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_86])])).
fof(f3593,plain,(
  ( ! [X0] : (c2_1(X0)) ) | (~spl0_122 | ~spl0_125 | ~spl0_139 | ~spl0_145 | ~spl0_147 | ~spl0_150)),
  inference(subsumption_resolution,[],[f3592,f3589])).
fof(f3589,plain,(
  ( ! [X1] : (c3_1(X1) | c2_1(X1)) ) | (~spl0_122 | ~spl0_125 | ~spl0_139 | ~spl0_147)),
  inference(subsumption_resolution,[],[f907,f3588])).
fof(f3588,plain,(
  ( ! [X0] : (c2_1(X0) | c1_1(X0)) ) | (~spl0_122 | ~spl0_125 | ~spl0_139)),
  inference(subsumption_resolution,[],[f799,f3587])).
fof(f3587,plain,(
  ( ! [X0] : (~c3_1(X0) | c1_1(X0)) ) | (~spl0_125 | ~spl0_139)),
  inference(subsumption_resolution,[],[f813,f876])).
fof(f876,plain,(
  ( ! [X0] : (~c3_1(X0) | c2_1(X0) | c1_1(X0)) ) | ~spl0_139),
  inference(avatar_component_clause,[],[f875])).
fof(f875,plain,(
  spl0_139 <=> ! [X0] : (c2_1(X0) | ~c3_1(X0) | c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_139])])).
fof(f813,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c2_1(X0) | c1_1(X0)) ) | ~spl0_125),
  inference(avatar_component_clause,[],[f812])).
fof(f812,plain,(
  spl0_125 <=> ! [X0] : (c1_1(X0) | ~c2_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_125])])).
fof(f799,plain,(
  ( ! [X0] : (c3_1(X0) | c1_1(X0) | c2_1(X0)) ) | ~spl0_122),
  inference(avatar_component_clause,[],[f798])).
fof(f798,plain,(
  spl0_122 <=> ! [X0] : (c3_1(X0) | c1_1(X0) | c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_122])])).
fof(f907,plain,(
  ( ! [X1] : (~c1_1(X1) | c3_1(X1) | c2_1(X1)) ) | ~spl0_147),
  inference(avatar_component_clause,[],[f906])).
fof(f906,plain,(
  spl0_147 <=> ! [X1] : (c3_1(X1) | ~c1_1(X1) | c2_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_147])])).
fof(f3592,plain,(
  ( ! [X0] : (c2_1(X0) | ~c3_1(X0)) ) | (~spl0_145 | ~spl0_150)),
  inference(subsumption_resolution,[],[f899,f918])).
fof(f3606,plain,(
  spl0_93 | ~spl0_122 | ~spl0_125 | ~spl0_139 | ~spl0_145 | ~spl0_147 | ~spl0_150),
  inference(avatar_contradiction_clause,[],[f3597])).
fof(f3597,plain,(
  $false | (spl0_93 | ~spl0_122 | ~spl0_125 | ~spl0_139 | ~spl0_145 | ~spl0_147 | ~spl0_150)),
  inference(resolution,[],[f3593,f655])).
fof(f655,plain,(
  ~c2_1(a118) | spl0_93),
  inference(avatar_component_clause,[],[f653])).
fof(f653,plain,(
  spl0_93 <=> c2_1(a118)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_93])])).
fof(f3586,plain,(
  ~spl0_155 | ~spl0_62 | spl0_99 | ~spl0_127),
  inference(avatar_split_clause,[],[f3585,f822,f683,f498,f2319])).
fof(f2319,plain,(
  spl0_155 <=> c1_1(a110)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_155])])).
fof(f498,plain,(
  spl0_62 <=> c2_1(a110)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_62])])).
fof(f683,plain,(
  spl0_99 <=> c0_1(a110)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_99])])).
fof(f822,plain,(
  spl0_127 <=> ! [X0] : (c0_1(X0) | ~c1_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_127])])).
fof(f3585,plain,(
  ~c1_1(a110) | (~spl0_62 | spl0_99 | ~spl0_127)),
  inference(subsumption_resolution,[],[f3455,f685])).
fof(f685,plain,(
  ~c0_1(a110) | spl0_99),
  inference(avatar_component_clause,[],[f683])).
fof(f3455,plain,(
  ~c1_1(a110) | c0_1(a110) | (~spl0_62 | ~spl0_127)),
  inference(resolution,[],[f823,f500])).
fof(f500,plain,(
  c2_1(a110) | ~spl0_62),
  inference(avatar_component_clause,[],[f498])).
fof(f823,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c1_1(X0) | c0_1(X0)) ) | ~spl0_127),
  inference(avatar_component_clause,[],[f822])).
fof(f3550,plain,(
  ~spl0_60 | ~spl0_61 | spl0_98 | ~spl0_128),
  inference(avatar_contradiction_clause,[],[f3549])).
fof(f3549,plain,(
  $false | (~spl0_60 | ~spl0_61 | spl0_98 | ~spl0_128)),
  inference(subsumption_resolution,[],[f3548,f495])).
fof(f495,plain,(
  c3_1(a112) | ~spl0_61),
  inference(avatar_component_clause,[],[f493])).
fof(f3548,plain,(
  ~c3_1(a112) | (~spl0_60 | spl0_98 | ~spl0_128)),
  inference(subsumption_resolution,[],[f3532,f680])).
fof(f680,plain,(
  ~c1_1(a112) | spl0_98),
  inference(avatar_component_clause,[],[f678])).
fof(f678,plain,(
  spl0_98 <=> c1_1(a112)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_98])])).
fof(f3532,plain,(
  c1_1(a112) | ~c3_1(a112) | (~spl0_60 | ~spl0_128)),
  inference(resolution,[],[f827,f490])).
fof(f827,plain,(
  ( ! [X0] : (~c0_1(X0) | c1_1(X0) | ~c3_1(X0)) ) | ~spl0_128),
  inference(avatar_component_clause,[],[f826])).
fof(f826,plain,(
  spl0_128 <=> ! [X0] : (c1_1(X0) | ~c0_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_128])])).
fof(f3509,plain,(
  ~spl0_41 | ~spl0_42 | ~spl0_43 | ~spl0_130 | ~spl0_131),
  inference(avatar_contradiction_clause,[],[f3508])).
fof(f3508,plain,(
  $false | (~spl0_41 | ~spl0_42 | ~spl0_43 | ~spl0_130 | ~spl0_131)),
  inference(subsumption_resolution,[],[f3507,f405])).
fof(f3507,plain,(
  ~c2_1(a94) | (~spl0_41 | ~spl0_42 | ~spl0_43 | ~spl0_130 | ~spl0_131)),
  inference(subsumption_resolution,[],[f3493,f3483])).
fof(f3483,plain,(
  c3_1(a94) | (~spl0_42 | ~spl0_43 | ~spl0_130)),
  inference(subsumption_resolution,[],[f3471,f400])).
fof(f3471,plain,(
  ~c1_1(a94) | c3_1(a94) | (~spl0_43 | ~spl0_130)),
  inference(resolution,[],[f840,f405])).
fof(f840,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c1_1(X0) | c3_1(X0)) ) | ~spl0_130),
  inference(avatar_component_clause,[],[f839])).
fof(f839,plain,(
  spl0_130 <=> ! [X0] : (c3_1(X0) | ~c1_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_130])])).
fof(f3493,plain,(
  ~c3_1(a94) | ~c2_1(a94) | (~spl0_41 | ~spl0_131)),
  inference(resolution,[],[f845,f395])).
fof(f3436,plain,(
  ~spl0_59 | ~spl0_146 | ~spl0_149 | ~spl0_163),
  inference(avatar_contradiction_clause,[],[f3435])).
fof(f3435,plain,(
  $false | (~spl0_59 | ~spl0_146 | ~spl0_149 | ~spl0_163)),
  inference(subsumption_resolution,[],[f3423,f2926])).
fof(f2926,plain,(
  c1_1(a115) | ~spl0_163),
  inference(avatar_component_clause,[],[f2925])).
fof(f2925,plain,(
  spl0_163 <=> c1_1(a115)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_163])])).
fof(f3423,plain,(
  ~c1_1(a115) | (~spl0_59 | ~spl0_146 | ~spl0_149)),
  inference(resolution,[],[f3373,f485])).
fof(f485,plain,(
  c0_1(a115) | ~spl0_59),
  inference(avatar_component_clause,[],[f483])).
fof(f483,plain,(
  spl0_59 <=> c0_1(a115)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_59])])).
fof(f3373,plain,(
  ( ! [X0] : (~c0_1(X0) | ~c1_1(X0)) ) | (~spl0_146 | ~spl0_149)),
  inference(subsumption_resolution,[],[f903,f915])).
fof(f3434,plain,(
  ~spl0_54 | ~spl0_55 | ~spl0_146 | ~spl0_149),
  inference(avatar_contradiction_clause,[],[f3433])).
fof(f3433,plain,(
  $false | (~spl0_54 | ~spl0_55 | ~spl0_146 | ~spl0_149)),
  inference(subsumption_resolution,[],[f3422,f465])).
fof(f3422,plain,(
  ~c1_1(a125) | (~spl0_54 | ~spl0_146 | ~spl0_149)),
  inference(resolution,[],[f3373,f460])).
fof(f460,plain,(
  c0_1(a125) | ~spl0_54),
  inference(avatar_component_clause,[],[f458])).
fof(f458,plain,(
  spl0_54 <=> c0_1(a125)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_54])])).
fof(f3430,plain,(
  ~spl0_41 | ~spl0_42 | ~spl0_146 | ~spl0_149),
  inference(avatar_contradiction_clause,[],[f3429])).
fof(f3429,plain,(
  $false | (~spl0_41 | ~spl0_42 | ~spl0_146 | ~spl0_149)),
  inference(subsumption_resolution,[],[f3418,f400])).
fof(f3418,plain,(
  ~c1_1(a94) | (~spl0_41 | ~spl0_146 | ~spl0_149)),
  inference(resolution,[],[f3373,f395])).
fof(f3404,plain,(
  spl0_100 | spl0_101 | ~spl0_135 | ~spl0_136),
  inference(avatar_contradiction_clause,[],[f3403])).
fof(f3403,plain,(
  $false | (spl0_100 | spl0_101 | ~spl0_135 | ~spl0_136)),
  inference(subsumption_resolution,[],[f3397,f695])).
fof(f3397,plain,(
  c3_1(a109) | (spl0_100 | ~spl0_135 | ~spl0_136)),
  inference(resolution,[],[f3372,f690])).
fof(f690,plain,(
  ~c0_1(a109) | spl0_100),
  inference(avatar_component_clause,[],[f688])).
fof(f3372,plain,(
  ( ! [X1] : (c0_1(X1) | c3_1(X1)) ) | (~spl0_135 | ~spl0_136)),
  inference(subsumption_resolution,[],[f861,f865])).
fof(f861,plain,(
  ( ! [X1] : (c3_1(X1) | c0_1(X1) | c1_1(X1)) ) | ~spl0_135),
  inference(avatar_component_clause,[],[f860])).
fof(f860,plain,(
  spl0_135 <=> ! [X1] : (c3_1(X1) | c0_1(X1) | c1_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_135])])).
fof(f3386,plain,(
  ~spl0_64 | spl0_101 | ~spl0_138 | ~spl0_141 | ~spl0_152),
  inference(avatar_contradiction_clause,[],[f3385])).
fof(f3385,plain,(
  $false | (~spl0_64 | spl0_101 | ~spl0_138 | ~spl0_141 | ~spl0_152)),
  inference(subsumption_resolution,[],[f3379,f695])).
fof(f3379,plain,(
  c3_1(a109) | (~spl0_64 | ~spl0_138 | ~spl0_141 | ~spl0_152)),
  inference(resolution,[],[f3371,f510])).
fof(f3371,plain,(
  ( ! [X1] : (~c2_1(X1) | c3_1(X1)) ) | (~spl0_138 | ~spl0_141 | ~spl0_152)),
  inference(subsumption_resolution,[],[f872,f3292])).
fof(f3292,plain,(
  ( ! [X0] : (~c0_1(X0) | ~c2_1(X0)) ) | (~spl0_141 | ~spl0_152)),
  inference(subsumption_resolution,[],[f935,f883])).
fof(f883,plain,(
  ( ! [X0] : (c1_1(X0) | ~c0_1(X0) | ~c2_1(X0)) ) | ~spl0_141),
  inference(avatar_component_clause,[],[f882])).
fof(f882,plain,(
  spl0_141 <=> ! [X0] : (c1_1(X0) | ~c0_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_141])])).
fof(f3324,plain,(
  ~spl0_53 | ~spl0_52 | ~spl0_141 | ~spl0_152),
  inference(avatar_split_clause,[],[f3303,f934,f882,f448,f453])).
fof(f453,plain,(
  spl0_53 <=> c2_1(a127)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_53])])).
fof(f448,plain,(
  spl0_52 <=> c0_1(a127)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_52])])).
fof(f3303,plain,(
  ~c2_1(a127) | (~spl0_52 | ~spl0_141 | ~spl0_152)),
  inference(resolution,[],[f3292,f450])).
fof(f450,plain,(
  c0_1(a127) | ~spl0_52),
  inference(avatar_component_clause,[],[f448])).
fof(f3323,plain,(
  ~spl0_43 | ~spl0_41 | ~spl0_141 | ~spl0_152),
  inference(avatar_split_clause,[],[f3298,f934,f882,f393,f403])).
fof(f3298,plain,(
  ~c2_1(a94) | (~spl0_41 | ~spl0_141 | ~spl0_152)),
  inference(resolution,[],[f3292,f395])).
fof(f3291,plain,(
  ~spl0_76 | ~spl0_77 | spl0_118 | ~spl0_126 | ~spl0_141 | ~spl0_143 | ~spl0_145),
  inference(avatar_contradiction_clause,[],[f3290])).
fof(f3290,plain,(
  $false | (~spl0_76 | ~spl0_77 | spl0_118 | ~spl0_126 | ~spl0_141 | ~spl0_143 | ~spl0_145)),
  inference(subsumption_resolution,[],[f3277,f780])).
fof(f780,plain,(
  ~c1_1(a97) | spl0_118),
  inference(avatar_component_clause,[],[f778])).
fof(f778,plain,(
  spl0_118 <=> c1_1(a97)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_118])])).
fof(f3277,plain,(
  c1_1(a97) | (~spl0_76 | ~spl0_77 | ~spl0_126 | ~spl0_141 | ~spl0_143 | ~spl0_145)),
  inference(resolution,[],[f3262,f2384])).
fof(f2384,plain,(
  c0_1(a97) | (~spl0_76 | ~spl0_77 | ~spl0_143)),
  inference(subsumption_resolution,[],[f2380,f570])).
fof(f570,plain,(
  c2_1(a97) | ~spl0_76),
  inference(avatar_component_clause,[],[f568])).
fof(f568,plain,(
  spl0_76 <=> c2_1(a97)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_76])])).
fof(f2380,plain,(
  ~c2_1(a97) | c0_1(a97) | (~spl0_77 | ~spl0_143)),
  inference(resolution,[],[f890,f575])).
fof(f575,plain,(
  c3_1(a97) | ~spl0_77),
  inference(avatar_component_clause,[],[f573])).
fof(f573,plain,(
  spl0_77 <=> c3_1(a97)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_77])])).
fof(f890,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c2_1(X0) | c0_1(X0)) ) | ~spl0_143),
  inference(avatar_component_clause,[],[f889])).
fof(f889,plain,(
  spl0_143 <=> ! [X0] : (c0_1(X0) | ~c2_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_143])])).
fof(f3262,plain,(
  ( ! [X0] : (~c0_1(X0) | c1_1(X0)) ) | (~spl0_126 | ~spl0_141 | ~spl0_145)),
  inference(subsumption_resolution,[],[f883,f3182])).
fof(f3182,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0)) ) | (~spl0_126 | ~spl0_145)),
  inference(subsumption_resolution,[],[f817,f899])).
fof(f817,plain,(
  ( ! [X0] : (~c0_1(X0) | c3_1(X0) | c2_1(X0)) ) | ~spl0_126),
  inference(avatar_component_clause,[],[f816])).
fof(f816,plain,(
  spl0_126 <=> ! [X0] : (c3_1(X0) | ~c0_1(X0) | c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_126])])).
fof(f3287,plain,(
  ~spl0_60 | spl0_98 | ~spl0_126 | ~spl0_141 | ~spl0_145),
  inference(avatar_contradiction_clause,[],[f3286])).
fof(f3286,plain,(
  $false | (~spl0_60 | spl0_98 | ~spl0_126 | ~spl0_141 | ~spl0_145)),
  inference(subsumption_resolution,[],[f3275,f680])).
fof(f3275,plain,(
  c1_1(a112) | (~spl0_60 | ~spl0_126 | ~spl0_141 | ~spl0_145)),
  inference(resolution,[],[f3262,f490])).
fof(f3256,plain,(
  ~spl0_56 | spl0_93 | spl0_94 | ~spl0_147),
  inference(avatar_contradiction_clause,[],[f3255])).
fof(f3255,plain,(
  $false | (~spl0_56 | spl0_93 | spl0_94 | ~spl0_147)),
  inference(subsumption_resolution,[],[f3254,f655])).
fof(f3254,plain,(
  c2_1(a118) | (~spl0_56 | spl0_94 | ~spl0_147)),
  inference(subsumption_resolution,[],[f3251,f660])).
fof(f660,plain,(
  ~c3_1(a118) | spl0_94),
  inference(avatar_component_clause,[],[f658])).
fof(f658,plain,(
  spl0_94 <=> c3_1(a118)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_94])])).
fof(f3251,plain,(
  c3_1(a118) | c2_1(a118) | (~spl0_56 | ~spl0_147)),
  inference(resolution,[],[f907,f470])).
fof(f470,plain,(
  c1_1(a118) | ~spl0_56),
  inference(avatar_component_clause,[],[f468])).
fof(f468,plain,(
  spl0_56 <=> c1_1(a118)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_56])])).
fof(f3171,plain,(
  ~spl0_72 | spl0_110 | spl0_111 | ~spl0_139),
  inference(avatar_contradiction_clause,[],[f3170])).
fof(f3170,plain,(
  $false | (~spl0_72 | spl0_110 | spl0_111 | ~spl0_139)),
  inference(subsumption_resolution,[],[f3169,f740])).
fof(f740,plain,(
  ~c1_1(a101) | spl0_110),
  inference(avatar_component_clause,[],[f738])).
fof(f738,plain,(
  spl0_110 <=> c1_1(a101)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_110])])).
fof(f3169,plain,(
  c1_1(a101) | (~spl0_72 | spl0_111 | ~spl0_139)),
  inference(subsumption_resolution,[],[f3150,f745])).
fof(f745,plain,(
  ~c2_1(a101) | spl0_111),
  inference(avatar_component_clause,[],[f743])).
fof(f743,plain,(
  spl0_111 <=> c2_1(a101)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_111])])).
fof(f3150,plain,(
  c2_1(a101) | c1_1(a101) | (~spl0_72 | ~spl0_139)),
  inference(resolution,[],[f876,f550])).
fof(f550,plain,(
  c3_1(a101) | ~spl0_72),
  inference(avatar_component_clause,[],[f548])).
fof(f548,plain,(
  spl0_72 <=> c3_1(a101)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_72])])).
fof(f3125,plain,(
  spl0_87 | ~spl0_50 | ~spl0_129 | ~spl0_153),
  inference(avatar_split_clause,[],[f3124,f1030,f835,f438,f623])).
fof(f623,plain,(
  spl0_87 <=> c2_1(a133)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_87])])).
fof(f438,plain,(
  spl0_50 <=> c0_1(a133)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_50])])).
fof(f835,plain,(
  spl0_129 <=> ! [X0] : (c2_1(X0) | ~c0_1(X0) | ~c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_129])])).
fof(f1030,plain,(
  spl0_153 <=> c1_1(a133)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_153])])).
fof(f3124,plain,(
  c2_1(a133) | (~spl0_50 | ~spl0_129 | ~spl0_153)),
  inference(subsumption_resolution,[],[f3100,f1032])).
fof(f1032,plain,(
  c1_1(a133) | ~spl0_153),
  inference(avatar_component_clause,[],[f1030])).
fof(f3100,plain,(
  c2_1(a133) | ~c1_1(a133) | (~spl0_50 | ~spl0_129)),
  inference(resolution,[],[f836,f440])).
fof(f440,plain,(
  c0_1(a133) | ~spl0_50),
  inference(avatar_component_clause,[],[f438])).
fof(f836,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0) | ~c1_1(X0)) ) | ~spl0_129),
  inference(avatar_component_clause,[],[f835])).
fof(f3044,plain,(
  spl0_110 | spl0_111 | ~spl0_122 | ~spl0_139),
  inference(avatar_contradiction_clause,[],[f3043])).
fof(f3043,plain,(
  $false | (spl0_110 | spl0_111 | ~spl0_122 | ~spl0_139)),
  inference(subsumption_resolution,[],[f3033,f740])).
fof(f3033,plain,(
  c1_1(a101) | (spl0_111 | ~spl0_122 | ~spl0_139)),
  inference(resolution,[],[f3011,f745])).
fof(f3011,plain,(
  ( ! [X0] : (c2_1(X0) | c1_1(X0)) ) | (~spl0_122 | ~spl0_139)),
  inference(subsumption_resolution,[],[f876,f799])).
fof(f3040,plain,(
  spl0_87 | ~spl0_122 | ~spl0_139 | spl0_153),
  inference(avatar_contradiction_clause,[],[f3039])).
fof(f3039,plain,(
  $false | (spl0_87 | ~spl0_122 | ~spl0_139 | spl0_153)),
  inference(subsumption_resolution,[],[f3027,f1031])).
fof(f1031,plain,(
  ~c1_1(a133) | spl0_153),
  inference(avatar_component_clause,[],[f1030])).
fof(f3027,plain,(
  c1_1(a133) | (spl0_87 | ~spl0_122 | ~spl0_139)),
  inference(resolution,[],[f3011,f625])).
fof(f625,plain,(
  ~c2_1(a133) | spl0_87),
  inference(avatar_component_clause,[],[f623])).
fof(f3018,plain,(
  ~spl0_159 | ~spl0_64 | spl0_100 | ~spl0_127),
  inference(avatar_split_clause,[],[f3017,f822,f688,f508,f2706])).
fof(f2706,plain,(
  spl0_159 <=> c1_1(a109)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_159])])).
fof(f3017,plain,(
  ~c1_1(a109) | (~spl0_64 | spl0_100 | ~spl0_127)),
  inference(subsumption_resolution,[],[f3013,f690])).
fof(f3013,plain,(
  ~c1_1(a109) | c0_1(a109) | (~spl0_64 | ~spl0_127)),
  inference(resolution,[],[f510,f823])).
fof(f3007,plain,(
  spl0_163 | spl0_97 | ~spl0_59 | ~spl0_142),
  inference(avatar_split_clause,[],[f2711,f885,f483,f673,f2925])).
fof(f673,plain,(
  spl0_97 <=> c3_1(a115)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_97])])).
fof(f885,plain,(
  spl0_142 <=> ! [X1] : (c3_1(X1) | ~c0_1(X1) | c1_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_142])])).
fof(f2711,plain,(
  c3_1(a115) | c1_1(a115) | (~spl0_59 | ~spl0_142)),
  inference(resolution,[],[f485,f886])).
fof(f886,plain,(
  ( ! [X1] : (~c0_1(X1) | c3_1(X1) | c1_1(X1)) ) | ~spl0_142),
  inference(avatar_component_clause,[],[f885])).
fof(f2996,plain,(
  ~spl0_57 | ~spl0_58 | ~spl0_151 | spl0_154),
  inference(avatar_contradiction_clause,[],[f2995])).
fof(f2995,plain,(
  $false | (~spl0_57 | ~spl0_58 | ~spl0_151 | spl0_154)),
  inference(subsumption_resolution,[],[f2994,f2266])).
fof(f2994,plain,(
  c2_1(a116) | (~spl0_57 | ~spl0_58 | ~spl0_151)),
  inference(subsumption_resolution,[],[f2975,f475])).
fof(f2975,plain,(
  ~c1_1(a116) | c2_1(a116) | (~spl0_58 | ~spl0_151)),
  inference(resolution,[],[f928,f480])).
fof(f480,plain,(
  c3_1(a116) | ~spl0_58),
  inference(avatar_component_clause,[],[f478])).
fof(f478,plain,(
  spl0_58 <=> c3_1(a116)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_58])])).
fof(f928,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c1_1(X0) | c2_1(X0)) ) | ~spl0_151),
  inference(avatar_component_clause,[],[f927])).
fof(f927,plain,(
  spl0_151 <=> ! [X0] : (c2_1(X0) | ~c1_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_151])])).
fof(f2945,plain,(
  spl0_159 | spl0_101 | ~spl0_135 | ~spl0_142),
  inference(avatar_split_clause,[],[f2942,f885,f860,f693,f2706])).
fof(f2942,plain,(
  c1_1(a109) | (spl0_101 | ~spl0_135 | ~spl0_142)),
  inference(resolution,[],[f695,f2785])).
fof(f2785,plain,(
  ( ! [X1] : (c3_1(X1) | c1_1(X1)) ) | (~spl0_135 | ~spl0_142)),
  inference(subsumption_resolution,[],[f861,f886])).
fof(f2877,plain,(
  ~spl0_153 | ~spl0_50 | ~spl0_51 | ~spl0_149),
  inference(avatar_split_clause,[],[f2876,f914,f443,f438,f1030])).
fof(f443,plain,(
  spl0_51 <=> c3_1(a133)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_51])])).
fof(f2876,plain,(
  ~c1_1(a133) | (~spl0_50 | ~spl0_51 | ~spl0_149)),
  inference(subsumption_resolution,[],[f2854,f445])).
fof(f445,plain,(
  c3_1(a133) | ~spl0_51),
  inference(avatar_component_clause,[],[f443])).
fof(f2854,plain,(
  ~c3_1(a133) | ~c1_1(a133) | (~spl0_50 | ~spl0_149)),
  inference(resolution,[],[f915,f440])).
fof(f2875,plain,(
  spl0_153 | ~spl0_50 | spl0_87 | ~spl0_140),
  inference(avatar_split_clause,[],[f2874,f878,f623,f438,f1030])).
fof(f878,plain,(
  spl0_140 <=> ! [X1] : (c2_1(X1) | ~c0_1(X1) | c1_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_140])])).
fof(f2874,plain,(
  c1_1(a133) | (~spl0_50 | spl0_87 | ~spl0_140)),
  inference(subsumption_resolution,[],[f2353,f625])).
fof(f2353,plain,(
  c2_1(a133) | c1_1(a133) | (~spl0_50 | ~spl0_140)),
  inference(resolution,[],[f879,f440])).
fof(f879,plain,(
  ( ! [X1] : (~c0_1(X1) | c2_1(X1) | c1_1(X1)) ) | ~spl0_140),
  inference(avatar_component_clause,[],[f878])).
fof(f2873,plain,(
  ~spl0_51 | spl0_87 | ~spl0_50 | ~spl0_145),
  inference(avatar_split_clause,[],[f2840,f898,f438,f623,f443])).
fof(f2840,plain,(
  c2_1(a133) | ~c3_1(a133) | (~spl0_50 | ~spl0_145)),
  inference(resolution,[],[f899,f440])).
fof(f2870,plain,(
  ~spl0_32 | ~spl0_34 | ~spl0_149 | ~spl0_158),
  inference(avatar_split_clause,[],[f2863,f2654,f914,f358,f348])).
fof(f348,plain,(
  spl0_32 <=> c1_1(a165)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_32])])).
fof(f358,plain,(
  spl0_34 <=> c3_1(a165)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_34])])).
fof(f2654,plain,(
  spl0_158 <=> c0_1(a165)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_158])])).
fof(f2863,plain,(
  ~c1_1(a165) | (~spl0_34 | ~spl0_149 | ~spl0_158)),
  inference(subsumption_resolution,[],[f2851,f360])).
fof(f360,plain,(
  c3_1(a165) | ~spl0_34),
  inference(avatar_component_clause,[],[f358])).
fof(f2851,plain,(
  ~c3_1(a165) | ~c1_1(a165) | (~spl0_149 | ~spl0_158)),
  inference(resolution,[],[f915,f2656])).
fof(f2656,plain,(
  c0_1(a165) | ~spl0_158),
  inference(avatar_component_clause,[],[f2654])).
fof(f2823,plain,(
  spl0_91 | spl0_92 | ~spl0_135 | ~spl0_142),
  inference(avatar_split_clause,[],[f2818,f885,f860,f648,f643])).
fof(f643,plain,(
  spl0_91 <=> c1_1(a124)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_91])])).
fof(f648,plain,(
  spl0_92 <=> c3_1(a124)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_92])])).
fof(f2818,plain,(
  c1_1(a124) | (spl0_92 | ~spl0_135 | ~spl0_142)),
  inference(resolution,[],[f2785,f650])).
fof(f650,plain,(
  ~c3_1(a124) | spl0_92),
  inference(avatar_component_clause,[],[f648])).
fof(f2802,plain,(
  spl0_105 | spl0_106 | ~spl0_122 | ~spl0_147),
  inference(avatar_contradiction_clause,[],[f2801])).
fof(f2801,plain,(
  $false | (spl0_105 | spl0_106 | ~spl0_122 | ~spl0_147)),
  inference(subsumption_resolution,[],[f2796,f715])).
fof(f715,plain,(
  ~c2_1(a104) | spl0_105),
  inference(avatar_component_clause,[],[f713])).
fof(f713,plain,(
  spl0_105 <=> c2_1(a104)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_105])])).
fof(f2796,plain,(
  c2_1(a104) | (spl0_106 | ~spl0_122 | ~spl0_147)),
  inference(resolution,[],[f2749,f720])).
fof(f720,plain,(
  ~c3_1(a104) | spl0_106),
  inference(avatar_component_clause,[],[f718])).
fof(f718,plain,(
  spl0_106 <=> c3_1(a104)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_106])])).
fof(f2749,plain,(
  ( ! [X0] : (c3_1(X0) | c2_1(X0)) ) | (~spl0_122 | ~spl0_147)),
  inference(subsumption_resolution,[],[f799,f907])).
fof(f2747,plain,(
  ~spl0_76 | ~spl0_77 | ~spl0_131 | ~spl0_143),
  inference(avatar_contradiction_clause,[],[f2746])).
fof(f2746,plain,(
  $false | (~spl0_76 | ~spl0_77 | ~spl0_131 | ~spl0_143)),
  inference(subsumption_resolution,[],[f2739,f570])).
fof(f2739,plain,(
  ~c2_1(a97) | (~spl0_77 | ~spl0_131 | ~spl0_143)),
  inference(resolution,[],[f2662,f575])).
fof(f2662,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c2_1(X0)) ) | (~spl0_131 | ~spl0_143)),
  inference(subsumption_resolution,[],[f845,f890])).
fof(f2719,plain,(
  ~spl0_59 | spl0_96 | spl0_97 | ~spl0_126),
  inference(avatar_contradiction_clause,[],[f2718])).
fof(f2718,plain,(
  $false | (~spl0_59 | spl0_96 | spl0_97 | ~spl0_126)),
  inference(subsumption_resolution,[],[f2717,f670])).
fof(f670,plain,(
  ~c2_1(a115) | spl0_96),
  inference(avatar_component_clause,[],[f668])).
fof(f668,plain,(
  spl0_96 <=> c2_1(a115)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_96])])).
fof(f2717,plain,(
  c2_1(a115) | (~spl0_59 | spl0_97 | ~spl0_126)),
  inference(subsumption_resolution,[],[f2713,f675])).
fof(f675,plain,(
  ~c3_1(a115) | spl0_97),
  inference(avatar_component_clause,[],[f673])).
fof(f2713,plain,(
  c3_1(a115) | c2_1(a115) | (~spl0_59 | ~spl0_126)),
  inference(resolution,[],[f485,f817])).
fof(f2709,plain,(
  spl0_159 | spl0_101 | ~spl0_64 | ~spl0_137),
  inference(avatar_split_clause,[],[f2639,f868,f508,f693,f2706])).
fof(f868,plain,(
  spl0_137 <=> ! [X0] : (c3_1(X0) | ~c2_1(X0) | c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_137])])).
fof(f2639,plain,(
  c3_1(a109) | c1_1(a109) | (~spl0_64 | ~spl0_137)),
  inference(resolution,[],[f869,f510])).
fof(f869,plain,(
  ( ! [X0] : (~c2_1(X0) | c3_1(X0) | c1_1(X0)) ) | ~spl0_137),
  inference(avatar_component_clause,[],[f868])).
fof(f2702,plain,(
  ~spl0_65 | ~spl0_66 | ~spl0_144 | ~spl0_149),
  inference(avatar_contradiction_clause,[],[f2701])).
fof(f2701,plain,(
  $false | (~spl0_65 | ~spl0_66 | ~spl0_144 | ~spl0_149)),
  inference(subsumption_resolution,[],[f2696,f515])).
fof(f2696,plain,(
  ~c1_1(a106) | (~spl0_66 | ~spl0_144 | ~spl0_149)),
  inference(resolution,[],[f2661,f520])).
fof(f2661,plain,(
  ( ! [X1] : (~c3_1(X1) | ~c1_1(X1)) ) | (~spl0_144 | ~spl0_149)),
  inference(subsumption_resolution,[],[f894,f915])).
fof(f894,plain,(
  ( ! [X1] : (~c3_1(X1) | ~c1_1(X1) | c0_1(X1)) ) | ~spl0_144),
  inference(avatar_component_clause,[],[f893])).
fof(f893,plain,(
  spl0_144 <=> ! [X1] : (c0_1(X1) | ~c1_1(X1) | ~c3_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_144])])).
fof(f2657,plain,(
  spl0_158 | ~spl0_33 | ~spl0_34 | ~spl0_143),
  inference(avatar_split_clause,[],[f2370,f889,f358,f353,f2654])).
fof(f353,plain,(
  spl0_33 <=> c2_1(a165)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_33])])).
fof(f2370,plain,(
  ~c2_1(a165) | c0_1(a165) | (~spl0_34 | ~spl0_143)),
  inference(resolution,[],[f890,f360])).
fof(f2652,plain,(
  ~spl0_74 | spl0_114 | spl0_115 | ~spl0_137),
  inference(avatar_contradiction_clause,[],[f2651])).
fof(f2651,plain,(
  $false | (~spl0_74 | spl0_114 | spl0_115 | ~spl0_137)),
  inference(subsumption_resolution,[],[f2650,f760])).
fof(f760,plain,(
  ~c1_1(a99) | spl0_114),
  inference(avatar_component_clause,[],[f758])).
fof(f758,plain,(
  spl0_114 <=> c1_1(a99)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_114])])).
fof(f2650,plain,(
  c1_1(a99) | (~spl0_74 | spl0_115 | ~spl0_137)),
  inference(subsumption_resolution,[],[f2643,f765])).
fof(f765,plain,(
  ~c3_1(a99) | spl0_115),
  inference(avatar_component_clause,[],[f763])).
fof(f763,plain,(
  spl0_115 <=> c3_1(a99)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_115])])).
fof(f2643,plain,(
  c3_1(a99) | c1_1(a99) | (~spl0_74 | ~spl0_137)),
  inference(resolution,[],[f869,f560])).
fof(f560,plain,(
  c2_1(a99) | ~spl0_74),
  inference(avatar_component_clause,[],[f558])).
fof(f558,plain,(
  spl0_74 <=> c2_1(a99)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_74])])).
fof(f2627,plain,(
  ~spl0_62 | ~spl0_63 | spl0_99 | ~spl0_143),
  inference(avatar_split_clause,[],[f2626,f889,f683,f503,f498])).
fof(f503,plain,(
  spl0_63 <=> c3_1(a110)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_63])])).
fof(f2626,plain,(
  ~c2_1(a110) | (~spl0_63 | spl0_99 | ~spl0_143)),
  inference(subsumption_resolution,[],[f2377,f685])).
fof(f2377,plain,(
  ~c2_1(a110) | c0_1(a110) | (~spl0_63 | ~spl0_143)),
  inference(resolution,[],[f890,f505])).
fof(f505,plain,(
  c3_1(a110) | ~spl0_63),
  inference(avatar_component_clause,[],[f503])).
fof(f2625,plain,(
  ~spl0_70 | ~spl0_71 | spl0_109 | ~spl0_130),
  inference(avatar_contradiction_clause,[],[f2624])).
fof(f2624,plain,(
  $false | (~spl0_70 | ~spl0_71 | spl0_109 | ~spl0_130)),
  inference(subsumption_resolution,[],[f2623,f735])).
fof(f2623,plain,(
  c3_1(a102) | (~spl0_70 | ~spl0_71 | ~spl0_130)),
  inference(subsumption_resolution,[],[f2607,f540])).
fof(f2607,plain,(
  ~c1_1(a102) | c3_1(a102) | (~spl0_71 | ~spl0_130)),
  inference(resolution,[],[f840,f545])).
fof(f545,plain,(
  c2_1(a102) | ~spl0_71),
  inference(avatar_component_clause,[],[f543])).
fof(f543,plain,(
  spl0_71 <=> c2_1(a102)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_71])])).
fof(f2622,plain,(
  ~spl0_54 | ~spl0_55 | spl0_89 | ~spl0_129 | ~spl0_130 | ~spl0_140),
  inference(avatar_contradiction_clause,[],[f2621])).
fof(f2621,plain,(
  $false | (~spl0_54 | ~spl0_55 | spl0_89 | ~spl0_129 | ~spl0_130 | ~spl0_140)),
  inference(subsumption_resolution,[],[f2620,f635])).
fof(f2620,plain,(
  c3_1(a125) | (~spl0_54 | ~spl0_55 | ~spl0_129 | ~spl0_130 | ~spl0_140)),
  inference(subsumption_resolution,[],[f2603,f465])).
fof(f2603,plain,(
  ~c1_1(a125) | c3_1(a125) | (~spl0_54 | ~spl0_129 | ~spl0_130 | ~spl0_140)),
  inference(resolution,[],[f840,f2130])).
fof(f2130,plain,(
  c2_1(a125) | (~spl0_54 | ~spl0_129 | ~spl0_140)),
  inference(resolution,[],[f2072,f460])).
fof(f2072,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0)) ) | (~spl0_129 | ~spl0_140)),
  inference(subsumption_resolution,[],[f836,f879])).
fof(f2619,plain,(
  ~spl0_52 | ~spl0_53 | spl0_88 | ~spl0_130 | ~spl0_142),
  inference(avatar_contradiction_clause,[],[f2618])).
fof(f2618,plain,(
  $false | (~spl0_52 | ~spl0_53 | spl0_88 | ~spl0_130 | ~spl0_142)),
  inference(subsumption_resolution,[],[f2617,f630])).
fof(f630,plain,(
  ~c3_1(a127) | spl0_88),
  inference(avatar_component_clause,[],[f628])).
fof(f628,plain,(
  spl0_88 <=> c3_1(a127)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_88])])).
fof(f2617,plain,(
  c3_1(a127) | (~spl0_52 | ~spl0_53 | spl0_88 | ~spl0_130 | ~spl0_142)),
  inference(subsumption_resolution,[],[f2602,f2364])).
fof(f2364,plain,(
  c1_1(a127) | (~spl0_52 | spl0_88 | ~spl0_142)),
  inference(subsumption_resolution,[],[f2361,f630])).
fof(f2361,plain,(
  c3_1(a127) | c1_1(a127) | (~spl0_52 | ~spl0_142)),
  inference(resolution,[],[f886,f450])).
fof(f2602,plain,(
  ~c1_1(a127) | c3_1(a127) | (~spl0_53 | ~spl0_130)),
  inference(resolution,[],[f840,f455])).
fof(f455,plain,(
  c2_1(a127) | ~spl0_53),
  inference(avatar_component_clause,[],[f453])).
fof(f2530,plain,(
  spl0_110 | spl0_111 | ~spl0_123 | ~spl0_129 | ~spl0_140),
  inference(avatar_split_clause,[],[f2507,f878,f835,f802,f743,f738])).
fof(f802,plain,(
  spl0_123 <=> ! [X0] : (c2_1(X0) | c0_1(X0) | c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_123])])).
fof(f2507,plain,(
  c1_1(a101) | (spl0_111 | ~spl0_123 | ~spl0_129 | ~spl0_140)),
  inference(resolution,[],[f2496,f745])).
fof(f2496,plain,(
  ( ! [X0] : (c2_1(X0) | c1_1(X0)) ) | (~spl0_123 | ~spl0_129 | ~spl0_140)),
  inference(subsumption_resolution,[],[f803,f2072])).
fof(f803,plain,(
  ( ! [X0] : (c0_1(X0) | c2_1(X0) | c1_1(X0)) ) | ~spl0_123),
  inference(avatar_component_clause,[],[f802])).
fof(f2527,plain,(
  spl0_95 | ~spl0_57 | ~spl0_127 | ~spl0_154),
  inference(avatar_split_clause,[],[f2526,f2264,f822,f473,f663])).
fof(f2526,plain,(
  c0_1(a116) | (~spl0_57 | ~spl0_127 | ~spl0_154)),
  inference(subsumption_resolution,[],[f2419,f475])).
fof(f2419,plain,(
  ~c1_1(a116) | c0_1(a116) | (~spl0_127 | ~spl0_154)),
  inference(resolution,[],[f2265,f823])).
fof(f2265,plain,(
  c2_1(a116) | ~spl0_154),
  inference(avatar_component_clause,[],[f2264])).
fof(f2478,plain,(
  ~spl0_65 | ~spl0_66 | spl0_102 | ~spl0_129 | ~spl0_140 | ~spl0_144),
  inference(avatar_contradiction_clause,[],[f2477])).
fof(f2477,plain,(
  $false | (~spl0_65 | ~spl0_66 | spl0_102 | ~spl0_129 | ~spl0_140 | ~spl0_144)),
  inference(subsumption_resolution,[],[f2476,f700])).
fof(f2476,plain,(
  c2_1(a106) | (~spl0_65 | ~spl0_66 | ~spl0_129 | ~spl0_140 | ~spl0_144)),
  inference(resolution,[],[f2408,f2072])).
fof(f2408,plain,(
  c0_1(a106) | (~spl0_65 | ~spl0_66 | ~spl0_144)),
  inference(subsumption_resolution,[],[f2394,f515])).
fof(f2394,plain,(
  ~c1_1(a106) | c0_1(a106) | (~spl0_66 | ~spl0_144)),
  inference(resolution,[],[f894,f520])).
fof(f2464,plain,(
  spl0_104 | spl0_105 | spl0_106 | ~spl0_122 | ~spl0_136 | ~spl0_144),
  inference(avatar_contradiction_clause,[],[f2463])).
fof(f2463,plain,(
  $false | (spl0_104 | spl0_105 | spl0_106 | ~spl0_122 | ~spl0_136 | ~spl0_144)),
  inference(subsumption_resolution,[],[f2453,f710])).
fof(f710,plain,(
  ~c0_1(a104) | spl0_104),
  inference(avatar_component_clause,[],[f708])).
fof(f708,plain,(
  spl0_104 <=> c0_1(a104)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_104])])).
fof(f2453,plain,(
  c0_1(a104) | (spl0_105 | spl0_106 | ~spl0_122 | ~spl0_136 | ~spl0_144)),
  inference(resolution,[],[f2437,f2269])).
fof(f2269,plain,(
  c1_1(a104) | (spl0_105 | spl0_106 | ~spl0_122)),
  inference(subsumption_resolution,[],[f2213,f715])).
fof(f2213,plain,(
  c1_1(a104) | c2_1(a104) | (spl0_106 | ~spl0_122)),
  inference(resolution,[],[f799,f720])).
fof(f2437,plain,(
  ( ! [X1] : (~c1_1(X1) | c0_1(X1)) ) | (~spl0_136 | ~spl0_144)),
  inference(subsumption_resolution,[],[f865,f894])).
fof(f2436,plain,(
  ~spl0_65 | ~spl0_66 | spl0_102 | ~spl0_151),
  inference(avatar_contradiction_clause,[],[f2435])).
fof(f2435,plain,(
  $false | (~spl0_65 | ~spl0_66 | spl0_102 | ~spl0_151)),
  inference(subsumption_resolution,[],[f2434,f700])).
fof(f2434,plain,(
  c2_1(a106) | (~spl0_65 | ~spl0_66 | ~spl0_151)),
  inference(subsumption_resolution,[],[f2429,f515])).
fof(f2429,plain,(
  ~c1_1(a106) | c2_1(a106) | (~spl0_66 | ~spl0_151)),
  inference(resolution,[],[f928,f520])).
fof(f2411,plain,(
  spl0_95 | ~spl0_57 | ~spl0_58 | ~spl0_144),
  inference(avatar_split_clause,[],[f2403,f893,f478,f473,f663])).
fof(f2403,plain,(
  c0_1(a116) | (~spl0_57 | ~spl0_58 | ~spl0_144)),
  inference(subsumption_resolution,[],[f2391,f475])).
fof(f2391,plain,(
  ~c1_1(a116) | c0_1(a116) | (~spl0_58 | ~spl0_144)),
  inference(resolution,[],[f894,f480])).
fof(f2407,plain,(
  ~spl0_155 | ~spl0_63 | spl0_99 | ~spl0_144),
  inference(avatar_split_clause,[],[f2406,f893,f683,f503,f2319])).
fof(f2406,plain,(
  ~c1_1(a110) | (~spl0_63 | spl0_99 | ~spl0_144)),
  inference(subsumption_resolution,[],[f2393,f685])).
fof(f2393,plain,(
  ~c1_1(a110) | c0_1(a110) | (~spl0_63 | ~spl0_144)),
  inference(resolution,[],[f894,f505])).
fof(f2323,plain,(
  spl0_120 | spl0_119 | spl0_121 | ~spl0_122),
  inference(avatar_split_clause,[],[f2214,f798,f793,f783,f788])).
fof(f788,plain,(
  spl0_120 <=> c2_1(a96)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_120])])).
fof(f783,plain,(
  spl0_119 <=> c1_1(a96)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_119])])).
fof(f793,plain,(
  spl0_121 <=> c3_1(a96)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_121])])).
fof(f2214,plain,(
  c1_1(a96) | c2_1(a96) | (spl0_121 | ~spl0_122)),
  inference(resolution,[],[f799,f795])).
fof(f795,plain,(
  ~c3_1(a96) | spl0_121),
  inference(avatar_component_clause,[],[f793])).
fof(f2322,plain,(
  spl0_155 | ~spl0_62 | ~spl0_63 | ~spl0_125),
  inference(avatar_split_clause,[],[f2224,f812,f503,f498,f2319])).
fof(f2224,plain,(
  ~c2_1(a110) | c1_1(a110) | (~spl0_63 | ~spl0_125)),
  inference(resolution,[],[f813,f505])).
fof(f2314,plain,(
  ~spl0_64 | spl0_100 | ~spl0_127 | ~spl0_148),
  inference(avatar_contradiction_clause,[],[f2313])).
fof(f2313,plain,(
  $false | (~spl0_64 | spl0_100 | ~spl0_127 | ~spl0_148)),
  inference(subsumption_resolution,[],[f2305,f690])).
fof(f2305,plain,(
  c0_1(a109) | (~spl0_64 | ~spl0_127 | ~spl0_148)),
  inference(resolution,[],[f2293,f510])).
fof(f2293,plain,(
  ( ! [X1] : (~c2_1(X1) | c0_1(X1)) ) | (~spl0_127 | ~spl0_148)),
  inference(subsumption_resolution,[],[f911,f823])).
fof(f911,plain,(
  ( ! [X1] : (c1_1(X1) | ~c2_1(X1) | c0_1(X1)) ) | ~spl0_148),
  inference(avatar_component_clause,[],[f910])).
fof(f910,plain,(
  spl0_148 <=> ! [X1] : (c1_1(X1) | ~c2_1(X1) | c0_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_148])])).
fof(f2285,plain,(
  ~spl0_33 | ~spl0_34 | ~spl0_131 | ~spl0_143),
  inference(avatar_contradiction_clause,[],[f2284])).
fof(f2284,plain,(
  $false | (~spl0_33 | ~spl0_34 | ~spl0_131 | ~spl0_143)),
  inference(subsumption_resolution,[],[f2274,f355])).
fof(f355,plain,(
  c2_1(a165) | ~spl0_33),
  inference(avatar_component_clause,[],[f353])).
fof(f2274,plain,(
  ~c2_1(a165) | (~spl0_34 | ~spl0_131 | ~spl0_143)),
  inference(resolution,[],[f2270,f360])).
fof(f2270,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c2_1(X0)) ) | (~spl0_131 | ~spl0_143)),
  inference(subsumption_resolution,[],[f845,f890])).
fof(f2258,plain,(
  ~spl0_76 | ~spl0_77 | spl0_118 | ~spl0_125),
  inference(avatar_contradiction_clause,[],[f2257])).
fof(f2257,plain,(
  $false | (~spl0_76 | ~spl0_77 | spl0_118 | ~spl0_125)),
  inference(subsumption_resolution,[],[f2256,f780])).
fof(f2256,plain,(
  c1_1(a97) | (~spl0_76 | ~spl0_77 | ~spl0_125)),
  inference(subsumption_resolution,[],[f2228,f570])).
fof(f2228,plain,(
  ~c2_1(a97) | c1_1(a97) | (~spl0_77 | ~spl0_125)),
  inference(resolution,[],[f813,f575])).
fof(f2200,plain,(
  ~spl0_58 | spl0_95 | ~spl0_143 | ~spl0_150),
  inference(avatar_contradiction_clause,[],[f2199])).
fof(f2199,plain,(
  $false | (~spl0_58 | spl0_95 | ~spl0_143 | ~spl0_150)),
  inference(subsumption_resolution,[],[f2185,f665])).
fof(f2185,plain,(
  c0_1(a116) | (~spl0_58 | ~spl0_143 | ~spl0_150)),
  inference(resolution,[],[f2168,f480])).
fof(f2168,plain,(
  ( ! [X0] : (~c3_1(X0) | c0_1(X0)) ) | (~spl0_143 | ~spl0_150)),
  inference(subsumption_resolution,[],[f890,f918])).
fof(f2132,plain,(
  spl0_87 | ~spl0_50 | ~spl0_129 | ~spl0_140),
  inference(avatar_split_clause,[],[f2128,f878,f835,f438,f623])).
fof(f2128,plain,(
  c2_1(a133) | (~spl0_50 | ~spl0_129 | ~spl0_140)),
  inference(resolution,[],[f2072,f440])).
fof(f2122,plain,(
  ~spl0_52 | ~spl0_129 | ~spl0_131 | ~spl0_140 | ~spl0_141 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f2119])).
fof(f2119,plain,(
  $false | (~spl0_52 | ~spl0_129 | ~spl0_131 | ~spl0_140 | ~spl0_141 | ~spl0_146)),
  inference(resolution,[],[f2115,f450])).
fof(f2115,plain,(
  ( ! [X0] : (~c0_1(X0)) ) | (~spl0_129 | ~spl0_131 | ~spl0_140 | ~spl0_141 | ~spl0_146)),
  inference(subsumption_resolution,[],[f2114,f2111])).
fof(f2111,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c0_1(X0)) ) | (~spl0_129 | ~spl0_131 | ~spl0_140)),
  inference(subsumption_resolution,[],[f845,f2072])).
fof(f2114,plain,(
  ( ! [X0] : (~c0_1(X0) | c3_1(X0)) ) | (~spl0_140 | ~spl0_141 | ~spl0_146)),
  inference(subsumption_resolution,[],[f903,f1479])).
fof(f1479,plain,(
  ( ! [X0] : (~c0_1(X0) | c1_1(X0)) ) | (~spl0_140 | ~spl0_141)),
  inference(subsumption_resolution,[],[f883,f879])).
fof(f2121,plain,(
  ~spl0_54 | ~spl0_129 | ~spl0_131 | ~spl0_140 | ~spl0_141 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f2120])).
fof(f2120,plain,(
  $false | (~spl0_54 | ~spl0_129 | ~spl0_131 | ~spl0_140 | ~spl0_141 | ~spl0_146)),
  inference(resolution,[],[f2115,f460])).
fof(f1891,plain,(
  spl0_110 | spl0_111 | ~spl0_123 | ~spl0_140 | ~spl0_141),
  inference(avatar_split_clause,[],[f1875,f882,f878,f802,f743,f738])).
fof(f1875,plain,(
  c1_1(a101) | (spl0_111 | ~spl0_123 | ~spl0_140 | ~spl0_141)),
  inference(resolution,[],[f1857,f745])).
fof(f1857,plain,(
  ( ! [X0] : (c2_1(X0) | c1_1(X0)) ) | (~spl0_123 | ~spl0_140 | ~spl0_141)),
  inference(subsumption_resolution,[],[f803,f1479])).
fof(f1700,plain,(
  ~spl0_66 | spl0_102 | ~spl0_145 | ~spl0_150),
  inference(avatar_contradiction_clause,[],[f1699])).
fof(f1699,plain,(
  $false | (~spl0_66 | spl0_102 | ~spl0_145 | ~spl0_150)),
  inference(subsumption_resolution,[],[f1696,f700])).
fof(f1696,plain,(
  c2_1(a106) | (~spl0_66 | ~spl0_145 | ~spl0_150)),
  inference(resolution,[],[f1668,f520])).
fof(f1668,plain,(
  ( ! [X0] : (~c3_1(X0) | c2_1(X0)) ) | (~spl0_145 | ~spl0_150)),
  inference(subsumption_resolution,[],[f899,f918])).
fof(f1608,plain,(
  spl0_91 | spl0_92 | ~spl0_135 | ~spl0_140 | ~spl0_141),
  inference(avatar_contradiction_clause,[],[f1607])).
fof(f1607,plain,(
  $false | (spl0_91 | spl0_92 | ~spl0_135 | ~spl0_140 | ~spl0_141)),
  inference(subsumption_resolution,[],[f1602,f645])).
fof(f645,plain,(
  ~c1_1(a124) | spl0_91),
  inference(avatar_component_clause,[],[f643])).
fof(f1602,plain,(
  c1_1(a124) | (spl0_92 | ~spl0_135 | ~spl0_140 | ~spl0_141)),
  inference(resolution,[],[f1542,f650])).
fof(f1542,plain,(
  ( ! [X1] : (c3_1(X1) | c1_1(X1)) ) | (~spl0_135 | ~spl0_140 | ~spl0_141)),
  inference(subsumption_resolution,[],[f861,f1479])).
fof(f1606,plain,(
  spl0_114 | spl0_115 | ~spl0_135 | ~spl0_140 | ~spl0_141),
  inference(avatar_contradiction_clause,[],[f1605])).
fof(f1605,plain,(
  $false | (spl0_114 | spl0_115 | ~spl0_135 | ~spl0_140 | ~spl0_141)),
  inference(subsumption_resolution,[],[f1600,f760])).
fof(f1600,plain,(
  c1_1(a99) | (spl0_115 | ~spl0_135 | ~spl0_140 | ~spl0_141)),
  inference(resolution,[],[f1542,f765])).
fof(f1587,plain,(
  spl0_98 | ~spl0_60 | ~spl0_140 | ~spl0_141),
  inference(avatar_split_clause,[],[f1579,f882,f878,f488,f678])).
fof(f1579,plain,(
  c1_1(a112) | (~spl0_60 | ~spl0_140 | ~spl0_141)),
  inference(resolution,[],[f1479,f490])).
fof(f1586,plain,(
  spl0_107 | ~spl0_69 | ~spl0_140 | ~spl0_141),
  inference(avatar_split_clause,[],[f1581,f882,f878,f533,f723])).
fof(f723,plain,(
  spl0_107 <=> c1_1(a103)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_107])])).
fof(f533,plain,(
  spl0_69 <=> c0_1(a103)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_69])])).
fof(f1581,plain,(
  c1_1(a103) | (~spl0_69 | ~spl0_140 | ~spl0_141)),
  inference(resolution,[],[f1479,f535])).
fof(f535,plain,(
  c0_1(a103) | ~spl0_69),
  inference(avatar_component_clause,[],[f533])).
fof(f1515,plain,(
  spl0_102 | ~spl0_65 | ~spl0_129 | ~spl0_133 | ~spl0_140),
  inference(avatar_split_clause,[],[f1507,f878,f853,f835,f513,f698])).
fof(f1507,plain,(
  c2_1(a106) | (~spl0_65 | ~spl0_129 | ~spl0_133 | ~spl0_140)),
  inference(resolution,[],[f1478,f515])).
fof(f1478,plain,(
  ( ! [X0] : (~c1_1(X0) | c2_1(X0)) ) | (~spl0_129 | ~spl0_133 | ~spl0_140)),
  inference(subsumption_resolution,[],[f854,f1477])).
fof(f1477,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0)) ) | (~spl0_129 | ~spl0_140)),
  inference(subsumption_resolution,[],[f836,f879])).
fof(f1353,plain,(
  ~spl0_54 | spl0_89 | ~spl0_142 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f1352])).
fof(f1352,plain,(
  $false | (~spl0_54 | spl0_89 | ~spl0_142 | ~spl0_146)),
  inference(subsumption_resolution,[],[f1345,f635])).
fof(f1345,plain,(
  c3_1(a125) | (~spl0_54 | ~spl0_142 | ~spl0_146)),
  inference(resolution,[],[f1313,f460])).
fof(f1313,plain,(
  ( ! [X0] : (~c0_1(X0) | c3_1(X0)) ) | (~spl0_142 | ~spl0_146)),
  inference(subsumption_resolution,[],[f903,f886])).
fof(f1351,plain,(
  ~spl0_52 | spl0_88 | ~spl0_142 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f1350])).
fof(f1350,plain,(
  $false | (~spl0_52 | spl0_88 | ~spl0_142 | ~spl0_146)),
  inference(subsumption_resolution,[],[f1344,f630])).
fof(f1344,plain,(
  c3_1(a127) | (~spl0_52 | ~spl0_142 | ~spl0_146)),
  inference(resolution,[],[f1313,f450])).
fof(f1058,plain,(
  ~spl0_57 | spl0_95 | ~spl0_127 | ~spl0_133),
  inference(avatar_contradiction_clause,[],[f1057])).
fof(f1057,plain,(
  $false | (~spl0_57 | spl0_95 | ~spl0_127 | ~spl0_133)),
  inference(subsumption_resolution,[],[f1054,f665])).
fof(f1054,plain,(
  c0_1(a116) | (~spl0_57 | ~spl0_127 | ~spl0_133)),
  inference(resolution,[],[f1002,f475])).
fof(f1002,plain,(
  ( ! [X0] : (~c1_1(X0) | c0_1(X0)) ) | (~spl0_127 | ~spl0_133)),
  inference(subsumption_resolution,[],[f854,f823])).
fof(f1045,plain,(
  ~spl0_74 | spl0_114 | spl0_115 | ~spl0_138 | ~spl0_142),
  inference(avatar_contradiction_clause,[],[f1044])).
fof(f1044,plain,(
  $false | (~spl0_74 | spl0_114 | spl0_115 | ~spl0_138 | ~spl0_142)),
  inference(subsumption_resolution,[],[f1043,f760])).
fof(f1043,plain,(
  c1_1(a99) | (~spl0_74 | spl0_115 | ~spl0_138 | ~spl0_142)),
  inference(subsumption_resolution,[],[f1042,f765])).
fof(f1042,plain,(
  c3_1(a99) | c1_1(a99) | (~spl0_74 | spl0_115 | ~spl0_138 | ~spl0_142)),
  inference(resolution,[],[f981,f886])).
fof(f981,plain,(
  c0_1(a99) | (~spl0_74 | spl0_115 | ~spl0_138)),
  inference(subsumption_resolution,[],[f979,f765])).
fof(f979,plain,(
  c3_1(a99) | c0_1(a99) | (~spl0_74 | ~spl0_138)),
  inference(resolution,[],[f872,f560])).
fof(f1000,plain,(
  spl0_108 | ~spl0_69 | spl0_107 | ~spl0_142),
  inference(avatar_split_clause,[],[f999,f885,f723,f533,f728])).
fof(f728,plain,(
  spl0_108 <=> c3_1(a103)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_108])])).
fof(f999,plain,(
  c3_1(a103) | (~spl0_69 | spl0_107 | ~spl0_142)),
  inference(subsumption_resolution,[],[f989,f725])).
fof(f725,plain,(
  ~c1_1(a103) | spl0_107),
  inference(avatar_component_clause,[],[f723])).
fof(f989,plain,(
  c3_1(a103) | c1_1(a103) | (~spl0_69 | ~spl0_142)),
  inference(resolution,[],[f886,f535])).
fof(f994,plain,(
  ~spl0_59 | spl0_96 | spl0_97 | ~spl0_129 | ~spl0_142),
  inference(avatar_contradiction_clause,[],[f993])).
fof(f993,plain,(
  $false | (~spl0_59 | spl0_96 | spl0_97 | ~spl0_129 | ~spl0_142)),
  inference(subsumption_resolution,[],[f992,f971])).
fof(f971,plain,(
  ~c1_1(a115) | (~spl0_59 | spl0_96 | ~spl0_129)),
  inference(subsumption_resolution,[],[f965,f670])).
fof(f965,plain,(
  c2_1(a115) | ~c1_1(a115) | (~spl0_59 | ~spl0_129)),
  inference(resolution,[],[f836,f485])).
fof(f992,plain,(
  c1_1(a115) | (~spl0_59 | spl0_97 | ~spl0_142)),
  inference(subsumption_resolution,[],[f986,f675])).
fof(f986,plain,(
  c3_1(a115) | c1_1(a115) | (~spl0_59 | ~spl0_142)),
  inference(resolution,[],[f886,f485])).
fof(f991,plain,(
  spl0_122 | ~spl0_123 | ~spl0_142),
  inference(avatar_split_clause,[],[f990,f885,f802,f798])).
fof(f990,plain,(
  ( ! [X0] : (c3_1(X0) | c1_1(X0) | c2_1(X0)) ) | (~spl0_123 | ~spl0_142)),
  inference(duplicate_literal_removal,[],[f983])).
fof(f983,plain,(
  ( ! [X0] : (c3_1(X0) | c1_1(X0) | c2_1(X0) | c1_1(X0)) ) | (~spl0_123 | ~spl0_142)),
  inference(resolution,[],[f886,f803])).
fof(f973,plain,(
  ~spl0_59 | spl0_96 | spl0_97 | ~spl0_122 | ~spl0_129),
  inference(avatar_contradiction_clause,[],[f972])).
fof(f972,plain,(
  $false | (~spl0_59 | spl0_96 | spl0_97 | ~spl0_122 | ~spl0_129)),
  inference(subsumption_resolution,[],[f971,f945])).
fof(f945,plain,(
  c1_1(a115) | (spl0_96 | spl0_97 | ~spl0_122)),
  inference(subsumption_resolution,[],[f939,f670])).
fof(f939,plain,(
  c1_1(a115) | c2_1(a115) | (spl0_97 | ~spl0_122)),
  inference(resolution,[],[f799,f675])).
fof(f937,plain,(
  spl0_149 | ~spl0_2 | spl0_151 | spl0_129),
  inference(avatar_split_clause,[],[f180,f835,f927,f186,f914])).
fof(f186,plain,(
  spl0_2 <=> ndr1_0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f180,axiom,(
  ( ! [X2,X0,X1] : (c2_1(X2) | c2_1(X1) | ~c1_1(X2) | ~c0_1(X2) | ~c3_1(X1) | ~c1_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f936,plain,(
  spl0_152 | ~spl0_2 | spl0_145 | spl0_127),
  inference(avatar_split_clause,[],[f179,f822,f898,f186,f934])).
fof(f179,axiom,(
  ( ! [X2,X0,X1] : (c0_1(X2) | c2_1(X1) | ~c2_1(X2) | ~c1_1(X2) | ~c3_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f932,plain,(
  spl0_149 | ~spl0_2 | spl0_129 | spl0_139),
  inference(avatar_split_clause,[],[f178,f875,f835,f186,f914])).
fof(f178,axiom,(
  ( ! [X2,X0,X1] : (c2_1(X2) | c1_1(X2) | c2_1(X1) | ~c3_1(X2) | ~c1_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f931,plain,(
  ~spl0_2 | spl0_144 | spl0_127 | spl0_138),
  inference(avatar_split_clause,[],[f177,f871,f822,f893,f186])).
fof(f177,axiom,(
  ( ! [X2,X0,X1] : (c3_1(X2) | c0_1(X2) | c0_1(X1) | c0_1(X0) | ~c2_1(X2) | ~c2_1(X1) | ~c1_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f930,plain,(
  ~spl0_2 | spl0_151 | spl0_145 | spl0_136),
  inference(avatar_split_clause,[],[f176,f864,f898,f927,f186])).
fof(f176,axiom,(
  ( ! [X2,X0,X1] : (c3_1(X2) | c0_1(X2) | c2_1(X1) | c2_1(X0) | ~c1_1(X2) | ~c3_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f929,plain,(
  ~spl0_2 | spl0_151 | spl0_139 | spl0_140),
  inference(avatar_split_clause,[],[f175,f878,f875,f927,f186])).
fof(f175,axiom,(
  ( ! [X2,X0,X1] : (c2_1(X2) | c1_1(X2) | c2_1(X1) | c1_1(X1) | c2_1(X0) | ~c0_1(X2) | ~c3_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f925,plain,(
  ~spl0_2 | spl0_142 | spl0_143 | spl0_150),
  inference(avatar_split_clause,[],[f174,f917,f889,f885,f186])).
fof(f174,axiom,(
  ( ! [X2,X0,X1] : (c2_1(X2) | c0_1(X2) | c0_1(X1) | c3_1(X0) | c1_1(X0) | ~c3_1(X2) | ~c3_1(X1) | ~c2_1(X1) | ~ndr1_0 | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f924,plain,(
  ~spl0_2 | spl0_137 | spl0_144 | spl0_123),
  inference(avatar_split_clause,[],[f173,f802,f893,f868,f186])).
fof(f173,axiom,(
  ( ! [X2,X0,X1] : (c2_1(X2) | c1_1(X2) | c0_1(X2) | c0_1(X1) | c3_1(X0) | c1_1(X0) | ~c3_1(X1) | ~c1_1(X1) | ~ndr1_0 | ~c2_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f923,plain,(
  ~spl0_2 | spl0_147 | spl0_122 | spl0_138),
  inference(avatar_split_clause,[],[f172,f871,f798,f906,f186])).
fof(f172,axiom,(
  ( ! [X2,X0,X1] : (c3_1(X1) | c0_1(X1) | c3_1(X2) | c2_1(X2) | c1_1(X2) | c3_1(X0) | c2_1(X0) | ~c2_1(X1) | ~ndr1_0 | ~c1_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f922,plain,(
  spl0_131 | ~spl0_2 | spl0_127 | spl0_30),
  inference(avatar_split_clause,[],[f171,f326,f822,f186,f844])).
fof(f326,plain,(
  spl0_30 <=> hskp1),
  introduced(avatar_definition,[new_symbols(naming,[spl0_30])])).
fof(f171,axiom,(
  ( ! [X0,X1] : (hskp1 | c0_1(X1) | ~c2_1(X1) | ~c1_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f920,plain,(
  ~spl0_2 | spl0_146 | spl0_128 | spl0_15),
  inference(avatar_split_clause,[],[f169,f251,f826,f902,f186])).
fof(f251,plain,(
  spl0_15 <=> hskp16),
  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).
fof(f169,axiom,(
  ( ! [X0,X1] : (hskp16 | c1_1(X1) | c3_1(X0) | ~c3_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f919,plain,(
  spl0_149 | ~spl0_2 | spl0_150 | spl0_23),
  inference(avatar_split_clause,[],[f168,f291,f917,f186,f914])).
fof(f291,plain,(
  spl0_23 <=> hskp8),
  introduced(avatar_definition,[new_symbols(naming,[spl0_23])])).
fof(f168,axiom,(
  ( ! [X0,X1] : (hskp8 | c2_1(X1) | c0_1(X1) | ~c3_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f912,plain,(
  spl0_131 | ~spl0_2 | spl0_148 | spl0_30),
  inference(avatar_split_clause,[],[f167,f326,f910,f186,f844])).
fof(f167,axiom,(
  ( ! [X0,X1] : (hskp1 | c1_1(X1) | c0_1(X1) | ~c2_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f908,plain,(
  ~spl0_2 | spl0_130 | spl0_147 | spl0_26),
  inference(avatar_split_clause,[],[f166,f306,f906,f839,f186])).
fof(f306,plain,(
  spl0_26 <=> hskp5),
  introduced(avatar_definition,[new_symbols(naming,[spl0_26])])).
fof(f166,axiom,(
  ( ! [X0,X1] : (hskp5 | c3_1(X1) | c2_1(X1) | c3_1(X0) | ~c1_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f904,plain,(
  ~spl0_2 | spl0_146 | spl0_137 | spl0_21),
  inference(avatar_split_clause,[],[f165,f281,f868,f902,f186])).
fof(f281,plain,(
  spl0_21 <=> hskp10),
  introduced(avatar_definition,[new_symbols(naming,[spl0_21])])).
fof(f165,axiom,(
  ( ! [X0,X1] : (hskp10 | c3_1(X1) | c1_1(X1) | c3_1(X0) | ~c2_1(X1) | ~ndr1_0 | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f900,plain,(
  ~spl0_2 | spl0_145 | spl0_140 | spl0_16),
  inference(avatar_split_clause,[],[f164,f256,f878,f898,f186])).
fof(f256,plain,(
  spl0_16 <=> hskp15),
  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).
fof(f164,axiom,(
  ( ! [X0,X1] : (hskp15 | c2_1(X1) | c1_1(X1) | c2_1(X0) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f896,plain,(
  spl0_131 | ~spl0_2 | spl0_122 | spl0_28),
  inference(avatar_split_clause,[],[f163,f316,f798,f186,f844])).
fof(f316,plain,(
  spl0_28 <=> hskp3),
  introduced(avatar_definition,[new_symbols(naming,[spl0_28])])).
fof(f163,axiom,(
  ( ! [X0,X1] : (hskp3 | c3_1(X1) | c2_1(X1) | c1_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f895,plain,(
  ~spl0_2 | spl0_142 | spl0_144 | spl0_19),
  inference(avatar_split_clause,[],[f162,f271,f893,f885,f186])).
fof(f271,plain,(
  spl0_19 <=> hskp12),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f162,axiom,(
  ( ! [X0,X1] : (hskp12 | c0_1(X1) | c3_1(X0) | c1_1(X0) | ~c3_1(X1) | ~c1_1(X1) | ~ndr1_0 | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f891,plain,(
  ~spl0_2 | spl0_143 | spl0_133 | spl0_24),
  inference(avatar_split_clause,[],[f161,f296,f853,f889,f186])).
fof(f296,plain,(
  spl0_24 <=> hskp7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_24])])).
fof(f161,axiom,(
  ( ! [X0,X1] : (hskp7 | c2_1(X1) | c0_1(X1) | c0_1(X0) | ~c1_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f887,plain,(
  ~spl0_2 | spl0_141 | spl0_142),
  inference(avatar_split_clause,[],[f160,f885,f882,f186])).
fof(f160,axiom,(
  ( ! [X0,X1] : (c3_1(X1) | c1_1(X1) | c1_1(X0) | ~c0_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f880,plain,(
  ~spl0_2 | spl0_139 | spl0_140 | spl0_17),
  inference(avatar_split_clause,[],[f159,f261,f878,f875,f186])).
fof(f261,plain,(
  spl0_17 <=> hskp14),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f159,axiom,(
  ( ! [X0,X1] : (hskp14 | c2_1(X1) | c1_1(X1) | c2_1(X0) | c1_1(X0) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f873,plain,(
  ~spl0_2 | spl0_137 | spl0_138 | spl0_21),
  inference(avatar_split_clause,[],[f158,f281,f871,f868,f186])).
fof(f158,axiom,(
  ( ! [X0,X1] : (hskp10 | c3_1(X1) | c0_1(X1) | c3_1(X0) | c1_1(X0) | ~c2_1(X1) | ~ndr1_0 | ~c2_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f862,plain,(
  ~spl0_2 | spl0_129 | spl0_135 | spl0_31),
  inference(avatar_split_clause,[],[f156,f331,f860,f835,f186])).
fof(f331,plain,(
  spl0_31 <=> hskp0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_31])])).
fof(f156,axiom,(
  ( ! [X0,X1] : (hskp0 | c3_1(X1) | c1_1(X1) | c0_1(X1) | c2_1(X0) | ~ndr1_0 | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f858,plain,(
  ~spl0_2 | spl0_133 | spl0_134 | spl0_25),
  inference(avatar_split_clause,[],[f155,f301,f856,f853,f186])).
fof(f301,plain,(
  spl0_25 <=> hskp6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_25])])).
fof(f155,axiom,(
  ( ! [X0,X1] : (hskp6 | c3_1(X1) | c2_1(X1) | c0_1(X1) | c2_1(X0) | c0_1(X0) | ~ndr1_0 | ~c1_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f837,plain,(
  ~spl0_2 | spl0_129 | spl0_25 | spl0_28),
  inference(avatar_split_clause,[],[f149,f316,f301,f835,f186])).
fof(f149,axiom,(
  ( ! [X0] : (hskp3 | hskp6 | c2_1(X0) | ~ndr1_0 | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f832,plain,(
  ~spl0_2 | spl0_125 | spl0_12 | spl0_26),
  inference(avatar_split_clause,[],[f147,f306,f236,f812,f186])).
fof(f236,plain,(
  spl0_12 <=> hskp19),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f147,axiom,(
  ( ! [X0] : (hskp5 | hskp19 | c1_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f831,plain,(
  ~spl0_2 | spl0_125 | spl0_13 | spl0_17),
  inference(avatar_split_clause,[],[f146,f261,f241,f812,f186])).
fof(f241,plain,(
  spl0_13 <=> hskp18),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f146,axiom,(
  ( ! [X0] : (hskp14 | hskp18 | c1_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f828,plain,(
  ~spl0_2 | spl0_128 | spl0_17 | spl0_21),
  inference(avatar_split_clause,[],[f143,f281,f261,f826,f186])).
fof(f143,axiom,(
  ( ! [X0] : (hskp10 | hskp14 | c1_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f824,plain,(
  ~spl0_2 | spl0_127 | spl0_5 | spl0_20),
  inference(avatar_split_clause,[],[f142,f276,f201,f822,f186])).
fof(f201,plain,(
  spl0_5 <=> hskp26),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f276,plain,(
  spl0_20 <=> hskp11),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f142,axiom,(
  ( ! [X0] : (hskp11 | hskp26 | c0_1(X0) | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f820,plain,(
  ~spl0_2 | spl0_126 | spl0_30 | spl0_14),
  inference(avatar_split_clause,[],[f141,f246,f326,f816,f186])).
fof(f246,plain,(
  spl0_14 <=> hskp17),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f141,axiom,(
  ( ! [X0] : (hskp17 | hskp1 | c3_1(X0) | c2_1(X0) | ~ndr1_0 | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f819,plain,(
  ~spl0_2 | spl0_126 | spl0_19 | spl0_26),
  inference(avatar_split_clause,[],[f140,f306,f271,f816,f186])).
fof(f140,axiom,(
  ( ! [X0] : (hskp5 | hskp12 | c3_1(X0) | c2_1(X0) | ~ndr1_0 | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f818,plain,(
  ~spl0_2 | spl0_126 | spl0_11 | spl0_21),
  inference(avatar_split_clause,[],[f139,f281,f231,f816,f186])).
fof(f231,plain,(
  spl0_11 <=> hskp20),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f139,axiom,(
  ( ! [X0] : (hskp10 | hskp20 | c3_1(X0) | c2_1(X0) | ~ndr1_0 | ~c0_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f814,plain,(
  ~spl0_2 | spl0_125 | spl0_18),
  inference(avatar_split_clause,[],[f138,f266,f812,f186])).
fof(f266,plain,(
  spl0_18 <=> hskp13),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f138,axiom,(
  ( ! [X0] : (hskp13 | c1_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0)) )),
  file('SYN454-1.tptp',unknown)).
fof(f800,plain,(
  ~spl0_2 | spl0_122 | spl0_30),
  inference(avatar_split_clause,[],[f133,f326,f798,f186])).
fof(f133,axiom,(
  ( ! [X0] : (hskp1 | c3_1(X0) | c2_1(X0) | c1_1(X0) | ~ndr1_0) )),
  file('SYN454-1.tptp',unknown)).
fof(f796,plain,(
  ~spl0_121 | ~spl0_31),
  inference(avatar_split_clause,[],[f132,f331,f793])).
fof(f132,axiom,(
  ~hskp0 | ~c3_1(a96)),
  file('SYN454-1.tptp',unknown)).
fof(f791,plain,(
  ~spl0_120 | ~spl0_31),
  inference(avatar_split_clause,[],[f131,f331,f788])).
fof(f131,axiom,(
  ~hskp0 | ~c2_1(a96)),
  file('SYN454-1.tptp',unknown)).
fof(f786,plain,(
  ~spl0_119 | ~spl0_31),
  inference(avatar_split_clause,[],[f130,f331,f783])).
fof(f130,axiom,(
  ~hskp0 | ~c1_1(a96)),
  file('SYN454-1.tptp',unknown)).
fof(f781,plain,(
  ~spl0_118 | ~spl0_30),
  inference(avatar_split_clause,[],[f129,f326,f778])).
fof(f129,axiom,(
  ~hskp1 | ~c1_1(a97)),
  file('SYN454-1.tptp',unknown)).
fof(f766,plain,(
  ~spl0_115 | ~spl0_28),
  inference(avatar_split_clause,[],[f126,f316,f763])).
fof(f126,axiom,(
  ~hskp3 | ~c3_1(a99)),
  file('SYN454-1.tptp',unknown)).
fof(f761,plain,(
  ~spl0_114 | ~spl0_28),
  inference(avatar_split_clause,[],[f125,f316,f758])).
fof(f125,axiom,(
  ~hskp3 | ~c1_1(a99)),
  file('SYN454-1.tptp',unknown)).
fof(f746,plain,(
  ~spl0_111 | ~spl0_26),
  inference(avatar_split_clause,[],[f122,f306,f743])).
fof(f122,axiom,(
  ~hskp5 | ~c2_1(a101)),
  file('SYN454-1.tptp',unknown)).
fof(f741,plain,(
  ~spl0_110 | ~spl0_26),
  inference(avatar_split_clause,[],[f121,f306,f738])).
fof(f121,axiom,(
  ~hskp5 | ~c1_1(a101)),
  file('SYN454-1.tptp',unknown)).
fof(f736,plain,(
  ~spl0_109 | ~spl0_25),
  inference(avatar_split_clause,[],[f120,f301,f733])).
fof(f120,axiom,(
  ~hskp6 | ~c3_1(a102)),
  file('SYN454-1.tptp',unknown)).
fof(f731,plain,(
  ~spl0_108 | ~spl0_24),
  inference(avatar_split_clause,[],[f119,f296,f728])).
fof(f119,axiom,(
  ~hskp7 | ~c3_1(a103)),
  file('SYN454-1.tptp',unknown)).
fof(f726,plain,(
  ~spl0_107 | ~spl0_24),
  inference(avatar_split_clause,[],[f118,f296,f723])).
fof(f118,axiom,(
  ~hskp7 | ~c1_1(a103)),
  file('SYN454-1.tptp',unknown)).
fof(f721,plain,(
  ~spl0_106 | ~spl0_23),
  inference(avatar_split_clause,[],[f117,f291,f718])).
fof(f117,axiom,(
  ~hskp8 | ~c3_1(a104)),
  file('SYN454-1.tptp',unknown)).
fof(f716,plain,(
  ~spl0_105 | ~spl0_23),
  inference(avatar_split_clause,[],[f116,f291,f713])).
fof(f116,axiom,(
  ~hskp8 | ~c2_1(a104)),
  file('SYN454-1.tptp',unknown)).
fof(f711,plain,(
  ~spl0_104 | ~spl0_23),
  inference(avatar_split_clause,[],[f115,f291,f708])).
fof(f115,axiom,(
  ~hskp8 | ~c0_1(a104)),
  file('SYN454-1.tptp',unknown)).
fof(f701,plain,(
  ~spl0_102 | ~spl0_21),
  inference(avatar_split_clause,[],[f113,f281,f698])).
fof(f113,axiom,(
  ~hskp10 | ~c2_1(a106)),
  file('SYN454-1.tptp',unknown)).
fof(f696,plain,(
  ~spl0_101 | ~spl0_20),
  inference(avatar_split_clause,[],[f112,f276,f693])).
fof(f112,axiom,(
  ~hskp11 | ~c3_1(a109)),
  file('SYN454-1.tptp',unknown)).
fof(f691,plain,(
  ~spl0_100 | ~spl0_20),
  inference(avatar_split_clause,[],[f111,f276,f688])).
fof(f111,axiom,(
  ~hskp11 | ~c0_1(a109)),
  file('SYN454-1.tptp',unknown)).
fof(f686,plain,(
  ~spl0_99 | ~spl0_19),
  inference(avatar_split_clause,[],[f110,f271,f683])).
fof(f110,axiom,(
  ~hskp12 | ~c0_1(a110)),
  file('SYN454-1.tptp',unknown)).
fof(f681,plain,(
  ~spl0_98 | ~spl0_18),
  inference(avatar_split_clause,[],[f109,f266,f678])).
fof(f109,axiom,(
  ~hskp13 | ~c1_1(a112)),
  file('SYN454-1.tptp',unknown)).
fof(f676,plain,(
  ~spl0_97 | ~spl0_17),
  inference(avatar_split_clause,[],[f108,f261,f673])).
fof(f108,axiom,(
  ~hskp14 | ~c3_1(a115)),
  file('SYN454-1.tptp',unknown)).
fof(f671,plain,(
  ~spl0_96 | ~spl0_17),
  inference(avatar_split_clause,[],[f107,f261,f668])).
fof(f107,axiom,(
  ~hskp14 | ~c2_1(a115)),
  file('SYN454-1.tptp',unknown)).
fof(f666,plain,(
  ~spl0_95 | ~spl0_16),
  inference(avatar_split_clause,[],[f106,f256,f663])).
fof(f106,axiom,(
  ~hskp15 | ~c0_1(a116)),
  file('SYN454-1.tptp',unknown)).
fof(f661,plain,(
  ~spl0_94 | ~spl0_15),
  inference(avatar_split_clause,[],[f105,f251,f658])).
fof(f105,axiom,(
  ~hskp16 | ~c3_1(a118)),
  file('SYN454-1.tptp',unknown)).
fof(f656,plain,(
  ~spl0_93 | ~spl0_15),
  inference(avatar_split_clause,[],[f104,f251,f653])).
fof(f104,axiom,(
  ~hskp16 | ~c2_1(a118)),
  file('SYN454-1.tptp',unknown)).
fof(f651,plain,(
  ~spl0_92 | ~spl0_14),
  inference(avatar_split_clause,[],[f103,f246,f648])).
fof(f103,axiom,(
  ~hskp17 | ~c3_1(a124)),
  file('SYN454-1.tptp',unknown)).
fof(f646,plain,(
  ~spl0_91 | ~spl0_14),
  inference(avatar_split_clause,[],[f102,f246,f643])).
fof(f102,axiom,(
  ~hskp17 | ~c1_1(a124)),
  file('SYN454-1.tptp',unknown)).
fof(f636,plain,(
  ~spl0_89 | ~spl0_13),
  inference(avatar_split_clause,[],[f100,f241,f633])).
fof(f100,axiom,(
  ~hskp18 | ~c3_1(a125)),
  file('SYN454-1.tptp',unknown)).
fof(f631,plain,(
  ~spl0_88 | ~spl0_12),
  inference(avatar_split_clause,[],[f99,f236,f628])).
fof(f99,axiom,(
  ~hskp19 | ~c3_1(a127)),
  file('SYN454-1.tptp',unknown)).
fof(f626,plain,(
  ~spl0_87 | ~spl0_11),
  inference(avatar_split_clause,[],[f98,f231,f623])).
fof(f98,axiom,(
  ~hskp20 | ~c2_1(a133)),
  file('SYN454-1.tptp',unknown)).
fof(f621,plain,(
  ~spl0_86 | ~spl0_10),
  inference(avatar_split_clause,[],[f97,f226,f618])).
fof(f226,plain,(
  spl0_10 <=> hskp21),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f97,axiom,(
  ~hskp21 | ~c2_1(a145)),
  file('SYN454-1.tptp',unknown)).
fof(f576,plain,(
  ~spl0_30 | spl0_77),
  inference(avatar_split_clause,[],[f88,f573,f326])).
fof(f88,axiom,(
  c3_1(a97) | ~hskp1),
  file('SYN454-1.tptp',unknown)).
fof(f571,plain,(
  ~spl0_30 | spl0_76),
  inference(avatar_split_clause,[],[f87,f568,f326])).
fof(f87,axiom,(
  c2_1(a97) | ~hskp1),
  file('SYN454-1.tptp',unknown)).
fof(f561,plain,(
  ~spl0_28 | spl0_74),
  inference(avatar_split_clause,[],[f85,f558,f316])).
fof(f85,axiom,(
  c2_1(a99) | ~hskp3),
  file('SYN454-1.tptp',unknown)).
fof(f551,plain,(
  ~spl0_26 | spl0_72),
  inference(avatar_split_clause,[],[f83,f548,f306])).
fof(f83,axiom,(
  c3_1(a101) | ~hskp5),
  file('SYN454-1.tptp',unknown)).
fof(f546,plain,(
  ~spl0_25 | spl0_71),
  inference(avatar_split_clause,[],[f82,f543,f301])).
fof(f82,axiom,(
  c2_1(a102) | ~hskp6),
  file('SYN454-1.tptp',unknown)).
fof(f541,plain,(
  ~spl0_25 | spl0_70),
  inference(avatar_split_clause,[],[f81,f538,f301])).
fof(f81,axiom,(
  c1_1(a102) | ~hskp6),
  file('SYN454-1.tptp',unknown)).
fof(f536,plain,(
  ~spl0_24 | spl0_69),
  inference(avatar_split_clause,[],[f80,f533,f296])).
fof(f80,axiom,(
  c0_1(a103) | ~hskp7),
  file('SYN454-1.tptp',unknown)).
fof(f521,plain,(
  ~spl0_21 | spl0_66),
  inference(avatar_split_clause,[],[f77,f518,f281])).
fof(f77,axiom,(
  c3_1(a106) | ~hskp10),
  file('SYN454-1.tptp',unknown)).
fof(f516,plain,(
  ~spl0_21 | spl0_65),
  inference(avatar_split_clause,[],[f76,f513,f281])).
fof(f76,axiom,(
  c1_1(a106) | ~hskp10),
  file('SYN454-1.tptp',unknown)).
fof(f511,plain,(
  ~spl0_20 | spl0_64),
  inference(avatar_split_clause,[],[f75,f508,f276])).
fof(f75,axiom,(
  c2_1(a109) | ~hskp11),
  file('SYN454-1.tptp',unknown)).
fof(f506,plain,(
  ~spl0_19 | spl0_63),
  inference(avatar_split_clause,[],[f74,f503,f271])).
fof(f74,axiom,(
  c3_1(a110) | ~hskp12),
  file('SYN454-1.tptp',unknown)).
fof(f501,plain,(
  ~spl0_19 | spl0_62),
  inference(avatar_split_clause,[],[f73,f498,f271])).
fof(f73,axiom,(
  c2_1(a110) | ~hskp12),
  file('SYN454-1.tptp',unknown)).
fof(f496,plain,(
  ~spl0_18 | spl0_61),
  inference(avatar_split_clause,[],[f72,f493,f266])).
fof(f72,axiom,(
  c3_1(a112) | ~hskp13),
  file('SYN454-1.tptp',unknown)).
fof(f491,plain,(
  ~spl0_18 | spl0_60),
  inference(avatar_split_clause,[],[f71,f488,f266])).
fof(f71,axiom,(
  c0_1(a112) | ~hskp13),
  file('SYN454-1.tptp',unknown)).
fof(f486,plain,(
  ~spl0_17 | spl0_59),
  inference(avatar_split_clause,[],[f70,f483,f261])).
fof(f70,axiom,(
  c0_1(a115) | ~hskp14),
  file('SYN454-1.tptp',unknown)).
fof(f481,plain,(
  ~spl0_16 | spl0_58),
  inference(avatar_split_clause,[],[f69,f478,f256])).
fof(f69,axiom,(
  c3_1(a116) | ~hskp15),
  file('SYN454-1.tptp',unknown)).
fof(f476,plain,(
  ~spl0_16 | spl0_57),
  inference(avatar_split_clause,[],[f68,f473,f256])).
fof(f68,axiom,(
  c1_1(a116) | ~hskp15),
  file('SYN454-1.tptp',unknown)).
fof(f471,plain,(
  ~spl0_15 | spl0_56),
  inference(avatar_split_clause,[],[f67,f468,f251])).
fof(f67,axiom,(
  c1_1(a118) | ~hskp16),
  file('SYN454-1.tptp',unknown)).
fof(f466,plain,(
  ~spl0_13 | spl0_55),
  inference(avatar_split_clause,[],[f66,f463,f241])).
fof(f66,axiom,(
  c1_1(a125) | ~hskp18),
  file('SYN454-1.tptp',unknown)).
fof(f461,plain,(
  ~spl0_13 | spl0_54),
  inference(avatar_split_clause,[],[f65,f458,f241])).
fof(f65,axiom,(
  c0_1(a125) | ~hskp18),
  file('SYN454-1.tptp',unknown)).
fof(f456,plain,(
  ~spl0_12 | spl0_53),
  inference(avatar_split_clause,[],[f64,f453,f236])).
fof(f64,axiom,(
  c2_1(a127) | ~hskp19),
  file('SYN454-1.tptp',unknown)).
fof(f451,plain,(
  ~spl0_12 | spl0_52),
  inference(avatar_split_clause,[],[f63,f448,f236])).
fof(f63,axiom,(
  c0_1(a127) | ~hskp19),
  file('SYN454-1.tptp',unknown)).
fof(f446,plain,(
  ~spl0_11 | spl0_51),
  inference(avatar_split_clause,[],[f62,f443,f231])).
fof(f62,axiom,(
  c3_1(a133) | ~hskp20),
  file('SYN454-1.tptp',unknown)).
fof(f441,plain,(
  ~spl0_11 | spl0_50),
  inference(avatar_split_clause,[],[f61,f438,f231])).
fof(f61,axiom,(
  c0_1(a133) | ~hskp20),
  file('SYN454-1.tptp',unknown)).
fof(f406,plain,(
  ~spl0_5 | spl0_43),
  inference(avatar_split_clause,[],[f54,f403,f201])).
fof(f54,axiom,(
  c2_1(a94) | ~hskp26),
  file('SYN454-1.tptp',unknown)).
fof(f401,plain,(
  ~spl0_5 | spl0_42),
  inference(avatar_split_clause,[],[f53,f398,f201])).
fof(f53,axiom,(
  c1_1(a94) | ~hskp26),
  file('SYN454-1.tptp',unknown)).
fof(f396,plain,(
  ~spl0_5 | spl0_41),
  inference(avatar_split_clause,[],[f52,f393,f201])).
fof(f52,axiom,(
  c0_1(a94) | ~hskp26),
  file('SYN454-1.tptp',unknown)).
fof(f361,plain,(
  ~spl0_1 | spl0_34),
  inference(avatar_split_clause,[],[f45,f358,f182])).
fof(f182,plain,(
  spl0_1 <=> hskp29),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f45,axiom,(
  c3_1(a165) | ~hskp29),
  file('SYN454-1.tptp',unknown)).
fof(f356,plain,(
  ~spl0_1 | spl0_33),
  inference(avatar_split_clause,[],[f44,f353,f182])).
fof(f44,axiom,(
  c2_1(a165) | ~hskp29),
  file('SYN454-1.tptp',unknown)).
fof(f351,plain,(
  ~spl0_1 | spl0_32),
  inference(avatar_split_clause,[],[f43,f348,f182])).
fof(f43,axiom,(
  c1_1(a165) | ~hskp29),
  file('SYN454-1.tptp',unknown)).
fof(f343,plain,(
  spl0_17 | spl0_30 | spl0_28),
  inference(avatar_split_clause,[],[f39,f316,f326,f261])).
fof(f39,axiom,(
  hskp3 | hskp1 | hskp14),
  file('SYN454-1.tptp',unknown)).
fof(f342,plain,(
  spl0_17 | spl0_19 | spl0_20),
  inference(avatar_split_clause,[],[f38,f276,f271,f261])).
fof(f38,axiom,(
  hskp11 | hskp12 | hskp14),
  file('SYN454-1.tptp',unknown)).
fof(f341,plain,(
  spl0_17 | spl0_15 | spl0_10),
  inference(avatar_split_clause,[],[f37,f226,f251,f261])).
fof(f37,axiom,(
  hskp21 | hskp16 | hskp14),
  file('SYN454-1.tptp',unknown)).
fof(f339,plain,(
  spl0_11 | spl0_1 | spl0_31),
  inference(avatar_split_clause,[],[f35,f331,f182,f231])).
fof(f35,axiom,(
  hskp0 | hskp29 | hskp20),
  file('SYN454-1.tptp',unknown)).
fof(f336,plain,(
  spl0_13 | spl0_12 | spl0_18),
  inference(avatar_split_clause,[],[f32,f266,f236,f241])).
fof(f32,axiom,(
  hskp13 | hskp19 | hskp18),
  file('SYN454-1.tptp',unknown)).
fof(f329,plain,(
  ~spl0_30 | spl0_2),
  inference(avatar_split_clause,[],[f29,f186,f326])).
fof(f29,axiom,(
  ndr1_0 | ~hskp1),
  file('SYN454-1.tptp',unknown)).
fof(f319,plain,(
  ~spl0_28 | spl0_2),
  inference(avatar_split_clause,[],[f27,f186,f316])).
fof(f27,axiom,(
  ndr1_0 | ~hskp3),
  file('SYN454-1.tptp',unknown)).
fof(f264,plain,(
  ~spl0_17 | spl0_2),
  inference(avatar_split_clause,[],[f16,f186,f261])).
fof(f16,axiom,(
  ndr1_0 | ~hskp14),
  file('SYN454-1.tptp',unknown)).
% SZS output end Proof for SYN454-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7036
% Time elapsed: 0.033 s
% ------------------------------
% ------------------------------
