% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET206-6
% SZS output start Proof for SET206-6
fof(f251890,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1372,f251803])).
fof(f251803,plain,(
  ~spl0_107),
  inference(avatar_contradiction_clause,[],[f251802])).
fof(f251802,plain,(
  $false | ~spl0_107),
  inference(trivial_inequality_removal,[],[f251691])).
fof(f251691,plain,(
  null_class != null_class | ~spl0_107),
  inference(superposition,[],[f113,f250590])).
fof(f250590,plain,(
  ( ! [X118] : (null_class = cross_product(null_class,X118)) ) | ~spl0_107),
  inference(resolution,[],[f23685,f1375])).
fof(f1375,plain,(
  ( ! [X22] : (~member(X22,null_class)) ) | ~spl0_107),
  inference(backward_demodulation,[],[f1370,f1300])).
fof(f1300,plain,(
  ( ! [X22] : (~member(X22,domain_of(null_class))) )),
  inference(trivial_inequality_removal,[],[f1294])).
fof(f1294,plain,(
  ( ! [X22] : (null_class != null_class | ~member(X22,domain_of(null_class))) )),
  inference(superposition,[],[f134,f1262])).
fof(f1262,plain,(
  ( ! [X0] : (null_class = intersection(null_class,X0)) )),
  inference(resolution,[],[f1183,f1212])).
fof(f1212,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X0)) )),
  inference(duplicate_literal_removal,[],[f1195])).
fof(f1195,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X0) | subclass(intersection(X0,X1),X0)) )),
  inference(resolution,[],[f187,f3])).
fof(f3,axiom,(
  ( ! [X0,X1] : (~member(not_subclass_element(X0,X1),X1) | subclass(X0,X1)) )),
  file('SET206-6.tptp',unknown)).
fof(f187,plain,(
  ( ! [X6,X4,X5] : (member(not_subclass_element(intersection(X4,X5),X6),X4) | subclass(intersection(X4,X5),X6)) )),
  inference(resolution,[],[f21,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (member(not_subclass_element(X0,X1),X0) | subclass(X0,X1)) )),
  file('SET206-6.tptp',unknown)).
fof(f21,axiom,(
  ( ! [X4,X0,X1] : (~member(X4,intersection(X0,X1)) | member(X4,X0)) )),
  file('SET206-6.tptp',unknown)).
fof(f1183,plain,(
  ( ! [X1] : (~subclass(X1,null_class) | null_class = X1) )),
  inference(resolution,[],[f1176,f7])).
fof(f7,axiom,(
  ( ! [X0,X1] : (~subclass(X1,X0) | X0 = X1 | ~subclass(X0,X1)) )),
  file('SET206-6.tptp',unknown)).
fof(f1176,plain,(
  ( ! [X0] : (subclass(null_class,X0)) )),
  inference(forward_demodulation,[],[f1173,f963])).
fof(f963,plain,(
  null_class = complement(universal_class)),
  inference(resolution,[],[f953,f4])).
fof(f4,axiom,(
  ( ! [X0] : (subclass(X0,universal_class)) )),
  file('SET206-6.tptp',unknown)).
fof(f953,plain,(
  ( ! [X0] : (~subclass(complement(X0),X0) | complement(X0) = null_class) )),
  inference(duplicate_literal_removal,[],[f942])).
fof(f942,plain,(
  ( ! [X0] : (~subclass(complement(X0),X0) | complement(X0) = null_class | complement(X0) = null_class) )),
  inference(resolution,[],[f194,f181])).
fof(f181,plain,(
  ( ! [X0] : (~member(regular(complement(X0)),X0) | complement(X0) = null_class) )),
  inference(resolution,[],[f66,f24])).
fof(f24,axiom,(
  ( ! [X4,X0] : (~member(X4,complement(X0)) | ~member(X4,X0)) )),
  file('SET206-6.tptp',unknown)).
fof(f66,axiom,(
  ( ! [X0] : (member(regular(X0),X0) | null_class = X0) )),
  file('SET206-6.tptp',unknown)).
fof(f194,plain,(
  ( ! [X6,X7] : (member(regular(X6),X7) | ~subclass(X6,X7) | null_class = X6) )),
  inference(resolution,[],[f1,f66])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (~member(X2,X0) | member(X2,X1) | ~subclass(X0,X1)) )),
  file('SET206-6.tptp',unknown)).
fof(f1173,plain,(
  ( ! [X0] : (subclass(complement(universal_class),X0)) )),
  inference(resolution,[],[f1167,f4])).
fof(f1167,plain,(
  ( ! [X2,X3] : (~subclass(complement(X2),X2) | subclass(complement(X2),X3)) )),
  inference(duplicate_literal_removal,[],[f1156])).
fof(f1156,plain,(
  ( ! [X2,X3] : (~subclass(complement(X2),X2) | subclass(complement(X2),X3) | subclass(complement(X2),X3)) )),
  inference(resolution,[],[f195,f182])).
fof(f182,plain,(
  ( ! [X0,X1] : (~member(not_subclass_element(complement(X0),X1),X0) | subclass(complement(X0),X1)) )),
  inference(resolution,[],[f2,f24])).
fof(f195,plain,(
  ( ! [X10,X8,X9] : (member(not_subclass_element(X8,X9),X10) | ~subclass(X8,X10) | subclass(X8,X9)) )),
  inference(resolution,[],[f1,f2])).
fof(f134,plain,(
  ( ! [X4,X0] : (null_class != intersection(X0,cross_product(unordered_pair(X4,X4),universal_class)) | ~member(X4,domain_of(X0))) )),
  inference(definition_unfolding,[],[f30,f28,f12])).
fof(f12,axiom,(
  ( ! [X0] : (unordered_pair(X0,X0) = singleton(X0)) )),
  file('SET206-6.tptp',unknown)).
fof(f28,axiom,(
  ( ! [X0,X5,X1] : (intersection(X5,cross_product(X0,X1)) = restrict(X5,X0,X1)) )),
  file('SET206-6.tptp',unknown)).
fof(f30,axiom,(
  ( ! [X4,X0] : (~member(X4,domain_of(X0)) | restrict(X0,singleton(X4),universal_class) != null_class) )),
  file('SET206-6.tptp',unknown)).
fof(f1370,plain,(
  null_class = domain_of(null_class) | ~spl0_107),
  inference(avatar_component_clause,[],[f1368])).
fof(f1368,plain,(
  spl0_107 <=> null_class = domain_of(null_class)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_107])])).
fof(f23685,plain,(
  ( ! [X2,X3] : (member(first(regular(cross_product(X2,X3))),X2) | null_class = cross_product(X2,X3)) )),
  inference(duplicate_literal_removal,[],[f23682])).
fof(f23682,plain,(
  ( ! [X2,X3] : (member(first(regular(cross_product(X2,X3))),X2) | null_class = cross_product(X2,X3) | null_class = cross_product(X2,X3)) )),
  inference(resolution,[],[f4186,f66])).
fof(f4186,plain,(
  ( ! [X2,X0,X3,X1] : (~member(regular(cross_product(X0,X1)),cross_product(X2,X3)) | member(first(regular(cross_product(X0,X1))),X2) | cross_product(X0,X1) = null_class) )),
  inference(superposition,[],[f127,f673])).
fof(f673,plain,(
  ( ! [X6,X7] : (regular(cross_product(X6,X7)) = unordered_pair(unordered_pair(first(regular(cross_product(X6,X7))),first(regular(cross_product(X6,X7)))),unordered_pair(first(regular(cross_product(X6,X7))),unordered_pair(second(regular(cross_product(X6,X7))),second(regular(cross_product(X6,X7)))))) | null_class = cross_product(X6,X7)) )),
  inference(resolution,[],[f130,f66])).
fof(f130,plain,(
  ( ! [X4,X0,X1] : (~member(X4,cross_product(X0,X1)) | unordered_pair(unordered_pair(first(X4),first(X4)),unordered_pair(first(X4),unordered_pair(second(X4),second(X4)))) = X4) )),
  inference(definition_unfolding,[],[f17,f118])).
fof(f118,plain,(
  ( ! [X0,X1] : (ordered_pair(X0,X1) = unordered_pair(unordered_pair(X0,X0),unordered_pair(X0,unordered_pair(X1,X1)))) )),
  inference(definition_unfolding,[],[f13,f12,f12])).
fof(f13,axiom,(
  ( ! [X0,X1] : (unordered_pair(singleton(X0),unordered_pair(X0,singleton(X1))) = ordered_pair(X0,X1)) )),
  file('SET206-6.tptp',unknown)).
fof(f17,axiom,(
  ( ! [X4,X0,X1] : (ordered_pair(first(X4),second(X4)) = X4 | ~member(X4,cross_product(X0,X1))) )),
  file('SET206-6.tptp',unknown)).
fof(f127,plain,(
  ( ! [X2,X0,X3,X1] : (~member(unordered_pair(unordered_pair(X2,X2),unordered_pair(X2,unordered_pair(X3,X3))),cross_product(X0,X1)) | member(X2,X0)) )),
  inference(definition_unfolding,[],[f14,f118])).
fof(f14,axiom,(
  ( ! [X2,X0,X3,X1] : (member(X2,X0) | ~member(ordered_pair(X2,X3),cross_product(X0,X1))) )),
  file('SET206-6.tptp',unknown)).
fof(f113,axiom,(
  null_class != cross_product(null_class,y)),
  file('SET206-6.tptp',unknown)).
fof(f1372,plain,(
  spl0_107),
  inference(avatar_split_clause,[],[f1356,f1368])).
fof(f1356,plain,(
  null_class = domain_of(null_class)),
  inference(resolution,[],[f1300,f66])).
% SZS output end Proof for SET206-6
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 206350
% Time elapsed: 14.452 s
% ------------------------------
% ------------------------------
