% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWV302-2
% SZS output start Proof for SWV302-2
fof(f42,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f17,f29,f30,f31,f39,f41])).
fof(f41,plain,(
  ~spl0_1 | ~spl0_4),
  inference(avatar_contradiction_clause,[],[f40])).
fof(f40,plain,(
  $false | (~spl0_1 | ~spl0_4)),
  inference(subsumption_resolution,[],[f12,f24])).
fof(f24,plain,(
  ( ! [X0] : (~c_in(c_Event_Oevent_OSays(v_B,c_Message_Oagent_OServer,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_B),c_Message_Omsg_OMPair(X0,c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OAgent(v_B)))))))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent)) ) | ~spl0_4),
  inference(avatar_component_clause,[],[f23])).
fof(f23,plain,(
  spl0_4 <=> ! [X0] : ~c_in(c_Event_Oevent_OSays(v_B,c_Message_Oagent_OServer,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_B),c_Message_Omsg_OMPair(X0,c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OAgent(v_B)))))))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f12,plain,(
  c_in(c_Event_Oevent_OSays(v_B,c_Message_Oagent_OServer,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_B),c_Message_Omsg_OMPair(v_x,c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OAgent(v_B)))))))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent) | ~spl0_1),
  inference(avatar_component_clause,[],[f10])).
fof(f10,plain,(
  spl0_1 <=> c_in(c_Event_Oevent_OSays(v_B,c_Message_Oagent_OServer,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_B),c_Message_Omsg_OMPair(v_x,c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OAgent(v_B)))))))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f39,plain,(
  ~spl0_2 | ~spl0_5),
  inference(avatar_contradiction_clause,[],[f38])).
fof(f38,plain,(
  $false | (~spl0_2 | ~spl0_5)),
  inference(subsumption_resolution,[],[f37,f32])).
fof(f32,plain,(
  ~c_in(v_NB,c_Event_Oused(v_evs2),tc_Message_Omsg) | ~spl0_2),
  inference(backward_demodulation,[],[f16,f3])).
fof(f3,axiom,(
  ~c_in(c_Message_Omsg_ONonce(v_NBa),c_Event_Oused(v_evs2),tc_Message_Omsg)),
  file('SWV302-2.tptp',unknown)).
fof(f16,plain,(
  v_NB = c_Message_Omsg_ONonce(v_NBa) | ~spl0_2),
  inference(avatar_component_clause,[],[f14])).
fof(f14,plain,(
  spl0_2 <=> v_NB = c_Message_Omsg_ONonce(v_NBa)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f37,plain,(
  c_in(v_NB,c_Event_Oused(v_evs2),tc_Message_Omsg) | ~spl0_5),
  inference(resolution,[],[f35,f6])).
fof(f6,axiom,(
  ( ! [X2,X1] : (~c_in(X1,c_Message_Oparts(c_Event_Oknows(c_Message_Oagent_OSpy,X2)),tc_Message_Omsg) | c_in(X1,c_Event_Oused(X2),tc_Message_Omsg)) )),
  file('SWV302-2.tptp',unknown)).
fof(f35,plain,(
  c_in(v_NB,c_Message_Oparts(c_Event_Oknows(c_Message_Oagent_OSpy,v_evs2)),tc_Message_Omsg) | ~spl0_5),
  inference(resolution,[],[f33,f7])).
fof(f7,axiom,(
  ( ! [X4,X5,X3] : (~c_in(c_Message_Omsg_OMPair(X3,X4),c_Message_Oparts(X5),tc_Message_Omsg) | c_in(X3,c_Message_Oparts(X5),tc_Message_Omsg)) )),
  file('SWV302-2.tptp',unknown)).
fof(f33,plain,(
  c_in(c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OKey(v_K)),c_Message_Oparts(c_Event_Oknows(c_Message_Oagent_OSpy,v_evs2)),tc_Message_Omsg) | ~spl0_5),
  inference(resolution,[],[f27,f8])).
fof(f8,axiom,(
  ( ! [X6,X5,X3] : (~c_in(c_Message_Omsg_OCrypt(X6,X3),c_Message_Oparts(X5),tc_Message_Omsg) | c_in(X3,c_Message_Oparts(X5),tc_Message_Omsg)) )),
  file('SWV302-2.tptp',unknown)).
fof(f27,plain,(
  c_in(c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OKey(v_K))),c_Message_Oparts(c_Event_Oknows(c_Message_Oagent_OSpy,v_evs2)),tc_Message_Omsg) | ~spl0_5),
  inference(avatar_component_clause,[],[f26])).
fof(f26,plain,(
  spl0_5 <=> c_in(c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OKey(v_K))),c_Message_Oparts(c_Event_Oknows(c_Message_Oagent_OSpy,v_evs2)),tc_Message_Omsg)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f31,plain,(
  ~spl0_3),
  inference(avatar_split_clause,[],[f5,f19])).
fof(f19,plain,(
  spl0_3 <=> c_in(c_Event_Oevent_OSays(c_Message_Oagent_OServer,v_B,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OCrypt(c_Public_OshrK(v_A),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OKey(v_K))),c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OKey(v_K)))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f5,axiom,(
  ~c_in(c_Event_Oevent_OSays(c_Message_Oagent_OServer,v_B,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OCrypt(c_Public_OshrK(v_A),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OKey(v_K))),c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OKey(v_K)))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent)),
  file('SWV302-2.tptp',unknown)).
fof(f30,plain,(
  spl0_5),
  inference(avatar_split_clause,[],[f4,f26])).
fof(f4,axiom,(
  c_in(c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OKey(v_K))),c_Message_Oparts(c_Event_Oknows(c_Message_Oagent_OSpy,v_evs2)),tc_Message_Omsg)),
  file('SWV302-2.tptp',unknown)).
fof(f29,plain,(
  spl0_3 | spl0_4 | ~spl0_5),
  inference(avatar_split_clause,[],[f2,f26,f23,f19])).
fof(f2,axiom,(
  ( ! [X0] : (~c_in(c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OKey(v_K))),c_Message_Oparts(c_Event_Oknows(c_Message_Oagent_OSpy,v_evs2)),tc_Message_Omsg) | ~c_in(c_Event_Oevent_OSays(v_B,c_Message_Oagent_OServer,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_B),c_Message_Omsg_OMPair(X0,c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OAgent(v_B)))))))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent) | c_in(c_Event_Oevent_OSays(c_Message_Oagent_OServer,v_B,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OCrypt(c_Public_OshrK(v_A),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OKey(v_K))),c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OKey(v_K)))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent)) )),
  file('SWV302-2.tptp',unknown)).
fof(f17,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f1,f14,f10])).
fof(f1,axiom,(
  v_NB = c_Message_Omsg_ONonce(v_NBa) | c_in(c_Event_Oevent_OSays(v_B,c_Message_Oagent_OServer,c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_B),c_Message_Omsg_OMPair(v_x,c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(v_NA,c_Message_Omsg_OMPair(v_NB,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OAgent(v_B)))))))))),c_List_Oset(v_evs2,tc_Event_Oevent),tc_Event_Oevent)),
  file('SWV302-2.tptp',unknown)).
% SZS output end Proof for SWV302-2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
