% Refutation found. Thanks to Tanya!
% SZS status Theorem for SYN068+1
% SZS output start Proof for SYN068+1
fof(f25,plain,(
  $false),
  inference(subsumption_resolution,[],[f24,f21])).
fof(f21,plain,(
  big_f(sK2)),
  inference(resolution,[],[f20,f18])).
fof(f18,plain,(
  big_j(sK2)),
  inference(cnf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X1] : (big_h(sK2,X1) | ~big_g(X1)) & big_j(sK2)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f7,f12])).
fof(f12,plain,(
  ? [X0] : (! [X1] : (big_h(X0,X1) | ~big_g(X1)) & big_j(X0)) => (! [X1] : (big_h(sK2,X1) | ~big_g(X1)) & big_j(sK2))),
  introduced(choice_axiom,[])).
fof(f7,plain,(
  ? [X0] : (! [X1] : (big_h(X0,X1) | ~big_g(X1)) & big_j(X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ? [X0] : (! [X1] : (big_g(X1) => big_h(X0,X1)) & big_j(X0))),
  file('SYN068+1.tptp',unknown)).
fof(f20,plain,(
  ( ! [X0] : (~big_j(X0) | big_f(X0)) )),
  inference(cnf_transformation,[],[f8])).
fof(f8,plain,(
  ! [X0] : (big_f(X0) | ~big_j(X0))),
  inference(ennf_transformation,[],[f4])).
fof(f4,negated_conjecture,(
  ~? [X0] : (~big_f(X0) & big_j(X0))),
  inference(negated_conjecture,[],[f3])).
fof(f3,conjecture,(
  ? [X0] : (~big_f(X0) & big_j(X0))),
  file('SYN068+1.tptp',unknown)).
fof(f24,plain,(
  ~big_f(sK2)),
  inference(resolution,[],[f23,f16])).
fof(f16,plain,(
  ( ! [X0] : (big_g(sK0(X0)) | ~big_f(X0)) )),
  inference(cnf_transformation,[],[f11])).
fof(f11,plain,(
  ! [X0] : (((~big_h(X0,sK0(X0)) & big_g(sK0(X0))) & (big_h(X0,sK1(X0)) & big_g(sK1(X0)))) | ~big_f(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f6,f10,f9])).
fof(f9,plain,(
  ! [X0] : (? [X1] : (~big_h(X0,X1) & big_g(X1)) => (~big_h(X0,sK0(X0)) & big_g(sK0(X0))))),
  introduced(choice_axiom,[])).
fof(f10,plain,(
  ! [X0] : (? [X2] : (big_h(X0,X2) & big_g(X2)) => (big_h(X0,sK1(X0)) & big_g(sK1(X0))))),
  introduced(choice_axiom,[])).
fof(f6,plain,(
  ! [X0] : ((? [X1] : (~big_h(X0,X1) & big_g(X1)) & ? [X2] : (big_h(X0,X2) & big_g(X2))) | ~big_f(X0))),
  inference(ennf_transformation,[],[f5])).
fof(f5,plain,(
  ! [X0] : (big_f(X0) => (? [X1] : (~big_h(X0,X1) & big_g(X1)) & ? [X2] : (big_h(X0,X2) & big_g(X2))))),
  inference(rectify,[],[f1])).
fof(f1,axiom,(
  ! [X0] : (big_f(X0) => (? [X2] : (~big_h(X0,X2) & big_g(X2)) & ? [X1] : (big_h(X0,X1) & big_g(X1))))),
  file('SYN068+1.tptp',unknown)).
fof(f23,plain,(
  ~big_g(sK0(sK2))),
  inference(subsumption_resolution,[],[f22,f21])).
fof(f22,plain,(
  ~big_f(sK2) | ~big_g(sK0(sK2))),
  inference(resolution,[],[f17,f19])).
fof(f19,plain,(
  ( ! [X1] : (big_h(sK2,X1) | ~big_g(X1)) )),
  inference(cnf_transformation,[],[f13])).
fof(f17,plain,(
  ( ! [X0] : (~big_h(X0,sK0(X0)) | ~big_f(X0)) )),
  inference(cnf_transformation,[],[f11])).
% SZS output end Proof for SYN068+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
