% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SYN510-1
% SZS output start Proof for SYN510-1
fof(f3857,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f312,f337,f352,f374,f379,f380,f381,f383,f384,f385,f386,f387,f392,f397,f402,f412,f427,f432,f457,f462,f467,f477,f482,f487,f492,f497,f502,f512,f517,f522,f527,f532,f537,f542,f547,f552,f557,f562,f567,f572,f587,f592,f607,f612,f617,f622,f627,f632,f637,f652,f657,f662,f667,f672,f677,f682,f687,f697,f702,f707,f712,f717,f722,f737,f742,f747,f752,f757,f762,f767,f777,f782,f787,f797,f802,f807,f812,f822,f827,f832,f837,f842,f847,f852,f857,f862,f867,f880,f892,f900,f912,f916,f920,f925,f929,f933,f937,f941,f945,f947,f948,f957,f961,f965,f966,f970,f974,f975,f979,f980,f984,f985,f994,f995,f996,f997,f998,f1003,f1004,f1005,f1006,f1010,f1011,f1012,f1013,f1014,f1015,f1017,f1018,f1019,f1023,f1024,f1025,f1026,f1027,f1028,f1046,f1080,f1082,f1092,f1097,f1132,f1134,f1136,f1157,f1159,f1184,f1189,f1256,f1260,f1262,f1291,f1437,f1439,f1655,f1656,f1743,f1744,f1791,f1799,f1814,f1818,f1820,f1844,f1869,f1883,f1903,f1936,f1947,f1955,f2017,f2099,f2187,f2193,f2245,f2282,f2374,f2403,f2422,f2425,f2470,f2473,f2504,f2535,f2545,f2597,f2619,f2638,f2648,f2671,f2696,f2714,f2718,f2720,f2797,f2804,f2858,f2860,f2865,f2886,f2889,f2927,f2957,f2980,f3008,f3031,f3032,f3101,f3144,f3146,f3184,f3208,f3267,f3270,f3293,f3301,f3323,f3454,f3491,f3544,f3550,f3552,f3556,f3640,f3644,f3666,f3668,f3694,f3712,f3770,f3852])).
fof(f3852,plain,(
  spl0_133 | ~spl0_132 | ~spl0_138 | ~spl0_145),
  inference(avatar_split_clause,[],[f3839,f931,f902,f877,f882])).
fof(f882,plain,(
  spl0_133 <=> ! [X0] : (c1_1(X0) | ~c3_1(X0) | c0_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_133])])).
fof(f877,plain,(
  spl0_132 <=> ! [X0] : (c3_1(X0) | c1_1(X0) | c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_132])])).
fof(f902,plain,(
  spl0_138 <=> ! [X0] : (c2_1(X0) | ~c3_1(X0) | c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_138])])).
fof(f931,plain,(
  spl0_145 <=> ! [X0] : (c0_1(X0) | ~c2_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_145])])).
fof(f3839,plain,(
  ( ! [X1] : (c0_1(X1) | ~c3_1(X1) | c1_1(X1)) ) | (~spl0_132 | ~spl0_138 | ~spl0_145)),
  inference(resolution,[],[f932,f3674])).
fof(f3674,plain,(
  ( ! [X0] : (c2_1(X0) | c1_1(X0)) ) | (~spl0_132 | ~spl0_138)),
  inference(subsumption_resolution,[],[f878,f903])).
fof(f903,plain,(
  ( ! [X0] : (~c3_1(X0) | c2_1(X0) | c1_1(X0)) ) | ~spl0_138),
  inference(avatar_component_clause,[],[f902])).
fof(f878,plain,(
  ( ! [X0] : (c3_1(X0) | c1_1(X0) | c2_1(X0)) ) | ~spl0_132),
  inference(avatar_component_clause,[],[f877])).
fof(f932,plain,(
  ( ! [X0] : (~c2_1(X0) | c0_1(X0) | ~c3_1(X0)) ) | ~spl0_145),
  inference(avatar_component_clause,[],[f931])).
fof(f3770,plain,(
  spl0_126 | spl0_125 | ~spl0_132 | ~spl0_138 | ~spl0_142),
  inference(avatar_split_clause,[],[f3764,f918,f902,f877,f844,f849])).
fof(f849,plain,(
  spl0_126 <=> c3_1(a538)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_126])])).
fof(f844,plain,(
  spl0_125 <=> c2_1(a538)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_125])])).
fof(f918,plain,(
  spl0_142 <=> ! [X0] : (c3_1(X0) | ~c1_1(X0) | c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_142])])).
fof(f3764,plain,(
  c3_1(a538) | (spl0_125 | ~spl0_132 | ~spl0_138 | ~spl0_142)),
  inference(resolution,[],[f3675,f846])).
fof(f846,plain,(
  ~c2_1(a538) | spl0_125),
  inference(avatar_component_clause,[],[f844])).
fof(f3675,plain,(
  ( ! [X0] : (c2_1(X0) | c3_1(X0)) ) | (~spl0_132 | ~spl0_138 | ~spl0_142)),
  inference(subsumption_resolution,[],[f919,f3674])).
fof(f919,plain,(
  ( ! [X0] : (~c1_1(X0) | c3_1(X0) | c2_1(X0)) ) | ~spl0_142),
  inference(avatar_component_clause,[],[f918])).
fof(f3712,plain,(
  ~spl0_63 | spl0_107 | ~spl0_141 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f3711])).
fof(f3711,plain,(
  $false | (~spl0_63 | spl0_107 | ~spl0_141 | ~spl0_146)),
  inference(subsumption_resolution,[],[f3701,f756])).
fof(f756,plain,(
  ~c2_1(a565) | spl0_107),
  inference(avatar_component_clause,[],[f754])).
fof(f754,plain,(
  spl0_107 <=> c2_1(a565)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_107])])).
fof(f3701,plain,(
  c2_1(a565) | (~spl0_63 | ~spl0_141 | ~spl0_146)),
  inference(resolution,[],[f3646,f536])).
fof(f536,plain,(
  c0_1(a565) | ~spl0_63),
  inference(avatar_component_clause,[],[f534])).
fof(f534,plain,(
  spl0_63 <=> c0_1(a565)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_63])])).
fof(f3646,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0)) ) | (~spl0_141 | ~spl0_146)),
  inference(subsumption_resolution,[],[f936,f915])).
fof(f915,plain,(
  ( ! [X0] : (~c0_1(X0) | c3_1(X0) | c2_1(X0)) ) | ~spl0_141),
  inference(avatar_component_clause,[],[f914])).
fof(f914,plain,(
  spl0_141 <=> ! [X0] : (c3_1(X0) | ~c0_1(X0) | c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_141])])).
fof(f936,plain,(
  ( ! [X0] : (c2_1(X0) | ~c0_1(X0) | ~c3_1(X0)) ) | ~spl0_146),
  inference(avatar_component_clause,[],[f935])).
fof(f935,plain,(
  spl0_146 <=> ! [X0] : (c2_1(X0) | ~c0_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_146])])).
fof(f3694,plain,(
  spl0_83 | ~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_148 | ~spl0_159),
  inference(avatar_contradiction_clause,[],[f3677])).
fof(f3677,plain,(
  $false | (spl0_83 | ~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_148 | ~spl0_159)),
  inference(resolution,[],[f3676,f636])).
fof(f636,plain,(
  ~c3_1(a604) | spl0_83),
  inference(avatar_component_clause,[],[f634])).
fof(f634,plain,(
  spl0_83 <=> c3_1(a604)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_83])])).
fof(f3676,plain,(
  ( ! [X0] : (c3_1(X0)) ) | (~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_148 | ~spl0_159)),
  inference(subsumption_resolution,[],[f3675,f3673])).
fof(f3673,plain,(
  ( ! [X0] : (c3_1(X0) | ~c2_1(X0)) ) | (~spl0_148 | ~spl0_159)),
  inference(subsumption_resolution,[],[f944,f1002])).
fof(f1002,plain,(
  ( ! [X1] : (c3_1(X1) | ~c2_1(X1) | c0_1(X1)) ) | ~spl0_159),
  inference(avatar_component_clause,[],[f1001])).
fof(f1001,plain,(
  spl0_159 <=> ! [X1] : (c3_1(X1) | ~c2_1(X1) | c0_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_159])])).
fof(f944,plain,(
  ( ! [X0] : (c3_1(X0) | ~c0_1(X0) | ~c2_1(X0)) ) | ~spl0_148),
  inference(avatar_component_clause,[],[f943])).
fof(f943,plain,(
  spl0_148 <=> ! [X0] : (c3_1(X0) | ~c0_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_148])])).
fof(f3668,plain,(
  ~spl0_55 | spl0_95 | ~spl0_142 | ~spl0_149),
  inference(avatar_contradiction_clause,[],[f3667])).
fof(f3667,plain,(
  $false | (~spl0_55 | spl0_95 | ~spl0_142 | ~spl0_149)),
  inference(subsumption_resolution,[],[f3659,f696])).
fof(f696,plain,(
  ~c3_1(a580) | spl0_95),
  inference(avatar_component_clause,[],[f694])).
fof(f694,plain,(
  spl0_95 <=> c3_1(a580)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_95])])).
fof(f3659,plain,(
  c3_1(a580) | (~spl0_55 | ~spl0_142 | ~spl0_149)),
  inference(resolution,[],[f3560,f496])).
fof(f496,plain,(
  c1_1(a580) | ~spl0_55),
  inference(avatar_component_clause,[],[f494])).
fof(f494,plain,(
  spl0_55 <=> c1_1(a580)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_55])])).
fof(f3560,plain,(
  ( ! [X0] : (~c1_1(X0) | c3_1(X0)) ) | (~spl0_142 | ~spl0_149)),
  inference(subsumption_resolution,[],[f919,f952])).
fof(f952,plain,(
  ( ! [X0] : (c3_1(X0) | ~c1_1(X0) | ~c2_1(X0)) ) | ~spl0_149),
  inference(avatar_component_clause,[],[f951])).
fof(f951,plain,(
  spl0_149 <=> ! [X0] : (c3_1(X0) | ~c1_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_149])])).
fof(f3666,plain,(
  ~spl0_48 | spl0_83 | ~spl0_142 | ~spl0_149),
  inference(avatar_contradiction_clause,[],[f3665])).
fof(f3665,plain,(
  $false | (~spl0_48 | spl0_83 | ~spl0_142 | ~spl0_149)),
  inference(subsumption_resolution,[],[f3657,f636])).
fof(f3657,plain,(
  c3_1(a604) | (~spl0_48 | ~spl0_142 | ~spl0_149)),
  inference(resolution,[],[f3560,f461])).
fof(f461,plain,(
  c1_1(a604) | ~spl0_48),
  inference(avatar_component_clause,[],[f459])).
fof(f459,plain,(
  spl0_48 <=> c1_1(a604)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_48])])).
fof(f3644,plain,(
  ~spl0_38 | ~spl0_145 | ~spl0_150 | ~spl0_156 | ~spl0_159),
  inference(avatar_contradiction_clause,[],[f3619])).
fof(f3619,plain,(
  $false | (~spl0_38 | ~spl0_145 | ~spl0_150 | ~spl0_156 | ~spl0_159)),
  inference(resolution,[],[f3583,f411])).
fof(f411,plain,(
  c2_1(a562) | ~spl0_38),
  inference(avatar_component_clause,[],[f409])).
fof(f409,plain,(
  spl0_38 <=> c2_1(a562)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_38])])).
fof(f3583,plain,(
  ( ! [X0] : (~c2_1(X0)) ) | (~spl0_145 | ~spl0_150 | ~spl0_156 | ~spl0_159)),
  inference(subsumption_resolution,[],[f3582,f3559])).
fof(f3559,plain,(
  ( ! [X0] : (~c2_1(X0) | c0_1(X0)) ) | (~spl0_145 | ~spl0_159)),
  inference(subsumption_resolution,[],[f932,f1002])).
fof(f3582,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c0_1(X0)) ) | (~spl0_150 | ~spl0_156)),
  inference(subsumption_resolution,[],[f956,f983])).
fof(f983,plain,(
  ( ! [X0] : (c1_1(X0) | ~c0_1(X0) | ~c2_1(X0)) ) | ~spl0_156),
  inference(avatar_component_clause,[],[f982])).
fof(f982,plain,(
  spl0_156 <=> ! [X0] : (c1_1(X0) | ~c0_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_156])])).
fof(f956,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c0_1(X0) | ~c1_1(X0)) ) | ~spl0_150),
  inference(avatar_component_clause,[],[f955])).
fof(f955,plain,(
  spl0_150 <=> ! [X0] : (~c2_1(X0) | ~c0_1(X0) | ~c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_150])])).
fof(f3640,plain,(
  ~spl0_49 | ~spl0_145 | ~spl0_150 | ~spl0_156 | ~spl0_159),
  inference(avatar_contradiction_clause,[],[f3623])).
fof(f3623,plain,(
  $false | (~spl0_49 | ~spl0_145 | ~spl0_150 | ~spl0_156 | ~spl0_159)),
  inference(resolution,[],[f3583,f466])).
fof(f466,plain,(
  c2_1(a604) | ~spl0_49),
  inference(avatar_component_clause,[],[f464])).
fof(f464,plain,(
  spl0_49 <=> c2_1(a604)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_49])])).
fof(f3556,plain,(
  ~spl0_79 | spl0_127 | ~spl0_156 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f3555])).
fof(f3555,plain,(
  $false | (~spl0_79 | spl0_127 | ~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f3538,f856])).
fof(f856,plain,(
  ~c1_1(a537) | spl0_127),
  inference(avatar_component_clause,[],[f854])).
fof(f854,plain,(
  spl0_127 <=> c1_1(a537)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_127])])).
fof(f3538,plain,(
  c1_1(a537) | (~spl0_79 | ~spl0_156 | ~spl0_161)),
  inference(resolution,[],[f3512,f616])).
fof(f616,plain,(
  c2_1(a537) | ~spl0_79),
  inference(avatar_component_clause,[],[f614])).
fof(f614,plain,(
  spl0_79 <=> c2_1(a537)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_79])])).
fof(f3512,plain,(
  ( ! [X0] : (~c2_1(X0) | c1_1(X0)) ) | (~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f983,f1022])).
fof(f1022,plain,(
  ( ! [X2] : (c1_1(X2) | ~c2_1(X2) | c0_1(X2)) ) | ~spl0_161),
  inference(avatar_component_clause,[],[f1021])).
fof(f1021,plain,(
  spl0_161 <=> ! [X2] : (c1_1(X2) | ~c2_1(X2) | c0_1(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_161])])).
fof(f3552,plain,(
  ~spl0_61 | ~spl0_62 | ~spl0_147 | ~spl0_153 | ~spl0_156 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f3551])).
fof(f3551,plain,(
  $false | (~spl0_61 | ~spl0_62 | ~spl0_147 | ~spl0_153 | ~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f3534,f3424])).
fof(f3424,plain,(
  ~c1_1(a566) | (~spl0_62 | ~spl0_147 | ~spl0_153)),
  inference(resolution,[],[f3402,f531])).
fof(f531,plain,(
  c3_1(a566) | ~spl0_62),
  inference(avatar_component_clause,[],[f529])).
fof(f529,plain,(
  spl0_62 <=> c3_1(a566)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_62])])).
fof(f3402,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c1_1(X0)) ) | (~spl0_147 | ~spl0_153)),
  inference(subsumption_resolution,[],[f969,f940])).
fof(f940,plain,(
  ( ! [X0] : (c2_1(X0) | ~c1_1(X0) | ~c3_1(X0)) ) | ~spl0_147),
  inference(avatar_component_clause,[],[f939])).
fof(f939,plain,(
  spl0_147 <=> ! [X0] : (c2_1(X0) | ~c1_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_147])])).
fof(f969,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c1_1(X0) | ~c2_1(X0)) ) | ~spl0_153),
  inference(avatar_component_clause,[],[f968])).
fof(f968,plain,(
  spl0_153 <=> ! [X0] : (~c3_1(X0) | ~c1_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_153])])).
fof(f3534,plain,(
  c1_1(a566) | (~spl0_61 | ~spl0_156 | ~spl0_161)),
  inference(resolution,[],[f3512,f526])).
fof(f526,plain,(
  c2_1(a566) | ~spl0_61),
  inference(avatar_component_clause,[],[f524])).
fof(f524,plain,(
  spl0_61 <=> c2_1(a566)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_61])])).
fof(f3550,plain,(
  ~spl0_53 | spl0_93 | ~spl0_156 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f3549])).
fof(f3549,plain,(
  $false | (~spl0_53 | spl0_93 | ~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f3530,f686])).
fof(f686,plain,(
  ~c1_1(a582) | spl0_93),
  inference(avatar_component_clause,[],[f684])).
fof(f684,plain,(
  spl0_93 <=> c1_1(a582)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_93])])).
fof(f3530,plain,(
  c1_1(a582) | (~spl0_53 | ~spl0_156 | ~spl0_161)),
  inference(resolution,[],[f3512,f486])).
fof(f486,plain,(
  c2_1(a582) | ~spl0_53),
  inference(avatar_component_clause,[],[f484])).
fof(f484,plain,(
  spl0_53 <=> c2_1(a582)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_53])])).
fof(f3544,plain,(
  ~spl0_47 | spl0_82 | ~spl0_156 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f3543])).
fof(f3543,plain,(
  $false | (~spl0_47 | spl0_82 | ~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f3526,f631])).
fof(f631,plain,(
  ~c1_1(a615) | spl0_82),
  inference(avatar_component_clause,[],[f629])).
fof(f629,plain,(
  spl0_82 <=> c1_1(a615)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_82])])).
fof(f3526,plain,(
  c1_1(a615) | (~spl0_47 | ~spl0_156 | ~spl0_161)),
  inference(resolution,[],[f3512,f456])).
fof(f456,plain,(
  c2_1(a615) | ~spl0_47),
  inference(avatar_component_clause,[],[f454])).
fof(f454,plain,(
  spl0_47 <=> c2_1(a615)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_47])])).
fof(f3491,plain,(
  ~spl0_73 | spl0_115 | ~spl0_143 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f3490])).
fof(f3490,plain,(
  $false | (~spl0_73 | spl0_115 | ~spl0_143 | ~spl0_161)),
  inference(subsumption_resolution,[],[f3472,f796])).
fof(f796,plain,(
  ~c0_1(a547) | spl0_115),
  inference(avatar_component_clause,[],[f794])).
fof(f794,plain,(
  spl0_115 <=> c0_1(a547)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_115])])).
fof(f3472,plain,(
  c0_1(a547) | (~spl0_73 | ~spl0_143 | ~spl0_161)),
  inference(resolution,[],[f3406,f586])).
fof(f586,plain,(
  c2_1(a547) | ~spl0_73),
  inference(avatar_component_clause,[],[f584])).
fof(f584,plain,(
  spl0_73 <=> c2_1(a547)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_73])])).
fof(f3406,plain,(
  ( ! [X2] : (~c2_1(X2) | c0_1(X2)) ) | (~spl0_143 | ~spl0_161)),
  inference(subsumption_resolution,[],[f1022,f924])).
fof(f924,plain,(
  ( ! [X0] : (c0_1(X0) | ~c1_1(X0) | ~c2_1(X0)) ) | ~spl0_143),
  inference(avatar_component_clause,[],[f923])).
fof(f923,plain,(
  spl0_143 <=> ! [X0] : (c0_1(X0) | ~c1_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_143])])).
fof(f3454,plain,(
  ~spl0_65 | spl0_108 | ~spl0_137 | ~spl0_156),
  inference(avatar_contradiction_clause,[],[f3453])).
fof(f3453,plain,(
  $false | (~spl0_65 | spl0_108 | ~spl0_137 | ~spl0_156)),
  inference(subsumption_resolution,[],[f3446,f761])).
fof(f761,plain,(
  ~c1_1(a564) | spl0_108),
  inference(avatar_component_clause,[],[f759])).
fof(f759,plain,(
  spl0_108 <=> c1_1(a564)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_108])])).
fof(f3446,plain,(
  c1_1(a564) | (~spl0_65 | ~spl0_137 | ~spl0_156)),
  inference(resolution,[],[f3405,f546])).
fof(f546,plain,(
  c0_1(a564) | ~spl0_65),
  inference(avatar_component_clause,[],[f544])).
fof(f544,plain,(
  spl0_65 <=> c0_1(a564)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_65])])).
fof(f3405,plain,(
  ( ! [X0] : (~c0_1(X0) | c1_1(X0)) ) | (~spl0_137 | ~spl0_156)),
  inference(subsumption_resolution,[],[f983,f899])).
fof(f899,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0) | c1_1(X0)) ) | ~spl0_137),
  inference(avatar_component_clause,[],[f898])).
fof(f898,plain,(
  spl0_137 <=> ! [X0] : (c2_1(X0) | ~c0_1(X0) | c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_137])])).
fof(f3323,plain,(
  spl0_113 | ~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_147),
  inference(avatar_contradiction_clause,[],[f3322])).
fof(f3322,plain,(
  $false | (spl0_113 | ~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_147)),
  inference(resolution,[],[f3288,f786])).
fof(f786,plain,(
  ~c2_1(a549) | spl0_113),
  inference(avatar_component_clause,[],[f784])).
fof(f784,plain,(
  spl0_113 <=> c2_1(a549)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_113])])).
fof(f3288,plain,(
  ( ! [X0] : (c2_1(X0)) ) | (~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_147)),
  inference(subsumption_resolution,[],[f3287,f3281])).
fof(f3281,plain,(
  ( ! [X0] : (c2_1(X0) | c1_1(X0)) ) | (~spl0_132 | ~spl0_138 | ~spl0_142)),
  inference(subsumption_resolution,[],[f903,f3151])).
fof(f3151,plain,(
  ( ! [X0] : (c2_1(X0) | c3_1(X0)) ) | (~spl0_132 | ~spl0_142)),
  inference(subsumption_resolution,[],[f878,f919])).
fof(f3287,plain,(
  ( ! [X0] : (c2_1(X0) | ~c1_1(X0)) ) | (~spl0_142 | ~spl0_147)),
  inference(subsumption_resolution,[],[f940,f919])).
fof(f3301,plain,(
  spl0_125 | ~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_147),
  inference(avatar_contradiction_clause,[],[f3300])).
fof(f3300,plain,(
  $false | (spl0_125 | ~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_147)),
  inference(subsumption_resolution,[],[f846,f3288])).
fof(f3293,plain,(
  spl0_89 | ~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_147),
  inference(avatar_contradiction_clause,[],[f3292])).
fof(f3292,plain,(
  $false | (spl0_89 | ~spl0_132 | ~spl0_138 | ~spl0_142 | ~spl0_147)),
  inference(subsumption_resolution,[],[f666,f3288])).
fof(f666,plain,(
  ~c2_1(a587) | spl0_89),
  inference(avatar_component_clause,[],[f664])).
fof(f664,plain,(
  spl0_89 <=> c2_1(a587)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_89])])).
fof(f3270,plain,(
  ~spl0_61 | ~spl0_62 | spl0_106 | ~spl0_145),
  inference(avatar_contradiction_clause,[],[f3269])).
fof(f3269,plain,(
  $false | (~spl0_61 | ~spl0_62 | spl0_106 | ~spl0_145)),
  inference(subsumption_resolution,[],[f3268,f531])).
fof(f3268,plain,(
  ~c3_1(a566) | (~spl0_61 | spl0_106 | ~spl0_145)),
  inference(subsumption_resolution,[],[f3243,f751])).
fof(f751,plain,(
  ~c0_1(a566) | spl0_106),
  inference(avatar_component_clause,[],[f749])).
fof(f749,plain,(
  spl0_106 <=> c0_1(a566)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_106])])).
fof(f3243,plain,(
  c0_1(a566) | ~c3_1(a566) | (~spl0_61 | ~spl0_145)),
  inference(resolution,[],[f932,f526])).
fof(f3267,plain,(
  ~spl0_53 | ~spl0_54 | ~spl0_145 | ~spl0_151 | ~spl0_155),
  inference(avatar_contradiction_clause,[],[f3266])).
fof(f3266,plain,(
  $false | (~spl0_53 | ~spl0_54 | ~spl0_145 | ~spl0_151 | ~spl0_155)),
  inference(subsumption_resolution,[],[f3265,f491])).
fof(f491,plain,(
  c3_1(a582) | ~spl0_54),
  inference(avatar_component_clause,[],[f489])).
fof(f489,plain,(
  spl0_54 <=> c3_1(a582)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_54])])).
fof(f3265,plain,(
  ~c3_1(a582) | (~spl0_53 | ~spl0_54 | ~spl0_145 | ~spl0_151 | ~spl0_155)),
  inference(subsumption_resolution,[],[f3239,f2989])).
fof(f2989,plain,(
  ~c0_1(a582) | (~spl0_54 | ~spl0_151 | ~spl0_155)),
  inference(resolution,[],[f2961,f491])).
fof(f2961,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c0_1(X0)) ) | (~spl0_151 | ~spl0_155)),
  inference(subsumption_resolution,[],[f960,f978])).
fof(f978,plain,(
  ( ! [X0] : (c1_1(X0) | ~c0_1(X0) | ~c3_1(X0)) ) | ~spl0_155),
  inference(avatar_component_clause,[],[f977])).
fof(f977,plain,(
  spl0_155 <=> ! [X0] : (c1_1(X0) | ~c0_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_155])])).
fof(f960,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c0_1(X0) | ~c1_1(X0)) ) | ~spl0_151),
  inference(avatar_component_clause,[],[f959])).
fof(f959,plain,(
  spl0_151 <=> ! [X0] : (~c3_1(X0) | ~c0_1(X0) | ~c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_151])])).
fof(f3239,plain,(
  c0_1(a582) | ~c3_1(a582) | (~spl0_53 | ~spl0_145)),
  inference(resolution,[],[f932,f486])).
fof(f3208,plain,(
  spl0_127 | ~spl0_137 | ~spl0_156 | ~spl0_166),
  inference(avatar_contradiction_clause,[],[f3207])).
fof(f3207,plain,(
  $false | (spl0_127 | ~spl0_137 | ~spl0_156 | ~spl0_166)),
  inference(subsumption_resolution,[],[f3204,f856])).
fof(f3204,plain,(
  c1_1(a537) | (~spl0_137 | ~spl0_156 | ~spl0_166)),
  inference(resolution,[],[f3154,f2368])).
fof(f2368,plain,(
  c0_1(a537) | ~spl0_166),
  inference(avatar_component_clause,[],[f2366])).
fof(f2366,plain,(
  spl0_166 <=> c0_1(a537)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_166])])).
fof(f3154,plain,(
  ( ! [X0] : (~c0_1(X0) | c1_1(X0)) ) | (~spl0_137 | ~spl0_156)),
  inference(subsumption_resolution,[],[f983,f899])).
fof(f3184,plain,(
  spl0_127 | spl0_128 | ~spl0_130 | ~spl0_139),
  inference(avatar_contradiction_clause,[],[f3183])).
fof(f3183,plain,(
  $false | (spl0_127 | spl0_128 | ~spl0_130 | ~spl0_139)),
  inference(subsumption_resolution,[],[f3180,f856])).
fof(f3180,plain,(
  c1_1(a537) | (spl0_128 | ~spl0_130 | ~spl0_139)),
  inference(resolution,[],[f3152,f861])).
fof(f861,plain,(
  ~c3_1(a537) | spl0_128),
  inference(avatar_component_clause,[],[f859])).
fof(f859,plain,(
  spl0_128 <=> c3_1(a537)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_128])])).
fof(f3152,plain,(
  ( ! [X0] : (c3_1(X0) | c1_1(X0)) ) | (~spl0_130 | ~spl0_139)),
  inference(subsumption_resolution,[],[f907,f870])).
fof(f870,plain,(
  ( ! [X0] : (c3_1(X0) | c0_1(X0) | c1_1(X0)) ) | ~spl0_130),
  inference(avatar_component_clause,[],[f869])).
fof(f869,plain,(
  spl0_130 <=> ! [X0] : (c3_1(X0) | c0_1(X0) | c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_130])])).
fof(f907,plain,(
  ( ! [X0] : (~c0_1(X0) | c3_1(X0) | c1_1(X0)) ) | ~spl0_139),
  inference(avatar_component_clause,[],[f906])).
fof(f906,plain,(
  spl0_139 <=> ! [X0] : (c3_1(X0) | ~c0_1(X0) | c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_139])])).
fof(f3146,plain,(
  ~spl0_79 | ~spl0_145 | ~spl0_159 | spl0_166),
  inference(avatar_contradiction_clause,[],[f3145])).
fof(f3145,plain,(
  $false | (~spl0_79 | ~spl0_145 | ~spl0_159 | spl0_166)),
  inference(subsumption_resolution,[],[f3123,f2367])).
fof(f2367,plain,(
  ~c0_1(a537) | spl0_166),
  inference(avatar_component_clause,[],[f2366])).
fof(f3123,plain,(
  c0_1(a537) | (~spl0_79 | ~spl0_145 | ~spl0_159)),
  inference(resolution,[],[f3077,f616])).
fof(f3077,plain,(
  ( ! [X1] : (~c2_1(X1) | c0_1(X1)) ) | (~spl0_145 | ~spl0_159)),
  inference(subsumption_resolution,[],[f1002,f932])).
fof(f3144,plain,(
  ~spl0_73 | spl0_115 | ~spl0_145 | ~spl0_159),
  inference(avatar_contradiction_clause,[],[f3143])).
fof(f3143,plain,(
  $false | (~spl0_73 | spl0_115 | ~spl0_145 | ~spl0_159)),
  inference(subsumption_resolution,[],[f3121,f796])).
fof(f3121,plain,(
  c0_1(a547) | (~spl0_73 | ~spl0_145 | ~spl0_159)),
  inference(resolution,[],[f3077,f586])).
fof(f3101,plain,(
  ~spl0_51 | spl0_87 | ~spl0_142 | ~spl0_147),
  inference(avatar_contradiction_clause,[],[f3100])).
fof(f3100,plain,(
  $false | (~spl0_51 | spl0_87 | ~spl0_142 | ~spl0_147)),
  inference(subsumption_resolution,[],[f3087,f656])).
fof(f656,plain,(
  ~c2_1(a590) | spl0_87),
  inference(avatar_component_clause,[],[f654])).
fof(f654,plain,(
  spl0_87 <=> c2_1(a590)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_87])])).
fof(f3087,plain,(
  c2_1(a590) | (~spl0_51 | ~spl0_142 | ~spl0_147)),
  inference(resolution,[],[f3033,f476])).
fof(f476,plain,(
  c1_1(a590) | ~spl0_51),
  inference(avatar_component_clause,[],[f474])).
fof(f474,plain,(
  spl0_51 <=> c1_1(a590)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_51])])).
fof(f3033,plain,(
  ( ! [X0] : (~c1_1(X0) | c2_1(X0)) ) | (~spl0_142 | ~spl0_147)),
  inference(subsumption_resolution,[],[f919,f940])).
fof(f3032,plain,(
  spl0_122 | ~spl0_78 | ~spl0_141 | ~spl0_151 | ~spl0_155),
  inference(avatar_split_clause,[],[f3015,f977,f959,f914,f609,f829])).
fof(f829,plain,(
  spl0_122 <=> c2_1(a539)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_122])])).
fof(f609,plain,(
  spl0_78 <=> c0_1(a539)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_78])])).
fof(f3015,plain,(
  c2_1(a539) | (~spl0_78 | ~spl0_141 | ~spl0_151 | ~spl0_155)),
  inference(resolution,[],[f2981,f611])).
fof(f611,plain,(
  c0_1(a539) | ~spl0_78),
  inference(avatar_component_clause,[],[f609])).
fof(f2981,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0)) ) | (~spl0_141 | ~spl0_151 | ~spl0_155)),
  inference(subsumption_resolution,[],[f915,f2961])).
fof(f3031,plain,(
  spl0_105 | ~spl0_59 | ~spl0_141 | ~spl0_151 | ~spl0_155),
  inference(avatar_split_clause,[],[f3013,f977,f959,f914,f514,f744])).
fof(f744,plain,(
  spl0_105 <=> c2_1(a567)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_105])])).
fof(f514,plain,(
  spl0_59 <=> c0_1(a567)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_59])])).
fof(f3013,plain,(
  c2_1(a567) | (~spl0_59 | ~spl0_141 | ~spl0_151 | ~spl0_155)),
  inference(resolution,[],[f2981,f516])).
fof(f516,plain,(
  c0_1(a567) | ~spl0_59),
  inference(avatar_component_clause,[],[f514])).
fof(f3008,plain,(
  spl0_96 | spl0_97 | ~spl0_140 | ~spl0_169),
  inference(avatar_contradiction_clause,[],[f3007])).
fof(f3007,plain,(
  $false | (spl0_96 | spl0_97 | ~spl0_140 | ~spl0_169)),
  inference(subsumption_resolution,[],[f3006,f701])).
fof(f701,plain,(
  ~c1_1(a576) | spl0_96),
  inference(avatar_component_clause,[],[f699])).
fof(f699,plain,(
  spl0_96 <=> c1_1(a576)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_96])])).
fof(f3006,plain,(
  c1_1(a576) | (spl0_97 | ~spl0_140 | ~spl0_169)),
  inference(subsumption_resolution,[],[f3005,f706])).
fof(f706,plain,(
  ~c3_1(a576) | spl0_97),
  inference(avatar_component_clause,[],[f704])).
fof(f704,plain,(
  spl0_97 <=> c3_1(a576)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_97])])).
fof(f3005,plain,(
  c3_1(a576) | c1_1(a576) | (~spl0_140 | ~spl0_169)),
  inference(resolution,[],[f2864,f911])).
fof(f911,plain,(
  ( ! [X0] : (~c2_1(X0) | c3_1(X0) | c1_1(X0)) ) | ~spl0_140),
  inference(avatar_component_clause,[],[f910])).
fof(f910,plain,(
  spl0_140 <=> ! [X0] : (c3_1(X0) | ~c2_1(X0) | c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_140])])).
fof(f2864,plain,(
  c2_1(a576) | ~spl0_169),
  inference(avatar_component_clause,[],[f2862])).
fof(f2862,plain,(
  spl0_169 <=> c2_1(a576)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_169])])).
fof(f2980,plain,(
  spl0_115 | spl0_116 | ~spl0_130 | ~spl0_136),
  inference(avatar_contradiction_clause,[],[f2979])).
fof(f2979,plain,(
  $false | (spl0_115 | spl0_116 | ~spl0_130 | ~spl0_136)),
  inference(subsumption_resolution,[],[f2970,f796])).
fof(f2970,plain,(
  c0_1(a547) | (spl0_116 | ~spl0_130 | ~spl0_136)),
  inference(resolution,[],[f2929,f801])).
fof(f801,plain,(
  ~c3_1(a547) | spl0_116),
  inference(avatar_component_clause,[],[f799])).
fof(f799,plain,(
  spl0_116 <=> c3_1(a547)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_116])])).
fof(f2929,plain,(
  ( ! [X0] : (c3_1(X0) | c0_1(X0)) ) | (~spl0_130 | ~spl0_136)),
  inference(subsumption_resolution,[],[f895,f870])).
fof(f895,plain,(
  ( ! [X0] : (~c1_1(X0) | c3_1(X0) | c0_1(X0)) ) | ~spl0_136),
  inference(avatar_component_clause,[],[f894])).
fof(f894,plain,(
  spl0_136 <=> ! [X0] : (c3_1(X0) | ~c1_1(X0) | c0_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_136])])).
fof(f2957,plain,(
  spl0_122 | spl0_123 | ~spl0_131 | ~spl0_141),
  inference(avatar_contradiction_clause,[],[f2956])).
fof(f2956,plain,(
  $false | (spl0_122 | spl0_123 | ~spl0_131 | ~spl0_141)),
  inference(subsumption_resolution,[],[f2950,f836])).
fof(f836,plain,(
  ~c3_1(a539) | spl0_123),
  inference(avatar_component_clause,[],[f834])).
fof(f834,plain,(
  spl0_123 <=> c3_1(a539)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_123])])).
fof(f2950,plain,(
  c3_1(a539) | (spl0_122 | ~spl0_131 | ~spl0_141)),
  inference(resolution,[],[f2928,f831])).
fof(f831,plain,(
  ~c2_1(a539) | spl0_122),
  inference(avatar_component_clause,[],[f829])).
fof(f2928,plain,(
  ( ! [X0] : (c2_1(X0) | c3_1(X0)) ) | (~spl0_131 | ~spl0_141)),
  inference(subsumption_resolution,[],[f915,f874])).
fof(f874,plain,(
  ( ! [X0] : (c2_1(X0) | c0_1(X0) | c3_1(X0)) ) | ~spl0_131),
  inference(avatar_component_clause,[],[f873])).
fof(f873,plain,(
  spl0_131 <=> ! [X0] : (c3_1(X0) | c0_1(X0) | c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_131])])).
fof(f2927,plain,(
  ~spl0_61 | ~spl0_62 | ~spl0_140 | ~spl0_153 | ~spl0_160),
  inference(avatar_contradiction_clause,[],[f2926])).
fof(f2926,plain,(
  $false | (~spl0_61 | ~spl0_62 | ~spl0_140 | ~spl0_153 | ~spl0_160)),
  inference(subsumption_resolution,[],[f2917,f531])).
fof(f2917,plain,(
  ~c3_1(a566) | (~spl0_61 | ~spl0_140 | ~spl0_153 | ~spl0_160)),
  inference(resolution,[],[f2867,f526])).
fof(f2867,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c3_1(X0)) ) | (~spl0_140 | ~spl0_153 | ~spl0_160)),
  inference(subsumption_resolution,[],[f969,f2866])).
fof(f2866,plain,(
  ( ! [X0] : (~c2_1(X0) | c1_1(X0)) ) | (~spl0_140 | ~spl0_160)),
  inference(subsumption_resolution,[],[f1009,f911])).
fof(f1009,plain,(
  ( ! [X0] : (c1_1(X0) | ~c2_1(X0) | ~c3_1(X0)) ) | ~spl0_160),
  inference(avatar_component_clause,[],[f1008])).
fof(f1008,plain,(
  spl0_160 <=> ! [X0] : (c1_1(X0) | ~c2_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_160])])).
fof(f2889,plain,(
  spl0_82 | ~spl0_47 | ~spl0_140 | ~spl0_160),
  inference(avatar_split_clause,[],[f2874,f1008,f910,f454,f629])).
fof(f2874,plain,(
  c1_1(a615) | (~spl0_47 | ~spl0_140 | ~spl0_160)),
  inference(resolution,[],[f2866,f456])).
fof(f2886,plain,(
  ~spl0_53 | spl0_93 | ~spl0_140 | ~spl0_160),
  inference(avatar_contradiction_clause,[],[f2885])).
fof(f2885,plain,(
  $false | (~spl0_53 | spl0_93 | ~spl0_140 | ~spl0_160)),
  inference(subsumption_resolution,[],[f2876,f686])).
fof(f2876,plain,(
  c1_1(a582) | (~spl0_53 | ~spl0_140 | ~spl0_160)),
  inference(resolution,[],[f2866,f486])).
fof(f2865,plain,(
  spl0_96 | spl0_169 | ~spl0_56 | ~spl0_137),
  inference(avatar_split_clause,[],[f2777,f898,f499,f2862,f699])).
fof(f499,plain,(
  spl0_56 <=> c0_1(a576)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_56])])).
fof(f2777,plain,(
  c2_1(a576) | c1_1(a576) | (~spl0_56 | ~spl0_137)),
  inference(resolution,[],[f899,f501])).
fof(f501,plain,(
  c0_1(a576) | ~spl0_56),
  inference(avatar_component_clause,[],[f499])).
fof(f2860,plain,(
  spl0_163 | ~spl0_78 | spl0_122 | ~spl0_137),
  inference(avatar_split_clause,[],[f2859,f898,f829,f609,f1933])).
fof(f1933,plain,(
  spl0_163 <=> c1_1(a539)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_163])])).
fof(f2859,plain,(
  c1_1(a539) | (~spl0_78 | spl0_122 | ~spl0_137)),
  inference(subsumption_resolution,[],[f2779,f831])).
fof(f2779,plain,(
  c2_1(a539) | c1_1(a539) | (~spl0_78 | ~spl0_137)),
  inference(resolution,[],[f899,f611])).
fof(f2858,plain,(
  spl0_122 | spl0_123 | ~spl0_142 | ~spl0_163),
  inference(avatar_contradiction_clause,[],[f2857])).
fof(f2857,plain,(
  $false | (spl0_122 | spl0_123 | ~spl0_142 | ~spl0_163)),
  inference(subsumption_resolution,[],[f2856,f831])).
fof(f2856,plain,(
  c2_1(a539) | (spl0_123 | ~spl0_142 | ~spl0_163)),
  inference(subsumption_resolution,[],[f2851,f836])).
fof(f2851,plain,(
  c3_1(a539) | c2_1(a539) | (~spl0_142 | ~spl0_163)),
  inference(resolution,[],[f919,f1935])).
fof(f1935,plain,(
  c1_1(a539) | ~spl0_163),
  inference(avatar_component_clause,[],[f1933])).
fof(f2804,plain,(
  spl0_125 | spl0_124 | spl0_126 | ~spl0_130 | ~spl0_137),
  inference(avatar_split_clause,[],[f2803,f898,f869,f849,f839,f844])).
fof(f839,plain,(
  spl0_124 <=> c1_1(a538)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_124])])).
fof(f2803,plain,(
  c2_1(a538) | (spl0_124 | spl0_126 | ~spl0_130 | ~spl0_137)),
  inference(subsumption_resolution,[],[f2783,f841])).
fof(f841,plain,(
  ~c1_1(a538) | spl0_124),
  inference(avatar_component_clause,[],[f839])).
fof(f2783,plain,(
  c2_1(a538) | c1_1(a538) | (spl0_124 | spl0_126 | ~spl0_130 | ~spl0_137)),
  inference(resolution,[],[f899,f2742])).
fof(f2742,plain,(
  c0_1(a538) | (spl0_124 | spl0_126 | ~spl0_130)),
  inference(subsumption_resolution,[],[f2684,f841])).
fof(f2684,plain,(
  c0_1(a538) | c1_1(a538) | (spl0_126 | ~spl0_130)),
  inference(resolution,[],[f870,f851])).
fof(f851,plain,(
  ~c3_1(a538) | spl0_126),
  inference(avatar_component_clause,[],[f849])).
fof(f2797,plain,(
  spl0_86 | ~spl0_51 | spl0_87 | ~spl0_134),
  inference(avatar_split_clause,[],[f2796,f886,f654,f474,f649])).
fof(f649,plain,(
  spl0_86 <=> c0_1(a590)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_86])])).
fof(f886,plain,(
  spl0_134 <=> ! [X0] : (c2_1(X0) | ~c1_1(X0) | c0_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_134])])).
fof(f2796,plain,(
  c0_1(a590) | (~spl0_51 | spl0_87 | ~spl0_134)),
  inference(subsumption_resolution,[],[f2759,f656])).
fof(f2759,plain,(
  c2_1(a590) | c0_1(a590) | (~spl0_51 | ~spl0_134)),
  inference(resolution,[],[f887,f476])).
fof(f887,plain,(
  ( ! [X0] : (~c1_1(X0) | c2_1(X0) | c0_1(X0)) ) | ~spl0_134),
  inference(avatar_component_clause,[],[f886])).
fof(f2720,plain,(
  ~spl0_81 | spl0_129 | ~spl0_140 | ~spl0_149),
  inference(avatar_contradiction_clause,[],[f2719])).
fof(f2719,plain,(
  $false | (~spl0_81 | spl0_129 | ~spl0_140 | ~spl0_149)),
  inference(subsumption_resolution,[],[f2711,f866])).
fof(f866,plain,(
  ~c3_1(a536) | spl0_129),
  inference(avatar_component_clause,[],[f864])).
fof(f864,plain,(
  spl0_129 <=> c3_1(a536)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_129])])).
fof(f2711,plain,(
  c3_1(a536) | (~spl0_81 | ~spl0_140 | ~spl0_149)),
  inference(resolution,[],[f2697,f626])).
fof(f626,plain,(
  c2_1(a536) | ~spl0_81),
  inference(avatar_component_clause,[],[f624])).
fof(f624,plain,(
  spl0_81 <=> c2_1(a536)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_81])])).
fof(f2697,plain,(
  ( ! [X0] : (~c2_1(X0) | c3_1(X0)) ) | (~spl0_140 | ~spl0_149)),
  inference(subsumption_resolution,[],[f952,f911])).
fof(f2718,plain,(
  ~spl0_73 | spl0_116 | ~spl0_140 | ~spl0_149),
  inference(avatar_contradiction_clause,[],[f2717])).
fof(f2717,plain,(
  $false | (~spl0_73 | spl0_116 | ~spl0_140 | ~spl0_149)),
  inference(subsumption_resolution,[],[f2710,f801])).
fof(f2710,plain,(
  c3_1(a547) | (~spl0_73 | ~spl0_140 | ~spl0_149)),
  inference(resolution,[],[f2697,f586])).
fof(f2714,plain,(
  ~spl0_49 | spl0_83 | ~spl0_140 | ~spl0_149),
  inference(avatar_contradiction_clause,[],[f2713])).
fof(f2713,plain,(
  $false | (~spl0_49 | spl0_83 | ~spl0_140 | ~spl0_149)),
  inference(subsumption_resolution,[],[f2703,f636])).
fof(f2703,plain,(
  c3_1(a604) | (~spl0_49 | ~spl0_140 | ~spl0_149)),
  inference(resolution,[],[f2697,f466])).
fof(f2696,plain,(
  spl0_98 | spl0_99 | spl0_100 | ~spl0_130),
  inference(avatar_contradiction_clause,[],[f2695])).
fof(f2695,plain,(
  $false | (spl0_98 | spl0_99 | spl0_100 | ~spl0_130)),
  inference(subsumption_resolution,[],[f2694,f716])).
fof(f716,plain,(
  ~c1_1(a573) | spl0_99),
  inference(avatar_component_clause,[],[f714])).
fof(f714,plain,(
  spl0_99 <=> c1_1(a573)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_99])])).
fof(f2694,plain,(
  c1_1(a573) | (spl0_98 | spl0_100 | ~spl0_130)),
  inference(subsumption_resolution,[],[f2683,f711])).
fof(f711,plain,(
  ~c0_1(a573) | spl0_98),
  inference(avatar_component_clause,[],[f709])).
fof(f709,plain,(
  spl0_98 <=> c0_1(a573)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_98])])).
fof(f2683,plain,(
  c0_1(a573) | c1_1(a573) | (spl0_100 | ~spl0_130)),
  inference(resolution,[],[f870,f721])).
fof(f721,plain,(
  ~c3_1(a573) | spl0_100),
  inference(avatar_component_clause,[],[f719])).
fof(f719,plain,(
  spl0_100 <=> c3_1(a573)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_100])])).
fof(f2671,plain,(
  spl0_86 | spl0_87 | ~spl0_131 | ~spl0_135),
  inference(avatar_split_clause,[],[f2650,f890,f873,f654,f649])).
fof(f890,plain,(
  spl0_135 <=> ! [X0] : (c2_1(X0) | ~c3_1(X0) | c0_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_135])])).
fof(f2650,plain,(
  c0_1(a590) | (spl0_87 | ~spl0_131 | ~spl0_135)),
  inference(resolution,[],[f2649,f656])).
fof(f2649,plain,(
  ( ! [X0] : (c2_1(X0) | c0_1(X0)) ) | (~spl0_131 | ~spl0_135)),
  inference(subsumption_resolution,[],[f891,f874])).
fof(f891,plain,(
  ( ! [X0] : (~c3_1(X0) | c2_1(X0) | c0_1(X0)) ) | ~spl0_135),
  inference(avatar_component_clause,[],[f890])).
fof(f2648,plain,(
  ~spl0_69 | ~spl0_70 | ~spl0_147 | ~spl0_153),
  inference(avatar_contradiction_clause,[],[f2647])).
fof(f2647,plain,(
  $false | (~spl0_69 | ~spl0_70 | ~spl0_147 | ~spl0_153)),
  inference(subsumption_resolution,[],[f2633,f566])).
fof(f566,plain,(
  c1_1(a551) | ~spl0_69),
  inference(avatar_component_clause,[],[f564])).
fof(f564,plain,(
  spl0_69 <=> c1_1(a551)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_69])])).
fof(f2633,plain,(
  ~c1_1(a551) | (~spl0_70 | ~spl0_147 | ~spl0_153)),
  inference(resolution,[],[f2620,f571])).
fof(f571,plain,(
  c3_1(a551) | ~spl0_70),
  inference(avatar_component_clause,[],[f569])).
fof(f569,plain,(
  spl0_70 <=> c3_1(a551)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_70])])).
fof(f2620,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c1_1(X0)) ) | (~spl0_147 | ~spl0_153)),
  inference(subsumption_resolution,[],[f940,f969])).
fof(f2638,plain,(
  ~spl0_35 | ~spl0_36 | ~spl0_147 | ~spl0_153),
  inference(avatar_contradiction_clause,[],[f2637])).
fof(f2637,plain,(
  $false | (~spl0_35 | ~spl0_36 | ~spl0_147 | ~spl0_153)),
  inference(subsumption_resolution,[],[f2623,f396])).
fof(f396,plain,(
  c1_1(a563) | ~spl0_35),
  inference(avatar_component_clause,[],[f394])).
fof(f394,plain,(
  spl0_35 <=> c1_1(a563)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_35])])).
fof(f2623,plain,(
  ~c1_1(a563) | (~spl0_36 | ~spl0_147 | ~spl0_153)),
  inference(resolution,[],[f2620,f401])).
fof(f401,plain,(
  c3_1(a563) | ~spl0_36),
  inference(avatar_component_clause,[],[f399])).
fof(f399,plain,(
  spl0_36 <=> c3_1(a563)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_36])])).
fof(f2619,plain,(
  spl0_124 | spl0_126 | ~spl0_132 | ~spl0_140),
  inference(avatar_contradiction_clause,[],[f2618])).
fof(f2618,plain,(
  $false | (spl0_124 | spl0_126 | ~spl0_132 | ~spl0_140)),
  inference(subsumption_resolution,[],[f2611,f841])).
fof(f2611,plain,(
  c1_1(a538) | (spl0_126 | ~spl0_132 | ~spl0_140)),
  inference(resolution,[],[f2579,f851])).
fof(f2579,plain,(
  ( ! [X0] : (c3_1(X0) | c1_1(X0)) ) | (~spl0_132 | ~spl0_140)),
  inference(subsumption_resolution,[],[f911,f878])).
fof(f2597,plain,(
  ~spl0_34 | ~spl0_36 | ~spl0_151 | ~spl0_155),
  inference(avatar_contradiction_clause,[],[f2596])).
fof(f2596,plain,(
  $false | (~spl0_34 | ~spl0_36 | ~spl0_151 | ~spl0_155)),
  inference(subsumption_resolution,[],[f2583,f391])).
fof(f391,plain,(
  c0_1(a563) | ~spl0_34),
  inference(avatar_component_clause,[],[f389])).
fof(f389,plain,(
  spl0_34 <=> c0_1(a563)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_34])])).
fof(f2583,plain,(
  ~c0_1(a563) | (~spl0_36 | ~spl0_151 | ~spl0_155)),
  inference(resolution,[],[f2550,f401])).
fof(f2550,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c0_1(X0)) ) | (~spl0_151 | ~spl0_155)),
  inference(subsumption_resolution,[],[f978,f960])).
fof(f2545,plain,(
  ~spl0_78 | ~spl0_137 | ~spl0_150 | ~spl0_156 | ~spl0_157),
  inference(avatar_contradiction_clause,[],[f2544])).
fof(f2544,plain,(
  $false | (~spl0_78 | ~spl0_137 | ~spl0_150 | ~spl0_156 | ~spl0_157)),
  inference(resolution,[],[f2538,f611])).
fof(f2538,plain,(
  ( ! [X0] : (~c0_1(X0)) ) | (~spl0_137 | ~spl0_150 | ~spl0_156 | ~spl0_157)),
  inference(subsumption_resolution,[],[f2537,f2474])).
fof(f2474,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c0_1(X0)) ) | (~spl0_150 | ~spl0_156)),
  inference(subsumption_resolution,[],[f956,f983])).
fof(f2537,plain,(
  ( ! [X0] : (c2_1(X0) | ~c0_1(X0)) ) | (~spl0_137 | ~spl0_156 | ~spl0_157)),
  inference(subsumption_resolution,[],[f989,f2475])).
fof(f2475,plain,(
  ( ! [X0] : (~c0_1(X0) | c1_1(X0)) ) | (~spl0_137 | ~spl0_156)),
  inference(subsumption_resolution,[],[f899,f983])).
fof(f989,plain,(
  ( ! [X0] : (c2_1(X0) | ~c0_1(X0) | ~c1_1(X0)) ) | ~spl0_157),
  inference(avatar_component_clause,[],[f988])).
fof(f988,plain,(
  spl0_157 <=> ! [X0] : (c2_1(X0) | ~c0_1(X0) | ~c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_157])])).
fof(f2535,plain,(
  ~spl0_80 | ~spl0_81 | ~spl0_150 | ~spl0_156),
  inference(avatar_split_clause,[],[f2530,f982,f955,f624,f619])).
fof(f619,plain,(
  spl0_80 <=> c0_1(a536)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_80])])).
fof(f2530,plain,(
  ~c0_1(a536) | (~spl0_81 | ~spl0_150 | ~spl0_156)),
  inference(resolution,[],[f2474,f626])).
fof(f2504,plain,(
  ~spl0_41 | ~spl0_143 | ~spl0_150 | ~spl0_156 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f2503])).
fof(f2503,plain,(
  $false | (~spl0_41 | ~spl0_143 | ~spl0_150 | ~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f426,f2479])).
fof(f2479,plain,(
  ( ! [X0] : (~c2_1(X0)) ) | (~spl0_143 | ~spl0_150 | ~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f2478,f2474])).
fof(f2478,plain,(
  ( ! [X0] : (c0_1(X0) | ~c2_1(X0)) ) | (~spl0_143 | ~spl0_150 | ~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f924,f2476])).
fof(f2476,plain,(
  ( ! [X2] : (c1_1(X2) | ~c2_1(X2)) ) | (~spl0_150 | ~spl0_156 | ~spl0_161)),
  inference(subsumption_resolution,[],[f1022,f2474])).
fof(f426,plain,(
  c2_1(a552) | ~spl0_41),
  inference(avatar_component_clause,[],[f424])).
fof(f424,plain,(
  spl0_41 <=> c2_1(a552)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_41])])).
fof(f2473,plain,(
  spl0_124 | spl0_125 | ~spl0_137 | ~spl0_154),
  inference(avatar_contradiction_clause,[],[f2472])).
fof(f2472,plain,(
  $false | (spl0_124 | spl0_125 | ~spl0_137 | ~spl0_154)),
  inference(subsumption_resolution,[],[f2461,f841])).
fof(f2461,plain,(
  c1_1(a538) | (spl0_125 | ~spl0_137 | ~spl0_154)),
  inference(resolution,[],[f2222,f846])).
fof(f2222,plain,(
  ( ! [X0] : (c2_1(X0) | c1_1(X0)) ) | (~spl0_137 | ~spl0_154)),
  inference(subsumption_resolution,[],[f899,f973])).
fof(f973,plain,(
  ( ! [X1] : (c2_1(X1) | c0_1(X1) | c1_1(X1)) ) | ~spl0_154),
  inference(avatar_component_clause,[],[f972])).
fof(f972,plain,(
  spl0_154 <=> ! [X1] : (c2_1(X1) | c0_1(X1) | c1_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_154])])).
fof(f2470,plain,(
  spl0_120 | spl0_121 | ~spl0_137 | ~spl0_154),
  inference(avatar_contradiction_clause,[],[f2469])).
fof(f2469,plain,(
  $false | (spl0_120 | spl0_121 | ~spl0_137 | ~spl0_154)),
  inference(subsumption_resolution,[],[f2456,f821])).
fof(f821,plain,(
  ~c1_1(a540) | spl0_120),
  inference(avatar_component_clause,[],[f819])).
fof(f819,plain,(
  spl0_120 <=> c1_1(a540)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_120])])).
fof(f2456,plain,(
  c1_1(a540) | (spl0_121 | ~spl0_137 | ~spl0_154)),
  inference(resolution,[],[f2222,f826])).
fof(f826,plain,(
  ~c2_1(a540) | spl0_121),
  inference(avatar_component_clause,[],[f824])).
fof(f824,plain,(
  spl0_121 <=> c2_1(a540)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_121])])).
fof(f2425,plain,(
  spl0_130 | ~spl0_131 | ~spl0_140),
  inference(avatar_split_clause,[],[f2362,f910,f873,f869])).
fof(f2362,plain,(
  ( ! [X2] : (c0_1(X2) | c3_1(X2) | c1_1(X2)) ) | (~spl0_131 | ~spl0_140)),
  inference(duplicate_literal_removal,[],[f2343])).
fof(f2343,plain,(
  ( ! [X2] : (c0_1(X2) | c3_1(X2) | c3_1(X2) | c1_1(X2)) ) | (~spl0_131 | ~spl0_140)),
  inference(resolution,[],[f874,f911])).
fof(f2422,plain,(
  ~spl0_58 | spl0_103 | spl0_104 | ~spl0_142),
  inference(avatar_contradiction_clause,[],[f2421])).
fof(f2421,plain,(
  $false | (~spl0_58 | spl0_103 | spl0_104 | ~spl0_142)),
  inference(subsumption_resolution,[],[f2420,f736])).
fof(f736,plain,(
  ~c2_1(a570) | spl0_103),
  inference(avatar_component_clause,[],[f734])).
fof(f734,plain,(
  spl0_103 <=> c2_1(a570)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_103])])).
fof(f2420,plain,(
  c2_1(a570) | (~spl0_58 | spl0_104 | ~spl0_142)),
  inference(subsumption_resolution,[],[f2382,f741])).
fof(f741,plain,(
  ~c3_1(a570) | spl0_104),
  inference(avatar_component_clause,[],[f739])).
fof(f739,plain,(
  spl0_104 <=> c3_1(a570)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_104])])).
fof(f2382,plain,(
  c3_1(a570) | c2_1(a570) | (~spl0_58 | ~spl0_142)),
  inference(resolution,[],[f919,f511])).
fof(f511,plain,(
  c1_1(a570) | ~spl0_58),
  inference(avatar_component_clause,[],[f509])).
fof(f509,plain,(
  spl0_58 <=> c1_1(a570)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_58])])).
fof(f2403,plain,(
  ~spl0_51 | spl0_86 | spl0_87 | ~spl0_131 | ~spl0_144),
  inference(avatar_contradiction_clause,[],[f2402])).
fof(f2402,plain,(
  $false | (~spl0_51 | spl0_86 | spl0_87 | ~spl0_131 | ~spl0_144)),
  inference(subsumption_resolution,[],[f2401,f651])).
fof(f651,plain,(
  ~c0_1(a590) | spl0_86),
  inference(avatar_component_clause,[],[f649])).
fof(f2401,plain,(
  c0_1(a590) | (~spl0_51 | spl0_86 | spl0_87 | ~spl0_131 | ~spl0_144)),
  inference(subsumption_resolution,[],[f2390,f476])).
fof(f2390,plain,(
  ~c1_1(a590) | c0_1(a590) | (spl0_86 | spl0_87 | ~spl0_131 | ~spl0_144)),
  inference(resolution,[],[f928,f2363])).
fof(f2363,plain,(
  c3_1(a590) | (spl0_86 | spl0_87 | ~spl0_131)),
  inference(subsumption_resolution,[],[f2347,f651])).
fof(f2347,plain,(
  c0_1(a590) | c3_1(a590) | (spl0_87 | ~spl0_131)),
  inference(resolution,[],[f874,f656])).
fof(f928,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c1_1(X0) | c0_1(X0)) ) | ~spl0_144),
  inference(avatar_component_clause,[],[f927])).
fof(f927,plain,(
  spl0_144 <=> ! [X0] : (c0_1(X0) | ~c1_1(X0) | ~c3_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_144])])).
fof(f2374,plain,(
  spl0_90 | spl0_88 | spl0_89 | ~spl0_131),
  inference(avatar_split_clause,[],[f2370,f873,f664,f659,f669])).
fof(f669,plain,(
  spl0_90 <=> c3_1(a587)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_90])])).
fof(f659,plain,(
  spl0_88 <=> c0_1(a587)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_88])])).
fof(f2370,plain,(
  c3_1(a587) | (spl0_88 | spl0_89 | ~spl0_131)),
  inference(subsumption_resolution,[],[f2359,f661])).
fof(f661,plain,(
  ~c0_1(a587) | spl0_88),
  inference(avatar_component_clause,[],[f659])).
fof(f2359,plain,(
  c0_1(a587) | c3_1(a587) | (spl0_89 | ~spl0_131)),
  inference(resolution,[],[f874,f666])).
fof(f2282,plain,(
  ~spl0_61 | ~spl0_62 | ~spl0_137 | ~spl0_148 | ~spl0_152 | ~spl0_153 | ~spl0_154 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f2281])).
fof(f2281,plain,(
  $false | (~spl0_61 | ~spl0_62 | ~spl0_137 | ~spl0_148 | ~spl0_152 | ~spl0_153 | ~spl0_154 | ~spl0_161)),
  inference(subsumption_resolution,[],[f2271,f531])).
fof(f2271,plain,(
  ~c3_1(a566) | (~spl0_61 | ~spl0_137 | ~spl0_148 | ~spl0_152 | ~spl0_153 | ~spl0_154 | ~spl0_161)),
  inference(resolution,[],[f2251,f526])).
fof(f2251,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c3_1(X0)) ) | (~spl0_137 | ~spl0_148 | ~spl0_152 | ~spl0_153 | ~spl0_154 | ~spl0_161)),
  inference(subsumption_resolution,[],[f969,f2223])).
fof(f2223,plain,(
  ( ! [X0] : (c1_1(X0)) ) | (~spl0_137 | ~spl0_148 | ~spl0_152 | ~spl0_154 | ~spl0_161)),
  inference(subsumption_resolution,[],[f2222,f2139])).
fof(f2139,plain,(
  ( ! [X2] : (~c2_1(X2) | c1_1(X2)) ) | (~spl0_148 | ~spl0_152 | ~spl0_161)),
  inference(subsumption_resolution,[],[f1022,f1821])).
fof(f1821,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c0_1(X0)) ) | (~spl0_148 | ~spl0_152)),
  inference(subsumption_resolution,[],[f944,f964])).
fof(f964,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c0_1(X0) | ~c2_1(X0)) ) | ~spl0_152),
  inference(avatar_component_clause,[],[f963])).
fof(f963,plain,(
  spl0_152 <=> ! [X0] : (~c3_1(X0) | ~c0_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_152])])).
fof(f2245,plain,(
  spl0_91 | spl0_92 | ~spl0_134 | ~spl0_154),
  inference(avatar_contradiction_clause,[],[f2244])).
fof(f2244,plain,(
  $false | (spl0_91 | spl0_92 | ~spl0_134 | ~spl0_154)),
  inference(subsumption_resolution,[],[f2230,f676])).
fof(f676,plain,(
  ~c0_1(a584) | spl0_91),
  inference(avatar_component_clause,[],[f674])).
fof(f674,plain,(
  spl0_91 <=> c0_1(a584)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_91])])).
fof(f2230,plain,(
  c0_1(a584) | (spl0_92 | ~spl0_134 | ~spl0_154)),
  inference(resolution,[],[f2194,f681])).
fof(f681,plain,(
  ~c2_1(a584) | spl0_92),
  inference(avatar_component_clause,[],[f679])).
fof(f679,plain,(
  spl0_92 <=> c2_1(a584)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_92])])).
fof(f2194,plain,(
  ( ! [X0] : (c2_1(X0) | c0_1(X0)) ) | (~spl0_134 | ~spl0_154)),
  inference(subsumption_resolution,[],[f887,f973])).
fof(f2193,plain,(
  ~spl0_77 | spl0_121 | ~spl0_138 | ~spl0_147),
  inference(avatar_contradiction_clause,[],[f2192])).
fof(f2192,plain,(
  $false | (~spl0_77 | spl0_121 | ~spl0_138 | ~spl0_147)),
  inference(subsumption_resolution,[],[f2183,f826])).
fof(f2183,plain,(
  c2_1(a540) | (~spl0_77 | ~spl0_138 | ~spl0_147)),
  inference(resolution,[],[f2100,f606])).
fof(f606,plain,(
  c3_1(a540) | ~spl0_77),
  inference(avatar_component_clause,[],[f604])).
fof(f604,plain,(
  spl0_77 <=> c3_1(a540)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_77])])).
fof(f2100,plain,(
  ( ! [X0] : (~c3_1(X0) | c2_1(X0)) ) | (~spl0_138 | ~spl0_147)),
  inference(subsumption_resolution,[],[f940,f903])).
fof(f2187,plain,(
  ~spl0_52 | spl0_92 | ~spl0_138 | ~spl0_147),
  inference(avatar_contradiction_clause,[],[f2186])).
fof(f2186,plain,(
  $false | (~spl0_52 | spl0_92 | ~spl0_138 | ~spl0_147)),
  inference(subsumption_resolution,[],[f2174,f681])).
fof(f2174,plain,(
  c2_1(a584) | (~spl0_52 | ~spl0_138 | ~spl0_147)),
  inference(resolution,[],[f2100,f481])).
fof(f481,plain,(
  c3_1(a584) | ~spl0_52),
  inference(avatar_component_clause,[],[f479])).
fof(f479,plain,(
  spl0_52 <=> c3_1(a584)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_52])])).
fof(f2099,plain,(
  ~spl0_68 | spl0_109 | ~spl0_135 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f2098])).
fof(f2098,plain,(
  $false | (~spl0_68 | spl0_109 | ~spl0_135 | ~spl0_146)),
  inference(subsumption_resolution,[],[f2091,f766])).
fof(f766,plain,(
  ~c2_1(a553) | spl0_109),
  inference(avatar_component_clause,[],[f764])).
fof(f764,plain,(
  spl0_109 <=> c2_1(a553)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_109])])).
fof(f2091,plain,(
  c2_1(a553) | (~spl0_68 | ~spl0_135 | ~spl0_146)),
  inference(resolution,[],[f2051,f561])).
fof(f561,plain,(
  c3_1(a553) | ~spl0_68),
  inference(avatar_component_clause,[],[f559])).
fof(f559,plain,(
  spl0_68 <=> c3_1(a553)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_68])])).
fof(f2051,plain,(
  ( ! [X0] : (~c3_1(X0) | c2_1(X0)) ) | (~spl0_135 | ~spl0_146)),
  inference(subsumption_resolution,[],[f936,f891])).
fof(f2017,plain,(
  ~spl0_59 | ~spl0_60 | ~spl0_148 | ~spl0_152 | ~spl0_157),
  inference(avatar_contradiction_clause,[],[f2016])).
fof(f2016,plain,(
  $false | (~spl0_59 | ~spl0_60 | ~spl0_148 | ~spl0_152 | ~spl0_157)),
  inference(subsumption_resolution,[],[f2004,f516])).
fof(f2004,plain,(
  ~c0_1(a567) | (~spl0_60 | ~spl0_148 | ~spl0_152 | ~spl0_157)),
  inference(resolution,[],[f1948,f521])).
fof(f521,plain,(
  c1_1(a567) | ~spl0_60),
  inference(avatar_component_clause,[],[f519])).
fof(f519,plain,(
  spl0_60 <=> c1_1(a567)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_60])])).
fof(f1948,plain,(
  ( ! [X0] : (~c1_1(X0) | ~c0_1(X0)) ) | (~spl0_148 | ~spl0_152 | ~spl0_157)),
  inference(subsumption_resolution,[],[f989,f1821])).
fof(f1955,plain,(
  spl0_127 | spl0_128 | ~spl0_79 | ~spl0_140),
  inference(avatar_split_clause,[],[f1920,f910,f614,f859,f854])).
fof(f1920,plain,(
  c3_1(a537) | c1_1(a537) | (~spl0_79 | ~spl0_140)),
  inference(resolution,[],[f911,f616])).
fof(f1947,plain,(
  ~spl0_74 | spl0_117 | spl0_118 | ~spl0_133),
  inference(avatar_contradiction_clause,[],[f1946])).
fof(f1946,plain,(
  $false | (~spl0_74 | spl0_117 | spl0_118 | ~spl0_133)),
  inference(subsumption_resolution,[],[f1945,f806])).
fof(f806,plain,(
  ~c0_1(a546) | spl0_117),
  inference(avatar_component_clause,[],[f804])).
fof(f804,plain,(
  spl0_117 <=> c0_1(a546)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_117])])).
fof(f1945,plain,(
  c0_1(a546) | (~spl0_74 | spl0_118 | ~spl0_133)),
  inference(subsumption_resolution,[],[f1944,f811])).
fof(f811,plain,(
  ~c1_1(a546) | spl0_118),
  inference(avatar_component_clause,[],[f809])).
fof(f809,plain,(
  spl0_118 <=> c1_1(a546)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_118])])).
fof(f1944,plain,(
  c1_1(a546) | c0_1(a546) | (~spl0_74 | ~spl0_133)),
  inference(resolution,[],[f591,f883])).
fof(f883,plain,(
  ( ! [X0] : (~c3_1(X0) | c1_1(X0) | c0_1(X0)) ) | ~spl0_133),
  inference(avatar_component_clause,[],[f882])).
fof(f591,plain,(
  c3_1(a546) | ~spl0_74),
  inference(avatar_component_clause,[],[f589])).
fof(f589,plain,(
  spl0_74 <=> c3_1(a546)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_74])])).
fof(f1936,plain,(
  spl0_163 | spl0_123 | ~spl0_78 | ~spl0_139),
  inference(avatar_split_clause,[],[f1913,f906,f609,f834,f1933])).
fof(f1913,plain,(
  c3_1(a539) | c1_1(a539) | (~spl0_78 | ~spl0_139)),
  inference(resolution,[],[f907,f611])).
fof(f1903,plain,(
  ~spl0_78 | spl0_123 | ~spl0_139 | ~spl0_158),
  inference(avatar_contradiction_clause,[],[f1902])).
fof(f1902,plain,(
  $false | (~spl0_78 | spl0_123 | ~spl0_139 | ~spl0_158)),
  inference(subsumption_resolution,[],[f1899,f836])).
fof(f1899,plain,(
  c3_1(a539) | (~spl0_78 | ~spl0_139 | ~spl0_158)),
  inference(resolution,[],[f1889,f611])).
fof(f1889,plain,(
  ( ! [X0] : (~c0_1(X0) | c3_1(X0)) ) | (~spl0_139 | ~spl0_158)),
  inference(subsumption_resolution,[],[f907,f993])).
fof(f993,plain,(
  ( ! [X0] : (c3_1(X0) | ~c0_1(X0) | ~c1_1(X0)) ) | ~spl0_158),
  inference(avatar_component_clause,[],[f992])).
fof(f992,plain,(
  spl0_158 <=> ! [X0] : (c3_1(X0) | ~c0_1(X0) | ~c1_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_158])])).
fof(f1883,plain,(
  ~spl0_52 | spl0_91 | spl0_92 | ~spl0_135),
  inference(avatar_contradiction_clause,[],[f1882])).
fof(f1882,plain,(
  $false | (~spl0_52 | spl0_91 | spl0_92 | ~spl0_135)),
  inference(subsumption_resolution,[],[f1881,f676])).
fof(f1881,plain,(
  c0_1(a584) | (~spl0_52 | spl0_92 | ~spl0_135)),
  inference(subsumption_resolution,[],[f1873,f681])).
fof(f1873,plain,(
  c2_1(a584) | c0_1(a584) | (~spl0_52 | ~spl0_135)),
  inference(resolution,[],[f891,f481])).
fof(f1869,plain,(
  spl0_93 | ~spl0_53 | ~spl0_54 | ~spl0_133 | ~spl0_148 | ~spl0_152),
  inference(avatar_split_clause,[],[f1868,f963,f943,f882,f489,f484,f684])).
fof(f1868,plain,(
  c1_1(a582) | (~spl0_53 | ~spl0_54 | ~spl0_133 | ~spl0_148 | ~spl0_152)),
  inference(subsumption_resolution,[],[f1855,f1825])).
fof(f1825,plain,(
  ~c0_1(a582) | (~spl0_53 | ~spl0_148 | ~spl0_152)),
  inference(resolution,[],[f1821,f486])).
fof(f1855,plain,(
  c1_1(a582) | c0_1(a582) | (~spl0_54 | ~spl0_133)),
  inference(resolution,[],[f883,f491])).
fof(f1844,plain,(
  spl0_117 | spl0_118 | ~spl0_130 | ~spl0_133),
  inference(avatar_contradiction_clause,[],[f1843])).
fof(f1843,plain,(
  $false | (spl0_117 | spl0_118 | ~spl0_130 | ~spl0_133)),
  inference(subsumption_resolution,[],[f1833,f806])).
fof(f1833,plain,(
  c0_1(a546) | (spl0_118 | ~spl0_130 | ~spl0_133)),
  inference(resolution,[],[f1822,f811])).
fof(f1822,plain,(
  ( ! [X0] : (c1_1(X0) | c0_1(X0)) ) | (~spl0_130 | ~spl0_133)),
  inference(subsumption_resolution,[],[f870,f883])).
fof(f1820,plain,(
  ~spl0_65 | ~spl0_66 | ~spl0_146 | ~spl0_152),
  inference(avatar_contradiction_clause,[],[f1819])).
fof(f1819,plain,(
  $false | (~spl0_65 | ~spl0_66 | ~spl0_146 | ~spl0_152)),
  inference(subsumption_resolution,[],[f1810,f546])).
fof(f1810,plain,(
  ~c0_1(a564) | (~spl0_66 | ~spl0_146 | ~spl0_152)),
  inference(resolution,[],[f1748,f551])).
fof(f551,plain,(
  c3_1(a564) | ~spl0_66),
  inference(avatar_component_clause,[],[f549])).
fof(f549,plain,(
  spl0_66 <=> c3_1(a564)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_66])])).
fof(f1748,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c0_1(X0)) ) | (~spl0_146 | ~spl0_152)),
  inference(subsumption_resolution,[],[f936,f964])).
fof(f1818,plain,(
  ~spl0_63 | ~spl0_64 | ~spl0_146 | ~spl0_152),
  inference(avatar_contradiction_clause,[],[f1817])).
fof(f1817,plain,(
  $false | (~spl0_63 | ~spl0_64 | ~spl0_146 | ~spl0_152)),
  inference(subsumption_resolution,[],[f1809,f536])).
fof(f1809,plain,(
  ~c0_1(a565) | (~spl0_64 | ~spl0_146 | ~spl0_152)),
  inference(resolution,[],[f1748,f541])).
fof(f541,plain,(
  c3_1(a565) | ~spl0_64),
  inference(avatar_component_clause,[],[f539])).
fof(f539,plain,(
  spl0_64 <=> c3_1(a565)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_64])])).
fof(f1814,plain,(
  ~spl0_34 | ~spl0_36 | ~spl0_146 | ~spl0_152),
  inference(avatar_contradiction_clause,[],[f1813])).
fof(f1813,plain,(
  $false | (~spl0_34 | ~spl0_36 | ~spl0_146 | ~spl0_152)),
  inference(subsumption_resolution,[],[f1804,f391])).
fof(f1804,plain,(
  ~c0_1(a563) | (~spl0_36 | ~spl0_146 | ~spl0_152)),
  inference(resolution,[],[f1748,f401])).
fof(f1799,plain,(
  ~spl0_59 | ~spl0_60 | ~spl0_150 | ~spl0_157),
  inference(avatar_contradiction_clause,[],[f1798])).
fof(f1798,plain,(
  $false | (~spl0_59 | ~spl0_60 | ~spl0_150 | ~spl0_157)),
  inference(subsumption_resolution,[],[f1784,f516])).
fof(f1784,plain,(
  ~c0_1(a567) | (~spl0_60 | ~spl0_150 | ~spl0_157)),
  inference(resolution,[],[f1746,f521])).
fof(f1746,plain,(
  ( ! [X0] : (~c1_1(X0) | ~c0_1(X0)) ) | (~spl0_150 | ~spl0_157)),
  inference(subsumption_resolution,[],[f956,f989])).
fof(f1791,plain,(
  ~spl0_34 | ~spl0_35 | ~spl0_150 | ~spl0_157),
  inference(avatar_contradiction_clause,[],[f1790])).
fof(f1790,plain,(
  $false | (~spl0_34 | ~spl0_35 | ~spl0_150 | ~spl0_157)),
  inference(subsumption_resolution,[],[f1776,f391])).
fof(f1776,plain,(
  ~c0_1(a563) | (~spl0_35 | ~spl0_150 | ~spl0_157)),
  inference(resolution,[],[f1746,f396])).
fof(f1744,plain,(
  spl0_93 | ~spl0_54 | ~spl0_133 | ~spl0_155),
  inference(avatar_split_clause,[],[f1645,f977,f882,f489,f684])).
fof(f1645,plain,(
  c1_1(a582) | (~spl0_54 | ~spl0_133 | ~spl0_155)),
  inference(resolution,[],[f1640,f491])).
fof(f1640,plain,(
  ( ! [X0] : (~c3_1(X0) | c1_1(X0)) ) | (~spl0_133 | ~spl0_155)),
  inference(subsumption_resolution,[],[f883,f978])).
fof(f1743,plain,(
  ~spl0_80 | ~spl0_81 | ~spl0_148 | ~spl0_152),
  inference(avatar_contradiction_clause,[],[f1742])).
fof(f1742,plain,(
  $false | (~spl0_80 | ~spl0_81 | ~spl0_148 | ~spl0_152)),
  inference(subsumption_resolution,[],[f1741,f621])).
fof(f621,plain,(
  c0_1(a536) | ~spl0_80),
  inference(avatar_component_clause,[],[f619])).
fof(f1741,plain,(
  ~c0_1(a536) | (~spl0_81 | ~spl0_148 | ~spl0_152)),
  inference(resolution,[],[f1688,f626])).
fof(f1688,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c0_1(X0)) ) | (~spl0_148 | ~spl0_152)),
  inference(subsumption_resolution,[],[f964,f944])).
fof(f1656,plain,(
  spl0_108 | ~spl0_66 | ~spl0_133 | ~spl0_155),
  inference(avatar_split_clause,[],[f1648,f977,f882,f549,f759])).
fof(f1648,plain,(
  c1_1(a564) | (~spl0_66 | ~spl0_133 | ~spl0_155)),
  inference(resolution,[],[f1640,f551])).
fof(f1655,plain,(
  ~spl0_77 | spl0_120 | ~spl0_133 | ~spl0_155),
  inference(avatar_contradiction_clause,[],[f1654])).
fof(f1654,plain,(
  $false | (~spl0_77 | spl0_120 | ~spl0_133 | ~spl0_155)),
  inference(subsumption_resolution,[],[f1651,f821])).
fof(f1651,plain,(
  c1_1(a540) | (~spl0_77 | ~spl0_133 | ~spl0_155)),
  inference(resolution,[],[f1640,f606])).
fof(f1439,plain,(
  ~spl0_62 | spl0_106 | ~spl0_133 | ~spl0_144),
  inference(avatar_contradiction_clause,[],[f1438])).
fof(f1438,plain,(
  $false | (~spl0_62 | spl0_106 | ~spl0_133 | ~spl0_144)),
  inference(subsumption_resolution,[],[f1429,f751])).
fof(f1429,plain,(
  c0_1(a566) | (~spl0_62 | ~spl0_133 | ~spl0_144)),
  inference(resolution,[],[f1398,f531])).
fof(f1398,plain,(
  ( ! [X0] : (~c3_1(X0) | c0_1(X0)) ) | (~spl0_133 | ~spl0_144)),
  inference(subsumption_resolution,[],[f883,f928])).
fof(f1437,plain,(
  ~spl0_52 | spl0_91 | ~spl0_133 | ~spl0_144),
  inference(avatar_contradiction_clause,[],[f1436])).
fof(f1436,plain,(
  $false | (~spl0_52 | spl0_91 | ~spl0_133 | ~spl0_144)),
  inference(subsumption_resolution,[],[f1427,f676])).
fof(f1427,plain,(
  c0_1(a584) | (~spl0_52 | ~spl0_133 | ~spl0_144)),
  inference(resolution,[],[f1398,f481])).
fof(f1291,plain,(
  ~spl0_52 | spl0_91 | ~spl0_140 | ~spl0_144 | ~spl0_145 | ~spl0_152 | ~spl0_154),
  inference(avatar_contradiction_clause,[],[f1290])).
fof(f1290,plain,(
  $false | (~spl0_52 | spl0_91 | ~spl0_140 | ~spl0_144 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(subsumption_resolution,[],[f1281,f676])).
fof(f1281,plain,(
  c0_1(a584) | (~spl0_52 | ~spl0_140 | ~spl0_144 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(resolution,[],[f1266,f481])).
fof(f1266,plain,(
  ( ! [X0] : (~c3_1(X0) | c0_1(X0)) ) | (~spl0_140 | ~spl0_144 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(subsumption_resolution,[],[f928,f1137])).
fof(f1137,plain,(
  ( ! [X1] : (c1_1(X1) | c0_1(X1)) ) | (~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(subsumption_resolution,[],[f973,f1108])).
fof(f1108,plain,(
  ( ! [X0] : (~c2_1(X0) | c1_1(X0)) ) | (~spl0_140 | ~spl0_145 | ~spl0_152)),
  inference(subsumption_resolution,[],[f911,f1083])).
fof(f1083,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c3_1(X0)) ) | (~spl0_145 | ~spl0_152)),
  inference(subsumption_resolution,[],[f964,f932])).
fof(f1262,plain,(
  ~spl0_69 | ~spl0_70 | ~spl0_144 | ~spl0_151),
  inference(avatar_contradiction_clause,[],[f1261])).
fof(f1261,plain,(
  $false | (~spl0_69 | ~spl0_70 | ~spl0_144 | ~spl0_151)),
  inference(subsumption_resolution,[],[f1252,f566])).
fof(f1252,plain,(
  ~c1_1(a551) | (~spl0_70 | ~spl0_144 | ~spl0_151)),
  inference(resolution,[],[f1212,f571])).
fof(f1212,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c1_1(X0)) ) | (~spl0_144 | ~spl0_151)),
  inference(subsumption_resolution,[],[f928,f960])).
fof(f1260,plain,(
  ~spl0_67 | ~spl0_68 | ~spl0_144 | ~spl0_151),
  inference(avatar_contradiction_clause,[],[f1259])).
fof(f1259,plain,(
  $false | (~spl0_67 | ~spl0_68 | ~spl0_144 | ~spl0_151)),
  inference(subsumption_resolution,[],[f1251,f556])).
fof(f556,plain,(
  c1_1(a553) | ~spl0_67),
  inference(avatar_component_clause,[],[f554])).
fof(f554,plain,(
  spl0_67 <=> c1_1(a553)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_67])])).
fof(f1251,plain,(
  ~c1_1(a553) | (~spl0_68 | ~spl0_144 | ~spl0_151)),
  inference(resolution,[],[f1212,f561])).
fof(f1256,plain,(
  ~spl0_35 | ~spl0_36 | ~spl0_144 | ~spl0_151),
  inference(avatar_contradiction_clause,[],[f1255])).
fof(f1255,plain,(
  $false | (~spl0_35 | ~spl0_36 | ~spl0_144 | ~spl0_151)),
  inference(subsumption_resolution,[],[f1243,f396])).
fof(f1243,plain,(
  ~c1_1(a563) | (~spl0_36 | ~spl0_144 | ~spl0_151)),
  inference(resolution,[],[f1212,f401])).
fof(f1189,plain,(
  spl0_108 | ~spl0_137 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154),
  inference(avatar_contradiction_clause,[],[f1176])).
fof(f1176,plain,(
  $false | (spl0_108 | ~spl0_137 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(resolution,[],[f1171,f761])).
fof(f1171,plain,(
  ( ! [X0] : (c1_1(X0)) ) | (~spl0_137 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(subsumption_resolution,[],[f1170,f1108])).
fof(f1170,plain,(
  ( ! [X0] : (c2_1(X0) | c1_1(X0)) ) | (~spl0_137 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(subsumption_resolution,[],[f899,f1137])).
fof(f1184,plain,(
  spl0_112 | ~spl0_137 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154),
  inference(avatar_contradiction_clause,[],[f1181])).
fof(f1181,plain,(
  $false | (spl0_112 | ~spl0_137 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(resolution,[],[f1171,f781])).
fof(f781,plain,(
  ~c1_1(a549) | spl0_112),
  inference(avatar_component_clause,[],[f779])).
fof(f779,plain,(
  spl0_112 <=> c1_1(a549)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_112])])).
fof(f1159,plain,(
  spl0_117 | spl0_118 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154),
  inference(avatar_split_clause,[],[f1142,f972,f963,f931,f910,f809,f804])).
fof(f1142,plain,(
  c0_1(a546) | (spl0_118 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(resolution,[],[f1137,f811])).
fof(f1157,plain,(
  spl0_98 | spl0_99 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154),
  inference(avatar_split_clause,[],[f1145,f972,f963,f931,f910,f714,f709])).
fof(f1145,plain,(
  c0_1(a573) | (spl0_99 | ~spl0_140 | ~spl0_145 | ~spl0_152 | ~spl0_154)),
  inference(resolution,[],[f1137,f716])).
fof(f1136,plain,(
  spl0_111 | spl0_112 | ~spl0_154 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f1135])).
fof(f1135,plain,(
  $false | (spl0_111 | spl0_112 | ~spl0_154 | ~spl0_161)),
  inference(subsumption_resolution,[],[f1127,f776])).
fof(f776,plain,(
  ~c0_1(a549) | spl0_111),
  inference(avatar_component_clause,[],[f774])).
fof(f774,plain,(
  spl0_111 <=> c0_1(a549)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_111])])).
fof(f1127,plain,(
  c0_1(a549) | (spl0_112 | ~spl0_154 | ~spl0_161)),
  inference(resolution,[],[f1118,f781])).
fof(f1118,plain,(
  ( ! [X1] : (c1_1(X1) | c0_1(X1)) ) | (~spl0_154 | ~spl0_161)),
  inference(subsumption_resolution,[],[f973,f1022])).
fof(f1134,plain,(
  spl0_98 | spl0_99 | ~spl0_154 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f1133])).
fof(f1133,plain,(
  $false | (spl0_98 | spl0_99 | ~spl0_154 | ~spl0_161)),
  inference(subsumption_resolution,[],[f1126,f711])).
fof(f1126,plain,(
  c0_1(a573) | (spl0_99 | ~spl0_154 | ~spl0_161)),
  inference(resolution,[],[f1118,f716])).
fof(f1132,plain,(
  spl0_117 | spl0_118 | ~spl0_154 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f1131])).
fof(f1131,plain,(
  $false | (spl0_117 | spl0_118 | ~spl0_154 | ~spl0_161)),
  inference(subsumption_resolution,[],[f1123,f806])).
fof(f1123,plain,(
  c0_1(a546) | (spl0_118 | ~spl0_154 | ~spl0_161)),
  inference(resolution,[],[f1118,f811])).
fof(f1097,plain,(
  ~spl0_53 | ~spl0_54 | ~spl0_145 | ~spl0_152),
  inference(avatar_contradiction_clause,[],[f1096])).
fof(f1096,plain,(
  $false | (~spl0_53 | ~spl0_54 | ~spl0_145 | ~spl0_152)),
  inference(subsumption_resolution,[],[f1095,f491])).
fof(f1095,plain,(
  ~c3_1(a582) | (~spl0_53 | ~spl0_145 | ~spl0_152)),
  inference(resolution,[],[f486,f1083])).
fof(f1092,plain,(
  ~spl0_41 | ~spl0_42 | ~spl0_145 | ~spl0_152),
  inference(avatar_contradiction_clause,[],[f1091])).
fof(f1091,plain,(
  $false | (~spl0_41 | ~spl0_42 | ~spl0_145 | ~spl0_152)),
  inference(subsumption_resolution,[],[f1084,f431])).
fof(f431,plain,(
  c3_1(a552) | ~spl0_42),
  inference(avatar_component_clause,[],[f429])).
fof(f429,plain,(
  spl0_42 <=> c3_1(a552)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_42])])).
fof(f1084,plain,(
  ~c3_1(a552) | (~spl0_41 | ~spl0_145 | ~spl0_152)),
  inference(resolution,[],[f1083,f426])).
fof(f1082,plain,(
  ~spl0_80 | spl0_129 | ~spl0_141 | ~spl0_148),
  inference(avatar_contradiction_clause,[],[f1081])).
fof(f1081,plain,(
  $false | (~spl0_80 | spl0_129 | ~spl0_141 | ~spl0_148)),
  inference(subsumption_resolution,[],[f1076,f866])).
fof(f1076,plain,(
  c3_1(a536) | (~spl0_80 | ~spl0_141 | ~spl0_148)),
  inference(resolution,[],[f1067,f621])).
fof(f1067,plain,(
  ( ! [X0] : (~c0_1(X0) | c3_1(X0)) ) | (~spl0_141 | ~spl0_148)),
  inference(subsumption_resolution,[],[f915,f944])).
fof(f1080,plain,(
  ~spl0_78 | spl0_123 | ~spl0_141 | ~spl0_148),
  inference(avatar_contradiction_clause,[],[f1079])).
fof(f1079,plain,(
  $false | (~spl0_78 | spl0_123 | ~spl0_141 | ~spl0_148)),
  inference(subsumption_resolution,[],[f1075,f836])).
fof(f1075,plain,(
  c3_1(a539) | (~spl0_78 | ~spl0_141 | ~spl0_148)),
  inference(resolution,[],[f1067,f611])).
fof(f1046,plain,(
  ~spl0_78 | spl0_122 | ~spl0_141 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f1045])).
fof(f1045,plain,(
  $false | (~spl0_78 | spl0_122 | ~spl0_141 | ~spl0_146)),
  inference(subsumption_resolution,[],[f1039,f831])).
fof(f1039,plain,(
  c2_1(a539) | (~spl0_78 | ~spl0_141 | ~spl0_146)),
  inference(resolution,[],[f1029,f611])).
fof(f1029,plain,(
  ( ! [X0] : (~c0_1(X0) | c2_1(X0)) ) | (~spl0_141 | ~spl0_146)),
  inference(subsumption_resolution,[],[f915,f936])).
fof(f1028,plain,(
  spl0_152 | ~spl0_2 | spl0_151 | spl0_136),
  inference(avatar_split_clause,[],[f208,f894,f959,f214,f963])).
fof(f214,plain,(
  spl0_2 <=> ndr1_0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f208,axiom,(
  ( ! [X2,X0,X1] : (c3_1(X2) | c0_1(X2) | ~c1_1(X2) | ~c3_1(X1) | ~c1_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1027,plain,(
  ~spl0_2 | spl0_148 | spl0_146 | spl0_155),
  inference(avatar_split_clause,[],[f207,f977,f935,f943,f214])).
fof(f207,axiom,(
  ( ! [X2,X0,X1] : (c1_1(X2) | c2_1(X1) | c3_1(X0) | ~c3_1(X2) | ~c0_1(X2) | ~c3_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1026,plain,(
  spl0_153 | ~spl0_2 | spl0_145 | spl0_134),
  inference(avatar_split_clause,[],[f206,f886,f931,f214,f968])).
fof(f206,axiom,(
  ( ! [X2,X0,X1] : (c2_1(X2) | c0_1(X2) | c0_1(X1) | ~c1_1(X2) | ~c3_1(X1) | ~c2_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1025,plain,(
  ~spl0_2 | spl0_146 | spl0_140 | spl0_145),
  inference(avatar_split_clause,[],[f205,f931,f910,f935,f214])).
fof(f205,axiom,(
  ( ! [X2,X0,X1] : (c0_1(X2) | c3_1(X1) | c1_1(X1) | c2_1(X0) | ~c3_1(X2) | ~c2_1(X2) | ~c2_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1024,plain,(
  spl0_151 | ~spl0_2 | spl0_157 | spl0_131),
  inference(avatar_split_clause,[],[f204,f873,f988,f214,f959])).
fof(f204,axiom,(
  ( ! [X2,X0,X1] : (c3_1(X2) | c2_1(X2) | c0_1(X2) | c2_1(X1) | ~c1_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1023,plain,(
  ~spl0_2 | spl0_145 | spl0_133 | spl0_161),
  inference(avatar_split_clause,[],[f203,f1021,f882,f931,f214])).
fof(f203,axiom,(
  ( ! [X2,X0,X1] : (c1_1(X2) | c0_1(X2) | c1_1(X1) | c0_1(X1) | c0_1(X0) | ~c2_1(X2) | ~c3_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1019,plain,(
  ~spl0_2 | spl0_139 | spl0_135 | spl0_134),
  inference(avatar_split_clause,[],[f202,f886,f890,f906,f214])).
fof(f202,axiom,(
  ( ! [X2,X0,X1] : (c2_1(X2) | c0_1(X2) | c2_1(X1) | c0_1(X1) | c3_1(X0) | c1_1(X0) | ~c1_1(X2) | ~c3_1(X1) | ~ndr1_0 | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1018,plain,(
  ~spl0_2 | spl0_157 | spl0_139 | spl0_130),
  inference(avatar_split_clause,[],[f201,f869,f906,f988,f214])).
fof(f201,axiom,(
  ( ! [X2,X0,X1] : (c3_1(X2) | c1_1(X2) | c0_1(X2) | c3_1(X1) | c1_1(X1) | c2_1(X0) | ~c0_1(X1) | ~ndr1_0 | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1017,plain,(
  spl0_152 | ~spl0_2 | spl0_151 | spl0_19),
  inference(avatar_split_clause,[],[f200,f299,f959,f214,f963])).
fof(f299,plain,(
  spl0_19 <=> hskp14),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f200,axiom,(
  ( ! [X0,X1] : (hskp14 | ~c3_1(X1) | ~c1_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1015,plain,(
  spl0_150 | ~spl0_2 | spl0_145 | spl0_21),
  inference(avatar_split_clause,[],[f198,f309,f931,f214,f955])).
fof(f309,plain,(
  spl0_21 <=> hskp12),
  introduced(avatar_definition,[new_symbols(naming,[spl0_21])])).
fof(f198,axiom,(
  ( ! [X0,X1] : (hskp12 | c0_1(X1) | ~c3_1(X1) | ~c2_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1014,plain,(
  spl0_152 | ~spl0_2 | spl0_142 | spl0_19),
  inference(avatar_split_clause,[],[f197,f299,f918,f214,f963])).
fof(f197,axiom,(
  ( ! [X0,X1] : (hskp14 | c3_1(X1) | c2_1(X1) | ~c1_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1013,plain,(
  spl0_151 | ~spl0_2 | spl0_142 | spl0_9),
  inference(avatar_split_clause,[],[f196,f249,f918,f214,f959])).
fof(f249,plain,(
  spl0_9 <=> hskp24),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f196,axiom,(
  ( ! [X0,X1] : (hskp24 | c3_1(X1) | c2_1(X1) | ~c1_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1012,plain,(
  ~spl0_2 | spl0_147 | spl0_156 | spl0_30),
  inference(avatar_split_clause,[],[f195,f354,f982,f939,f214])).
fof(f354,plain,(
  spl0_30 <=> hskp3),
  introduced(avatar_definition,[new_symbols(naming,[spl0_30])])).
fof(f195,axiom,(
  ( ! [X0,X1] : (hskp3 | c1_1(X1) | c2_1(X0) | ~c2_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1011,plain,(
  ~spl0_2 | spl0_157 | spl0_156 | spl0_26),
  inference(avatar_split_clause,[],[f194,f334,f982,f988,f214])).
fof(f334,plain,(
  spl0_26 <=> hskp7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_26])])).
fof(f194,axiom,(
  ( ! [X0,X1] : (hskp7 | c1_1(X1) | c2_1(X0) | ~c2_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1010,plain,(
  ~spl0_2 | spl0_160 | spl0_156 | spl0_14),
  inference(avatar_split_clause,[],[f193,f274,f982,f1008,f214])).
fof(f274,plain,(
  spl0_14 <=> hskp19),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f193,axiom,(
  ( ! [X0,X1] : (hskp19 | c1_1(X1) | c1_1(X0) | ~c2_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1006,plain,(
  spl0_150 | ~spl0_2 | spl0_139 | spl0_20),
  inference(avatar_split_clause,[],[f192,f304,f906,f214,f955])).
fof(f304,plain,(
  spl0_20 <=> hskp13),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f192,axiom,(
  ( ! [X0,X1] : (hskp13 | c3_1(X1) | c1_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1005,plain,(
  ~spl0_2 | spl0_156 | spl0_145 | spl0_1),
  inference(avatar_split_clause,[],[f191,f210,f931,f982,f214])).
fof(f210,plain,(
  spl0_1 <=> hskp31),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f191,axiom,(
  ( ! [X0,X1] : (hskp31 | c0_1(X1) | c1_1(X0) | ~c3_1(X1) | ~c2_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1004,plain,(
  ~spl0_2 | spl0_157 | spl0_144 | spl0_27),
  inference(avatar_split_clause,[],[f190,f339,f927,f988,f214])).
fof(f339,plain,(
  spl0_27 <=> hskp6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_27])])).
fof(f190,axiom,(
  ( ! [X0,X1] : (hskp6 | c0_1(X1) | c2_1(X0) | ~c3_1(X1) | ~c1_1(X1) | ~ndr1_0 | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f1003,plain,(
  spl0_152 | ~spl0_2 | spl0_159 | spl0_33),
  inference(avatar_split_clause,[],[f189,f369,f1001,f214,f963])).
fof(f369,plain,(
  spl0_33 <=> hskp0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_33])])).
fof(f189,axiom,(
  ( ! [X0,X1] : (hskp0 | c3_1(X1) | c0_1(X1) | ~c2_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f998,plain,(
  spl0_150 | ~spl0_2 | spl0_133 | spl0_32),
  inference(avatar_split_clause,[],[f187,f364,f882,f214,f955])).
fof(f364,plain,(
  spl0_32 <=> hskp1),
  introduced(avatar_definition,[new_symbols(naming,[spl0_32])])).
fof(f187,axiom,(
  ( ! [X0,X1] : (hskp1 | c1_1(X1) | c0_1(X1) | ~c3_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f997,plain,(
  spl0_152 | ~spl0_2 | spl0_141),
  inference(avatar_split_clause,[],[f186,f914,f214,f963])).
fof(f186,axiom,(
  ( ! [X0,X1] : (c3_1(X1) | c2_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f996,plain,(
  ~spl0_2 | spl0_142 | spl0_156 | spl0_11),
  inference(avatar_split_clause,[],[f185,f259,f982,f918,f214])).
fof(f259,plain,(
  spl0_11 <=> hskp22),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f185,axiom,(
  ( ! [X0,X1] : (hskp22 | c1_1(X1) | c3_1(X0) | c2_1(X0) | ~c2_1(X1) | ~c0_1(X1) | ~ndr1_0 | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f995,plain,(
  ~spl0_2 | spl0_149 | spl0_137 | spl0_15),
  inference(avatar_split_clause,[],[f184,f279,f898,f951,f214])).
fof(f279,plain,(
  spl0_15 <=> hskp18),
  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).
fof(f184,axiom,(
  ( ! [X0,X1] : (hskp18 | c2_1(X1) | c1_1(X1) | c3_1(X0) | ~c0_1(X1) | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f994,plain,(
  ~spl0_2 | spl0_158 | spl0_137 | spl0_21),
  inference(avatar_split_clause,[],[f183,f309,f898,f992,f214])).
fof(f183,axiom,(
  ( ! [X0,X1] : (hskp12 | c2_1(X1) | c1_1(X1) | c3_1(X0) | ~c0_1(X1) | ~ndr1_0 | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f985,plain,(
  ~spl0_2 | spl0_146 | spl0_131 | spl0_26),
  inference(avatar_split_clause,[],[f180,f334,f873,f935,f214])).
fof(f180,axiom,(
  ( ! [X0,X1] : (hskp7 | c3_1(X1) | c2_1(X1) | c0_1(X1) | c2_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f984,plain,(
  ~spl0_2 | spl0_156 | spl0_131 | spl0_27),
  inference(avatar_split_clause,[],[f179,f339,f873,f982,f214])).
fof(f179,axiom,(
  ( ! [X0,X1] : (hskp6 | c3_1(X1) | c2_1(X1) | c0_1(X1) | c1_1(X0) | ~ndr1_0 | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f980,plain,(
  ~spl0_2 | spl0_146 | spl0_130 | spl0_29),
  inference(avatar_split_clause,[],[f178,f349,f869,f935,f214])).
fof(f349,plain,(
  spl0_29 <=> hskp4),
  introduced(avatar_definition,[new_symbols(naming,[spl0_29])])).
fof(f178,axiom,(
  ( ! [X0,X1] : (hskp4 | c3_1(X1) | c1_1(X1) | c0_1(X1) | c2_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f979,plain,(
  ~spl0_2 | spl0_155 | spl0_154 | spl0_32),
  inference(avatar_split_clause,[],[f177,f364,f972,f977,f214])).
fof(f177,axiom,(
  ( ! [X0,X1] : (hskp1 | c2_1(X1) | c1_1(X1) | c0_1(X1) | c1_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f975,plain,(
  ~spl0_2 | spl0_138 | spl0_130 | spl0_30),
  inference(avatar_split_clause,[],[f176,f354,f869,f902,f214])).
fof(f176,axiom,(
  ( ! [X0,X1] : (hskp3 | c3_1(X1) | c1_1(X1) | c0_1(X1) | c2_1(X0) | c1_1(X0) | ~ndr1_0 | ~c3_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f974,plain,(
  ~spl0_2 | spl0_133 | spl0_154 | spl0_33),
  inference(avatar_split_clause,[],[f175,f369,f972,f882,f214])).
fof(f175,axiom,(
  ( ! [X0,X1] : (hskp0 | c2_1(X1) | c1_1(X1) | c0_1(X1) | c1_1(X0) | c0_1(X0) | ~ndr1_0 | ~c3_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f970,plain,(
  spl0_153 | ~spl0_2 | spl0_13 | spl0_19),
  inference(avatar_split_clause,[],[f174,f299,f269,f214,f968])).
fof(f269,plain,(
  spl0_13 <=> hskp20),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f174,axiom,(
  ( ! [X0] : (hskp14 | hskp20 | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f966,plain,(
  spl0_152 | ~spl0_2 | spl0_23 | spl0_19),
  inference(avatar_split_clause,[],[f173,f299,f319,f214,f963])).
fof(f319,plain,(
  spl0_23 <=> hskp10),
  introduced(avatar_definition,[new_symbols(naming,[spl0_23])])).
fof(f173,axiom,(
  ( ! [X0] : (hskp14 | hskp10 | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f965,plain,(
  spl0_152 | ~spl0_2 | spl0_7 | spl0_29),
  inference(avatar_split_clause,[],[f172,f349,f239,f214,f963])).
fof(f239,plain,(
  spl0_7 <=> hskp26),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f172,axiom,(
  ( ! [X0] : (hskp4 | hskp26 | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f961,plain,(
  spl0_151 | ~spl0_2 | spl0_6 | spl0_7),
  inference(avatar_split_clause,[],[f171,f239,f234,f214,f959])).
fof(f234,plain,(
  spl0_6 <=> hskp27),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f171,axiom,(
  ( ! [X0] : (hskp26 | hskp27 | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f957,plain,(
  spl0_150 | ~spl0_2 | spl0_7 | spl0_19),
  inference(avatar_split_clause,[],[f170,f299,f239,f214,f955])).
fof(f170,axiom,(
  ( ! [X0] : (hskp14 | hskp26 | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f948,plain,(
  ~spl0_2 | spl0_148 | spl0_19 | spl0_12),
  inference(avatar_split_clause,[],[f167,f264,f299,f943,f214])).
fof(f264,plain,(
  spl0_12 <=> hskp21),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f167,axiom,(
  ( ! [X0] : (hskp21 | hskp14 | c3_1(X0) | ~ndr1_0 | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f947,plain,(
  ~spl0_2 | spl0_148 | spl0_7 | spl0_9),
  inference(avatar_split_clause,[],[f166,f249,f239,f943,f214])).
fof(f166,axiom,(
  ( ! [X0] : (hskp24 | hskp26 | c3_1(X0) | ~ndr1_0 | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f945,plain,(
  ~spl0_2 | spl0_148 | spl0_21 | spl0_12),
  inference(avatar_split_clause,[],[f164,f264,f309,f943,f214])).
fof(f164,axiom,(
  ( ! [X0] : (hskp21 | hskp12 | c3_1(X0) | ~ndr1_0 | ~c2_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f941,plain,(
  ~spl0_2 | spl0_147 | spl0_30 | spl0_27),
  inference(avatar_split_clause,[],[f163,f339,f354,f939,f214])).
fof(f163,axiom,(
  ( ! [X0] : (hskp6 | hskp3 | c2_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f937,plain,(
  ~spl0_2 | spl0_146 | spl0_18 | spl0_33),
  inference(avatar_split_clause,[],[f162,f369,f294,f935,f214])).
fof(f294,plain,(
  spl0_18 <=> hskp15),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f162,axiom,(
  ( ! [X0] : (hskp0 | hskp15 | c2_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f933,plain,(
  ~spl0_2 | spl0_145 | spl0_20 | spl0_19),
  inference(avatar_split_clause,[],[f161,f299,f304,f931,f214])).
fof(f161,axiom,(
  ( ! [X0] : (hskp14 | hskp13 | c0_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c2_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f929,plain,(
  ~spl0_2 | spl0_144 | spl0_33 | spl0_3),
  inference(avatar_split_clause,[],[f160,f219,f369,f927,f214])).
fof(f219,plain,(
  spl0_3 <=> hskp30),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f160,axiom,(
  ( ! [X0] : (hskp30 | hskp0 | c0_1(X0) | ~ndr1_0 | ~c3_1(X0) | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f925,plain,(
  ~spl0_2 | spl0_143 | spl0_23 | spl0_24),
  inference(avatar_split_clause,[],[f159,f324,f319,f923,f214])).
fof(f324,plain,(
  spl0_24 <=> hskp9),
  introduced(avatar_definition,[new_symbols(naming,[spl0_24])])).
fof(f159,axiom,(
  ( ! [X0] : (hskp9 | hskp10 | c0_1(X0) | ~ndr1_0 | ~c2_1(X0) | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f920,plain,(
  ~spl0_2 | spl0_142 | spl0_30 | spl0_27),
  inference(avatar_split_clause,[],[f157,f339,f354,f918,f214])).
fof(f157,axiom,(
  ( ! [X0] : (hskp6 | hskp3 | c3_1(X0) | c2_1(X0) | ~ndr1_0 | ~c1_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f916,plain,(
  ~spl0_2 | spl0_141 | spl0_12 | spl0_11),
  inference(avatar_split_clause,[],[f156,f259,f264,f914,f214])).
fof(f156,axiom,(
  ( ! [X0] : (hskp22 | hskp21 | c3_1(X0) | c2_1(X0) | ~ndr1_0 | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f912,plain,(
  ~spl0_2 | spl0_140 | spl0_4 | spl0_12),
  inference(avatar_split_clause,[],[f155,f264,f224,f910,f214])).
fof(f224,plain,(
  spl0_4 <=> hskp29),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f155,axiom,(
  ( ! [X0] : (hskp21 | hskp29 | c3_1(X0) | c1_1(X0) | ~ndr1_0 | ~c2_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f900,plain,(
  ~spl0_2 | spl0_137 | spl0_21 | spl0_22),
  inference(avatar_split_clause,[],[f152,f314,f309,f898,f214])).
fof(f314,plain,(
  spl0_22 <=> hskp11),
  introduced(avatar_definition,[new_symbols(naming,[spl0_22])])).
fof(f152,axiom,(
  ( ! [X0] : (hskp11 | hskp12 | c2_1(X0) | c1_1(X0) | ~ndr1_0 | ~c0_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f892,plain,(
  ~spl0_2 | spl0_135 | spl0_4 | spl0_22),
  inference(avatar_split_clause,[],[f150,f314,f224,f890,f214])).
fof(f150,axiom,(
  ( ! [X0] : (hskp11 | hskp29 | c2_1(X0) | c0_1(X0) | ~ndr1_0 | ~c3_1(X0)) )),
  file('SYN510-1.tptp',unknown)).
fof(f880,plain,(
  ~spl0_2 | spl0_132 | spl0_30 | spl0_17),
  inference(avatar_split_clause,[],[f147,f289,f354,f877,f214])).
fof(f289,plain,(
  spl0_17 <=> hskp16),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f147,axiom,(
  ( ! [X0] : (hskp16 | hskp3 | c3_1(X0) | c2_1(X0) | c1_1(X0) | ~ndr1_0) )),
  file('SYN510-1.tptp',unknown)).
fof(f867,plain,(
  ~spl0_129 | ~spl0_33),
  inference(avatar_split_clause,[],[f143,f369,f864])).
fof(f143,axiom,(
  ~hskp0 | ~c3_1(a536)),
  file('SYN510-1.tptp',unknown)).
fof(f862,plain,(
  ~spl0_128 | ~spl0_32),
  inference(avatar_split_clause,[],[f142,f364,f859])).
fof(f142,axiom,(
  ~hskp1 | ~c3_1(a537)),
  file('SYN510-1.tptp',unknown)).
fof(f857,plain,(
  ~spl0_127 | ~spl0_32),
  inference(avatar_split_clause,[],[f141,f364,f854])).
fof(f141,axiom,(
  ~hskp1 | ~c1_1(a537)),
  file('SYN510-1.tptp',unknown)).
fof(f852,plain,(
  ~spl0_126 | ~spl0_31),
  inference(avatar_split_clause,[],[f140,f359,f849])).
fof(f359,plain,(
  spl0_31 <=> hskp2),
  introduced(avatar_definition,[new_symbols(naming,[spl0_31])])).
fof(f140,axiom,(
  ~hskp2 | ~c3_1(a538)),
  file('SYN510-1.tptp',unknown)).
fof(f847,plain,(
  ~spl0_125 | ~spl0_31),
  inference(avatar_split_clause,[],[f139,f359,f844])).
fof(f139,axiom,(
  ~hskp2 | ~c2_1(a538)),
  file('SYN510-1.tptp',unknown)).
fof(f842,plain,(
  ~spl0_124 | ~spl0_31),
  inference(avatar_split_clause,[],[f138,f359,f839])).
fof(f138,axiom,(
  ~hskp2 | ~c1_1(a538)),
  file('SYN510-1.tptp',unknown)).
fof(f837,plain,(
  ~spl0_123 | ~spl0_30),
  inference(avatar_split_clause,[],[f137,f354,f834])).
fof(f137,axiom,(
  ~hskp3 | ~c3_1(a539)),
  file('SYN510-1.tptp',unknown)).
fof(f832,plain,(
  ~spl0_122 | ~spl0_30),
  inference(avatar_split_clause,[],[f136,f354,f829])).
fof(f136,axiom,(
  ~hskp3 | ~c2_1(a539)),
  file('SYN510-1.tptp',unknown)).
fof(f827,plain,(
  ~spl0_121 | ~spl0_29),
  inference(avatar_split_clause,[],[f135,f349,f824])).
fof(f135,axiom,(
  ~hskp4 | ~c2_1(a540)),
  file('SYN510-1.tptp',unknown)).
fof(f822,plain,(
  ~spl0_120 | ~spl0_29),
  inference(avatar_split_clause,[],[f134,f349,f819])).
fof(f134,axiom,(
  ~hskp4 | ~c1_1(a540)),
  file('SYN510-1.tptp',unknown)).
fof(f812,plain,(
  ~spl0_118 | ~spl0_27),
  inference(avatar_split_clause,[],[f132,f339,f809])).
fof(f132,axiom,(
  ~hskp6 | ~c1_1(a546)),
  file('SYN510-1.tptp',unknown)).
fof(f807,plain,(
  ~spl0_117 | ~spl0_27),
  inference(avatar_split_clause,[],[f131,f339,f804])).
fof(f131,axiom,(
  ~hskp6 | ~c0_1(a546)),
  file('SYN510-1.tptp',unknown)).
fof(f802,plain,(
  ~spl0_116 | ~spl0_26),
  inference(avatar_split_clause,[],[f130,f334,f799])).
fof(f130,axiom,(
  ~hskp7 | ~c3_1(a547)),
  file('SYN510-1.tptp',unknown)).
fof(f797,plain,(
  ~spl0_115 | ~spl0_26),
  inference(avatar_split_clause,[],[f129,f334,f794])).
fof(f129,axiom,(
  ~hskp7 | ~c0_1(a547)),
  file('SYN510-1.tptp',unknown)).
fof(f787,plain,(
  ~spl0_113 | ~spl0_24),
  inference(avatar_split_clause,[],[f127,f324,f784])).
fof(f127,axiom,(
  ~hskp9 | ~c2_1(a549)),
  file('SYN510-1.tptp',unknown)).
fof(f782,plain,(
  ~spl0_112 | ~spl0_24),
  inference(avatar_split_clause,[],[f126,f324,f779])).
fof(f126,axiom,(
  ~hskp9 | ~c1_1(a549)),
  file('SYN510-1.tptp',unknown)).
fof(f777,plain,(
  ~spl0_111 | ~spl0_24),
  inference(avatar_split_clause,[],[f125,f324,f774])).
fof(f125,axiom,(
  ~hskp9 | ~c0_1(a549)),
  file('SYN510-1.tptp',unknown)).
fof(f767,plain,(
  ~spl0_109 | ~spl0_22),
  inference(avatar_split_clause,[],[f123,f314,f764])).
fof(f123,axiom,(
  ~hskp11 | ~c2_1(a553)),
  file('SYN510-1.tptp',unknown)).
fof(f762,plain,(
  ~spl0_108 | ~spl0_21),
  inference(avatar_split_clause,[],[f122,f309,f759])).
fof(f122,axiom,(
  ~hskp12 | ~c1_1(a564)),
  file('SYN510-1.tptp',unknown)).
fof(f757,plain,(
  ~spl0_107 | ~spl0_20),
  inference(avatar_split_clause,[],[f121,f304,f754])).
fof(f121,axiom,(
  ~hskp13 | ~c2_1(a565)),
  file('SYN510-1.tptp',unknown)).
fof(f752,plain,(
  ~spl0_106 | ~spl0_19),
  inference(avatar_split_clause,[],[f120,f299,f749])).
fof(f120,axiom,(
  ~hskp14 | ~c0_1(a566)),
  file('SYN510-1.tptp',unknown)).
fof(f747,plain,(
  ~spl0_105 | ~spl0_18),
  inference(avatar_split_clause,[],[f119,f294,f744])).
fof(f119,axiom,(
  ~hskp15 | ~c2_1(a567)),
  file('SYN510-1.tptp',unknown)).
fof(f742,plain,(
  ~spl0_104 | ~spl0_17),
  inference(avatar_split_clause,[],[f118,f289,f739])).
fof(f118,axiom,(
  ~hskp16 | ~c3_1(a570)),
  file('SYN510-1.tptp',unknown)).
fof(f737,plain,(
  ~spl0_103 | ~spl0_17),
  inference(avatar_split_clause,[],[f117,f289,f734])).
fof(f117,axiom,(
  ~hskp16 | ~c2_1(a570)),
  file('SYN510-1.tptp',unknown)).
fof(f722,plain,(
  ~spl0_100 | ~spl0_15),
  inference(avatar_split_clause,[],[f114,f279,f719])).
fof(f114,axiom,(
  ~hskp18 | ~c3_1(a573)),
  file('SYN510-1.tptp',unknown)).
fof(f717,plain,(
  ~spl0_99 | ~spl0_15),
  inference(avatar_split_clause,[],[f113,f279,f714])).
fof(f113,axiom,(
  ~hskp18 | ~c1_1(a573)),
  file('SYN510-1.tptp',unknown)).
fof(f712,plain,(
  ~spl0_98 | ~spl0_15),
  inference(avatar_split_clause,[],[f112,f279,f709])).
fof(f112,axiom,(
  ~hskp18 | ~c0_1(a573)),
  file('SYN510-1.tptp',unknown)).
fof(f707,plain,(
  ~spl0_97 | ~spl0_14),
  inference(avatar_split_clause,[],[f111,f274,f704])).
fof(f111,axiom,(
  ~hskp19 | ~c3_1(a576)),
  file('SYN510-1.tptp',unknown)).
fof(f702,plain,(
  ~spl0_96 | ~spl0_14),
  inference(avatar_split_clause,[],[f110,f274,f699])).
fof(f110,axiom,(
  ~hskp19 | ~c1_1(a576)),
  file('SYN510-1.tptp',unknown)).
fof(f697,plain,(
  ~spl0_95 | ~spl0_13),
  inference(avatar_split_clause,[],[f109,f269,f694])).
fof(f109,axiom,(
  ~hskp20 | ~c3_1(a580)),
  file('SYN510-1.tptp',unknown)).
fof(f687,plain,(
  ~spl0_93 | ~spl0_12),
  inference(avatar_split_clause,[],[f107,f264,f684])).
fof(f107,axiom,(
  ~hskp21 | ~c1_1(a582)),
  file('SYN510-1.tptp',unknown)).
fof(f682,plain,(
  ~spl0_92 | ~spl0_11),
  inference(avatar_split_clause,[],[f106,f259,f679])).
fof(f106,axiom,(
  ~hskp22 | ~c2_1(a584)),
  file('SYN510-1.tptp',unknown)).
fof(f677,plain,(
  ~spl0_91 | ~spl0_11),
  inference(avatar_split_clause,[],[f105,f259,f674])).
fof(f105,axiom,(
  ~hskp22 | ~c0_1(a584)),
  file('SYN510-1.tptp',unknown)).
fof(f672,plain,(
  ~spl0_90 | ~spl0_10),
  inference(avatar_split_clause,[],[f104,f254,f669])).
fof(f254,plain,(
  spl0_10 <=> hskp23),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f104,axiom,(
  ~hskp23 | ~c3_1(a587)),
  file('SYN510-1.tptp',unknown)).
fof(f667,plain,(
  ~spl0_89 | ~spl0_10),
  inference(avatar_split_clause,[],[f103,f254,f664])).
fof(f103,axiom,(
  ~hskp23 | ~c2_1(a587)),
  file('SYN510-1.tptp',unknown)).
fof(f662,plain,(
  ~spl0_88 | ~spl0_10),
  inference(avatar_split_clause,[],[f102,f254,f659])).
fof(f102,axiom,(
  ~hskp23 | ~c0_1(a587)),
  file('SYN510-1.tptp',unknown)).
fof(f657,plain,(
  ~spl0_87 | ~spl0_9),
  inference(avatar_split_clause,[],[f101,f249,f654])).
fof(f101,axiom,(
  ~hskp24 | ~c2_1(a590)),
  file('SYN510-1.tptp',unknown)).
fof(f652,plain,(
  ~spl0_86 | ~spl0_9),
  inference(avatar_split_clause,[],[f100,f249,f649])).
fof(f100,axiom,(
  ~hskp24 | ~c0_1(a590)),
  file('SYN510-1.tptp',unknown)).
fof(f637,plain,(
  ~spl0_83 | ~spl0_7),
  inference(avatar_split_clause,[],[f97,f239,f634])).
fof(f97,axiom,(
  ~hskp26 | ~c3_1(a604)),
  file('SYN510-1.tptp',unknown)).
fof(f632,plain,(
  ~spl0_82 | ~spl0_6),
  inference(avatar_split_clause,[],[f96,f234,f629])).
fof(f96,axiom,(
  ~hskp27 | ~c1_1(a615)),
  file('SYN510-1.tptp',unknown)).
fof(f627,plain,(
  ~spl0_33 | spl0_81),
  inference(avatar_split_clause,[],[f95,f624,f369])).
fof(f95,axiom,(
  c2_1(a536) | ~hskp0),
  file('SYN510-1.tptp',unknown)).
fof(f622,plain,(
  ~spl0_33 | spl0_80),
  inference(avatar_split_clause,[],[f94,f619,f369])).
fof(f94,axiom,(
  c0_1(a536) | ~hskp0),
  file('SYN510-1.tptp',unknown)).
fof(f617,plain,(
  ~spl0_32 | spl0_79),
  inference(avatar_split_clause,[],[f93,f614,f364])).
fof(f93,axiom,(
  c2_1(a537) | ~hskp1),
  file('SYN510-1.tptp',unknown)).
fof(f612,plain,(
  ~spl0_30 | spl0_78),
  inference(avatar_split_clause,[],[f92,f609,f354])).
fof(f92,axiom,(
  c0_1(a539) | ~hskp3),
  file('SYN510-1.tptp',unknown)).
fof(f607,plain,(
  ~spl0_29 | spl0_77),
  inference(avatar_split_clause,[],[f91,f604,f349])).
fof(f91,axiom,(
  c3_1(a540) | ~hskp4),
  file('SYN510-1.tptp',unknown)).
fof(f592,plain,(
  ~spl0_27 | spl0_74),
  inference(avatar_split_clause,[],[f88,f589,f339])).
fof(f88,axiom,(
  c3_1(a546) | ~hskp6),
  file('SYN510-1.tptp',unknown)).
fof(f587,plain,(
  ~spl0_26 | spl0_73),
  inference(avatar_split_clause,[],[f87,f584,f334])).
fof(f87,axiom,(
  c2_1(a547) | ~hskp7),
  file('SYN510-1.tptp',unknown)).
fof(f572,plain,(
  ~spl0_23 | spl0_70),
  inference(avatar_split_clause,[],[f84,f569,f319])).
fof(f84,axiom,(
  c3_1(a551) | ~hskp10),
  file('SYN510-1.tptp',unknown)).
fof(f567,plain,(
  ~spl0_23 | spl0_69),
  inference(avatar_split_clause,[],[f83,f564,f319])).
fof(f83,axiom,(
  c1_1(a551) | ~hskp10),
  file('SYN510-1.tptp',unknown)).
fof(f562,plain,(
  ~spl0_22 | spl0_68),
  inference(avatar_split_clause,[],[f82,f559,f314])).
fof(f82,axiom,(
  c3_1(a553) | ~hskp11),
  file('SYN510-1.tptp',unknown)).
fof(f557,plain,(
  ~spl0_22 | spl0_67),
  inference(avatar_split_clause,[],[f81,f554,f314])).
fof(f81,axiom,(
  c1_1(a553) | ~hskp11),
  file('SYN510-1.tptp',unknown)).
fof(f552,plain,(
  ~spl0_21 | spl0_66),
  inference(avatar_split_clause,[],[f80,f549,f309])).
fof(f80,axiom,(
  c3_1(a564) | ~hskp12),
  file('SYN510-1.tptp',unknown)).
fof(f547,plain,(
  ~spl0_21 | spl0_65),
  inference(avatar_split_clause,[],[f79,f544,f309])).
fof(f79,axiom,(
  c0_1(a564) | ~hskp12),
  file('SYN510-1.tptp',unknown)).
fof(f542,plain,(
  ~spl0_20 | spl0_64),
  inference(avatar_split_clause,[],[f78,f539,f304])).
fof(f78,axiom,(
  c3_1(a565) | ~hskp13),
  file('SYN510-1.tptp',unknown)).
fof(f537,plain,(
  ~spl0_20 | spl0_63),
  inference(avatar_split_clause,[],[f77,f534,f304])).
fof(f77,axiom,(
  c0_1(a565) | ~hskp13),
  file('SYN510-1.tptp',unknown)).
fof(f532,plain,(
  ~spl0_19 | spl0_62),
  inference(avatar_split_clause,[],[f76,f529,f299])).
fof(f76,axiom,(
  c3_1(a566) | ~hskp14),
  file('SYN510-1.tptp',unknown)).
fof(f527,plain,(
  ~spl0_19 | spl0_61),
  inference(avatar_split_clause,[],[f75,f524,f299])).
fof(f75,axiom,(
  c2_1(a566) | ~hskp14),
  file('SYN510-1.tptp',unknown)).
fof(f522,plain,(
  ~spl0_18 | spl0_60),
  inference(avatar_split_clause,[],[f74,f519,f294])).
fof(f74,axiom,(
  c1_1(a567) | ~hskp15),
  file('SYN510-1.tptp',unknown)).
fof(f517,plain,(
  ~spl0_18 | spl0_59),
  inference(avatar_split_clause,[],[f73,f514,f294])).
fof(f73,axiom,(
  c0_1(a567) | ~hskp15),
  file('SYN510-1.tptp',unknown)).
fof(f512,plain,(
  ~spl0_17 | spl0_58),
  inference(avatar_split_clause,[],[f72,f509,f289])).
fof(f72,axiom,(
  c1_1(a570) | ~hskp16),
  file('SYN510-1.tptp',unknown)).
fof(f502,plain,(
  ~spl0_14 | spl0_56),
  inference(avatar_split_clause,[],[f70,f499,f274])).
fof(f70,axiom,(
  c0_1(a576) | ~hskp19),
  file('SYN510-1.tptp',unknown)).
fof(f497,plain,(
  ~spl0_13 | spl0_55),
  inference(avatar_split_clause,[],[f69,f494,f269])).
fof(f69,axiom,(
  c1_1(a580) | ~hskp20),
  file('SYN510-1.tptp',unknown)).
fof(f492,plain,(
  ~spl0_12 | spl0_54),
  inference(avatar_split_clause,[],[f68,f489,f264])).
fof(f68,axiom,(
  c3_1(a582) | ~hskp21),
  file('SYN510-1.tptp',unknown)).
fof(f487,plain,(
  ~spl0_12 | spl0_53),
  inference(avatar_split_clause,[],[f67,f484,f264])).
fof(f67,axiom,(
  c2_1(a582) | ~hskp21),
  file('SYN510-1.tptp',unknown)).
fof(f482,plain,(
  ~spl0_11 | spl0_52),
  inference(avatar_split_clause,[],[f66,f479,f259])).
fof(f66,axiom,(
  c3_1(a584) | ~hskp22),
  file('SYN510-1.tptp',unknown)).
fof(f477,plain,(
  ~spl0_9 | spl0_51),
  inference(avatar_split_clause,[],[f65,f474,f249])).
fof(f65,axiom,(
  c1_1(a590) | ~hskp24),
  file('SYN510-1.tptp',unknown)).
fof(f467,plain,(
  ~spl0_7 | spl0_49),
  inference(avatar_split_clause,[],[f63,f464,f239])).
fof(f63,axiom,(
  c2_1(a604) | ~hskp26),
  file('SYN510-1.tptp',unknown)).
fof(f462,plain,(
  ~spl0_7 | spl0_48),
  inference(avatar_split_clause,[],[f62,f459,f239])).
fof(f62,axiom,(
  c1_1(a604) | ~hskp26),
  file('SYN510-1.tptp',unknown)).
fof(f457,plain,(
  ~spl0_6 | spl0_47),
  inference(avatar_split_clause,[],[f61,f454,f234])).
fof(f61,axiom,(
  c2_1(a615) | ~hskp27),
  file('SYN510-1.tptp',unknown)).
fof(f432,plain,(
  ~spl0_4 | spl0_42),
  inference(avatar_split_clause,[],[f56,f429,f224])).
fof(f56,axiom,(
  c3_1(a552) | ~hskp29),
  file('SYN510-1.tptp',unknown)).
fof(f427,plain,(
  ~spl0_4 | spl0_41),
  inference(avatar_split_clause,[],[f55,f424,f224])).
fof(f55,axiom,(
  c2_1(a552) | ~hskp29),
  file('SYN510-1.tptp',unknown)).
fof(f412,plain,(
  ~spl0_3 | spl0_38),
  inference(avatar_split_clause,[],[f52,f409,f219])).
fof(f52,axiom,(
  c2_1(a562) | ~hskp30),
  file('SYN510-1.tptp',unknown)).
fof(f402,plain,(
  ~spl0_1 | spl0_36),
  inference(avatar_split_clause,[],[f50,f399,f210])).
fof(f50,axiom,(
  c3_1(a563) | ~hskp31),
  file('SYN510-1.tptp',unknown)).
fof(f397,plain,(
  ~spl0_1 | spl0_35),
  inference(avatar_split_clause,[],[f49,f394,f210])).
fof(f49,axiom,(
  c1_1(a563) | ~hskp31),
  file('SYN510-1.tptp',unknown)).
fof(f392,plain,(
  ~spl0_1 | spl0_34),
  inference(avatar_split_clause,[],[f48,f389,f210])).
fof(f48,axiom,(
  c0_1(a563) | ~hskp31),
  file('SYN510-1.tptp',unknown)).
fof(f387,plain,(
  spl0_12 | spl0_11 | spl0_10),
  inference(avatar_split_clause,[],[f47,f254,f259,f264])).
fof(f47,axiom,(
  hskp23 | hskp22 | hskp21),
  file('SYN510-1.tptp',unknown)).
fof(f386,plain,(
  spl0_9 | spl0_24 | spl0_31),
  inference(avatar_split_clause,[],[f46,f359,f324,f249])).
fof(f46,axiom,(
  hskp2 | hskp9 | hskp24),
  file('SYN510-1.tptp',unknown)).
fof(f385,plain,(
  spl0_7 | spl0_27 | spl0_31),
  inference(avatar_split_clause,[],[f45,f359,f339,f239])).
fof(f45,axiom,(
  hskp2 | hskp6 | hskp26),
  file('SYN510-1.tptp',unknown)).
fof(f384,plain,(
  spl0_7 | spl0_26 | spl0_11),
  inference(avatar_split_clause,[],[f44,f259,f334,f239])).
fof(f44,axiom,(
  hskp22 | hskp7 | hskp26),
  file('SYN510-1.tptp',unknown)).
fof(f383,plain,(
  spl0_30 | spl0_9 | spl0_15),
  inference(avatar_split_clause,[],[f43,f279,f249,f354])).
fof(f43,axiom,(
  hskp18 | hskp24 | hskp3),
  file('SYN510-1.tptp',unknown)).
fof(f381,plain,(
  spl0_20 | spl0_19 | spl0_27),
  inference(avatar_split_clause,[],[f41,f339,f299,f304])).
fof(f41,axiom,(
  hskp6 | hskp14 | hskp13),
  file('SYN510-1.tptp',unknown)).
fof(f380,plain,(
  spl0_20 | spl0_9 | spl0_27),
  inference(avatar_split_clause,[],[f40,f339,f249,f304])).
fof(f40,axiom,(
  hskp6 | hskp24 | hskp13),
  file('SYN510-1.tptp',unknown)).
fof(f379,plain,(
  spl0_21 | spl0_26 | spl0_29),
  inference(avatar_split_clause,[],[f39,f349,f334,f309])).
fof(f39,axiom,(
  hskp4 | hskp7 | hskp12),
  file('SYN510-1.tptp',unknown)).
fof(f374,plain,(
  spl0_1 | spl0_21 | spl0_19),
  inference(avatar_split_clause,[],[f34,f299,f309,f210])).
fof(f34,axiom,(
  hskp14 | hskp12 | hskp31),
  file('SYN510-1.tptp',unknown)).
fof(f352,plain,(
  ~spl0_29 | spl0_2),
  inference(avatar_split_clause,[],[f28,f214,f349])).
fof(f28,axiom,(
  ndr1_0 | ~hskp4),
  file('SYN510-1.tptp',unknown)).
fof(f337,plain,(
  ~spl0_26 | spl0_2),
  inference(avatar_split_clause,[],[f25,f214,f334])).
fof(f25,axiom,(
  ndr1_0 | ~hskp7),
  file('SYN510-1.tptp',unknown)).
fof(f312,plain,(
  ~spl0_21 | spl0_2),
  inference(avatar_split_clause,[],[f20,f214,f309])).
fof(f20,axiom,(
  ndr1_0 | ~hskp12),
  file('SYN510-1.tptp',unknown)).
% SZS output end Proof for SYN510-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7036
% Time elapsed: 0.032 s
% ------------------------------
% ------------------------------
