% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWW410-1
% SZS output start Proof for SWW410-1
fof(f13108,plain,(
  $false),
  inference(subsumption_resolution,[],[f13106,f13078])).
fof(f13078,plain,(
  ~heap(sep(lseg(x2,nil),emp))),
  inference(forward_demodulation,[],[f13077,f2])).
fof(f2,axiom,(
  ( ! [X2,X3] : (sep(lseg(X3,X3),X2) = X2) )),
  file('SWW410-1.tptp',unknown)).
fof(f13077,plain,(
  ~heap(sep(lseg(nil,nil),sep(lseg(x2,nil),emp)))),
  inference(forward_demodulation,[],[f13071,f1])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (sep(X0,sep(X1,X2)) = sep(X1,sep(X0,X2))) )),
  file('SWW410-1.tptp',unknown)).
fof(f13071,plain,(
  ~heap(sep(lseg(x2,nil),sep(lseg(nil,nil),emp)))),
  inference(backward_demodulation,[],[f13069,f18])).
fof(f18,plain,(
  ~heap(sep(lseg(x2,x3),sep(lseg(x3,nil),emp)))),
  inference(forward_demodulation,[],[f17,f1])).
fof(f17,axiom,(
  ~heap(sep(lseg(x3,nil),sep(lseg(x2,x3),emp)))),
  file('SWW410-1.tptp',unknown)).
fof(f13069,plain,(
  nil = x3),
  inference(subsumption_resolution,[],[f13062,f18])).
fof(f13062,plain,(
  nil = x3 | heap(sep(lseg(x2,x3),sep(lseg(x3,nil),emp)))),
  inference(resolution,[],[f13051,f124])).
fof(f124,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~heap(sep(lseg(X5,X6),sep(lseg(X4,X5),sep(lseg(X6,X7),X8)))) | X6 = X7 | heap(sep(lseg(X4,X6),sep(lseg(X6,X7),X8)))) )),
  inference(superposition,[],[f11,f1])).
fof(f11,axiom,(
  ( ! [X6,X4,X2,X5,X3] : (~heap(sep(lseg(X3,X4),sep(lseg(X4,X5),sep(lseg(X5,X6),X2)))) | X5 = X6 | heap(sep(lseg(X3,X5),sep(lseg(X5,X6),X2)))) )),
  file('SWW410-1.tptp',unknown)).
fof(f13051,plain,(
  heap(sep(lseg(x1,x3),sep(lseg(x2,x1),sep(lseg(x3,nil),emp))))),
  inference(subsumption_resolution,[],[f12956,f14])).
fof(f14,axiom,(
  x1 != x3),
  file('SWW410-1.tptp',unknown)).
fof(f12956,plain,(
  x1 = x3 | heap(sep(lseg(x1,x3),sep(lseg(x2,x1),sep(lseg(x3,nil),emp))))),
  inference(resolution,[],[f310,f16])).
fof(f16,axiom,(
  heap(sep(lseg(x2,x1),sep(lseg(x3,nil),sep(next(x1,x3),emp))))),
  file('SWW410-1.tptp',unknown)).
fof(f310,plain,(
  ( ! [X14,X17,X15,X13,X16] : (~heap(sep(X15,sep(X16,sep(next(X13,X14),X17)))) | X13 = X14 | heap(sep(lseg(X13,X14),sep(X15,sep(X16,X17))))) )),
  inference(superposition,[],[f88,f21])).
fof(f21,plain,(
  ( ! [X6,X4,X5,X3] : (sep(X3,sep(X6,sep(X4,X5))) = sep(X6,sep(X4,sep(X3,X5)))) )),
  inference(superposition,[],[f1,f1])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~heap(sep(next(X2,X0),X1)) | X0 = X2 | heap(sep(lseg(X2,X0),X1))) )),
  inference(superposition,[],[f8,f2])).
fof(f8,axiom,(
  ( ! [X4,X2,X5,X3] : (~heap(sep(next(X3,X4),sep(lseg(X4,X5),X2))) | X3 = X4 | heap(sep(lseg(X3,X5),X2))) )),
  file('SWW410-1.tptp',unknown)).
fof(f13106,plain,(
  heap(sep(lseg(x2,nil),emp))),
  inference(resolution,[],[f13076,f79])).
fof(f79,plain,(
  ( ! [X4,X2,X3] : (~heap(sep(lseg(X3,nil),sep(lseg(X2,X3),X4))) | heap(sep(lseg(X2,nil),X4))) )),
  inference(superposition,[],[f9,f1])).
fof(f9,axiom,(
  ( ! [X4,X2,X3] : (~heap(sep(lseg(X3,X4),sep(lseg(X4,nil),X2))) | heap(sep(lseg(X3,nil),X2))) )),
  file('SWW410-1.tptp',unknown)).
fof(f13076,plain,(
  heap(sep(lseg(x1,nil),sep(lseg(x2,x1),emp)))),
  inference(forward_demodulation,[],[f13075,f2])).
fof(f13075,plain,(
  heap(sep(lseg(nil,nil),sep(lseg(x1,nil),sep(lseg(x2,x1),emp))))),
  inference(forward_demodulation,[],[f13070,f21])).
fof(f13070,plain,(
  heap(sep(lseg(x1,nil),sep(lseg(x2,x1),sep(lseg(nil,nil),emp))))),
  inference(backward_demodulation,[],[f13069,f13051])).
% SZS output end Proof for SWW410-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 478201
% Time elapsed: 5.986 s
% ------------------------------
% ------------------------------
