% Refutation found. Thanks to Tanya!
% SZS status Theorem for SEU385+1
% SZS output start Proof for SEU385+1
fof(f89867,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f437,f600,f610,f1080,f1477,f1528,f2443,f8195,f89851])).
fof(f89851,plain,(
  ~spl28_4 | ~spl28_7 | ~spl28_12 | ~spl28_16 | ~spl28_18),
  inference(avatar_contradiction_clause,[],[f89850])).
fof(f89850,plain,(
  $false | (~spl28_4 | ~spl28_7 | ~spl28_12 | ~spl28_16 | ~spl28_18)),
  inference(subsumption_resolution,[],[f89849,f361])).
fof(f361,plain,(
  apply_netmap(sK23,sK24,sK26) != apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26)),
  inference(backward_demodulation,[],[f346,f347])).
fof(f347,plain,(
  apply_netmap(sK23,sK24,sK26) != apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK27)),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ((((apply_netmap(sK23,sK24,sK26) != apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK27) & sK26 = sK27 & element(sK27,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))) & element(sK26,the_carrier(sK24))) & element(sK25,the_carrier(sK24))) & net_str(sK24,sK23) & directed_relstr(sK24) & ~empty_carrier(sK24)) & one_sorted_str(sK23) & ~empty_carrier(sK23)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24,sK25,sK26,sK27])],[f156,f222,f221,f220,f219,f218])).
fof(f218,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (apply_netmap(X0,X1,X3) != apply_netmap(X0,netstr_restr_to_element(X0,X1,X2),X4) & X3 = X4 & element(X4,the_carrier(netstr_restr_to_element(X0,X1,X2)))) & element(X3,the_carrier(X1))) & element(X2,the_carrier(X1))) & net_str(X1,X0) & directed_relstr(X1) & ~empty_carrier(X1)) & one_sorted_str(X0) & ~empty_carrier(X0)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (apply_netmap(sK23,X1,X3) != apply_netmap(sK23,netstr_restr_to_element(sK23,X1,X2),X4) & X3 = X4 & element(X4,the_carrier(netstr_restr_to_element(sK23,X1,X2)))) & element(X3,the_carrier(X1))) & element(X2,the_carrier(X1))) & net_str(X1,sK23) & directed_relstr(X1) & ~empty_carrier(X1)) & one_sorted_str(sK23) & ~empty_carrier(sK23))),
  introduced(choice_axiom,[])).
fof(f219,plain,(
  ? [X1] : (? [X2] : (? [X3] : (? [X4] : (apply_netmap(sK23,X1,X3) != apply_netmap(sK23,netstr_restr_to_element(sK23,X1,X2),X4) & X3 = X4 & element(X4,the_carrier(netstr_restr_to_element(sK23,X1,X2)))) & element(X3,the_carrier(X1))) & element(X2,the_carrier(X1))) & net_str(X1,sK23) & directed_relstr(X1) & ~empty_carrier(X1)) => (? [X2] : (? [X3] : (? [X4] : (apply_netmap(sK23,sK24,X3) != apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,X2),X4) & X3 = X4 & element(X4,the_carrier(netstr_restr_to_element(sK23,sK24,X2)))) & element(X3,the_carrier(sK24))) & element(X2,the_carrier(sK24))) & net_str(sK24,sK23) & directed_relstr(sK24) & ~empty_carrier(sK24))),
  introduced(choice_axiom,[])).
fof(f220,plain,(
  ? [X2] : (? [X3] : (? [X4] : (apply_netmap(sK23,sK24,X3) != apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,X2),X4) & X3 = X4 & element(X4,the_carrier(netstr_restr_to_element(sK23,sK24,X2)))) & element(X3,the_carrier(sK24))) & element(X2,the_carrier(sK24))) => (? [X3] : (? [X4] : (apply_netmap(sK23,sK24,X3) != apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),X4) & X3 = X4 & element(X4,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))) & element(X3,the_carrier(sK24))) & element(sK25,the_carrier(sK24)))),
  introduced(choice_axiom,[])).
fof(f221,plain,(
  ? [X3] : (? [X4] : (apply_netmap(sK23,sK24,X3) != apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),X4) & X3 = X4 & element(X4,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))) & element(X3,the_carrier(sK24))) => (? [X4] : (apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),X4) != apply_netmap(sK23,sK24,sK26) & sK26 = X4 & element(X4,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))) & element(sK26,the_carrier(sK24)))),
  introduced(choice_axiom,[])).
fof(f222,plain,(
  ? [X4] : (apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),X4) != apply_netmap(sK23,sK24,sK26) & sK26 = X4 & element(X4,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))) => (apply_netmap(sK23,sK24,sK26) != apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK27) & sK26 = sK27 & element(sK27,the_carrier(netstr_restr_to_element(sK23,sK24,sK25))))),
  introduced(choice_axiom,[])).
fof(f156,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (apply_netmap(X0,X1,X3) != apply_netmap(X0,netstr_restr_to_element(X0,X1,X2),X4) & X3 = X4 & element(X4,the_carrier(netstr_restr_to_element(X0,X1,X2)))) & element(X3,the_carrier(X1))) & element(X2,the_carrier(X1))) & net_str(X1,X0) & directed_relstr(X1) & ~empty_carrier(X1)) & one_sorted_str(X0) & ~empty_carrier(X0))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((apply_netmap(X0,X1,X3) != apply_netmap(X0,netstr_restr_to_element(X0,X1,X2),X4) & X3 = X4) & element(X4,the_carrier(netstr_restr_to_element(X0,X1,X2)))) & element(X3,the_carrier(X1))) & element(X2,the_carrier(X1))) & (net_str(X1,X0) & directed_relstr(X1) & ~empty_carrier(X1))) & (one_sorted_str(X0) & ~empty_carrier(X0)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,negated_conjecture,(
  ~! [X0] : ((one_sorted_str(X0) & ~empty_carrier(X0)) => ! [X1] : ((net_str(X1,X0) & directed_relstr(X1) & ~empty_carrier(X1)) => ! [X2] : (element(X2,the_carrier(X1)) => ! [X3] : (element(X3,the_carrier(X1)) => ! [X4] : (element(X4,the_carrier(netstr_restr_to_element(X0,X1,X2))) => (X3 = X4 => apply_netmap(X0,X1,X3) = apply_netmap(X0,netstr_restr_to_element(X0,X1,X2),X4)))))))),
  inference(negated_conjecture,[],[f70])).
fof(f70,conjecture,(
  ! [X0] : ((one_sorted_str(X0) & ~empty_carrier(X0)) => ! [X1] : ((net_str(X1,X0) & directed_relstr(X1) & ~empty_carrier(X1)) => ! [X2] : (element(X2,the_carrier(X1)) => ! [X3] : (element(X3,the_carrier(X1)) => ! [X4] : (element(X4,the_carrier(netstr_restr_to_element(X0,X1,X2))) => (X3 = X4 => apply_netmap(X0,X1,X3) = apply_netmap(X0,netstr_restr_to_element(X0,X1,X2),X4)))))))),
  file('SEU385+1.tptp',unknown)).
fof(f346,plain,(
  sK26 = sK27),
  inference(cnf_transformation,[],[f223])).
fof(f89849,plain,(
  apply_netmap(sK23,sK24,sK26) = apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) | (~spl28_4 | ~spl28_7 | ~spl28_12 | ~spl28_16 | ~spl28_18)),
  inference(backward_demodulation,[],[f89848,f3437])).
fof(f3437,plain,(
  apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | (~spl28_7 | ~spl28_12)),
  inference(forward_demodulation,[],[f3436,f1130])).
fof(f1130,plain,(
  apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~spl28_7),
  inference(subsumption_resolution,[],[f1129,f340])).
fof(f340,plain,(
  ~empty_carrier(sK24)),
  inference(cnf_transformation,[],[f223])).
fof(f1129,plain,(
  apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | empty_carrier(sK24) | ~spl28_7),
  inference(subsumption_resolution,[],[f1128,f342])).
fof(f342,plain,(
  net_str(sK24,sK23)),
  inference(cnf_transformation,[],[f223])).
fof(f1128,plain,(
  apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~spl28_7),
  inference(subsumption_resolution,[],[f1127,f343])).
fof(f343,plain,(
  element(sK25,the_carrier(sK24))),
  inference(cnf_transformation,[],[f223])).
fof(f1127,plain,(
  apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~spl28_7),
  inference(subsumption_resolution,[],[f1126,f339])).
fof(f339,plain,(
  one_sorted_str(sK23)),
  inference(cnf_transformation,[],[f223])).
fof(f1126,plain,(
  ~one_sorted_str(sK23) | apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~spl28_7),
  inference(subsumption_resolution,[],[f1125,f338])).
fof(f338,plain,(
  ~empty_carrier(sK23)),
  inference(cnf_transformation,[],[f223])).
fof(f1125,plain,(
  empty_carrier(sK23) | ~one_sorted_str(sK23) | apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~spl28_7),
  inference(duplicate_literal_removal,[],[f1124])).
fof(f1124,plain,(
  empty_carrier(sK23) | ~one_sorted_str(sK23) | apply_netmap(sK23,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~one_sorted_str(sK23) | empty_carrier(sK23) | ~spl28_7),
  inference(resolution,[],[f599,f267])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (net_str(netstr_restr_to_element(X0,X1,X2),X0) | ~element(X2,the_carrier(X1)) | ~net_str(X1,X0) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : ((net_str(netstr_restr_to_element(X0,X1,X2),X0) & strict_net_str(netstr_restr_to_element(X0,X1,X2),X0)) | ~element(X2,the_carrier(X1)) | ~net_str(X1,X0) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((net_str(netstr_restr_to_element(X0,X1,X2),X0) & strict_net_str(netstr_restr_to_element(X0,X1,X2),X0)) | (~element(X2,the_carrier(X1)) | ~net_str(X1,X0) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2] : ((element(X2,the_carrier(X1)) & net_str(X1,X0) & ~empty_carrier(X1) & one_sorted_str(X0) & ~empty_carrier(X0)) => (net_str(netstr_restr_to_element(X0,X1,X2),X0) & strict_net_str(netstr_restr_to_element(X0,X1,X2),X0)))),
  file('SEU385+1.tptp',unknown)).
fof(f599,plain,(
  ( ! [X0] : (~net_str(netstr_restr_to_element(sK23,sK24,sK25),X0) | empty_carrier(X0) | ~one_sorted_str(X0) | apply_netmap(X0,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),X0,the_mapping(X0,netstr_restr_to_element(sK23,sK24,sK25)),sK26)) ) | ~spl28_7),
  inference(avatar_component_clause,[],[f598])).
fof(f598,plain,(
  spl28_7 <=> ! [X0] : (apply_netmap(X0,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),X0,the_mapping(X0,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | empty_carrier(X0) | ~one_sorted_str(X0) | ~net_str(netstr_restr_to_element(sK23,sK24,sK25),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_7])])).
fof(f3436,plain,(
  apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~spl28_12),
  inference(subsumption_resolution,[],[f3435,f340])).
fof(f3435,plain,(
  apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | empty_carrier(sK24) | ~spl28_12),
  inference(subsumption_resolution,[],[f3434,f342])).
fof(f3434,plain,(
  apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~spl28_12),
  inference(subsumption_resolution,[],[f3433,f343])).
fof(f3433,plain,(
  apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~spl28_12),
  inference(subsumption_resolution,[],[f3432,f339])).
fof(f3432,plain,(
  apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~one_sorted_str(sK23) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~spl28_12),
  inference(subsumption_resolution,[],[f3431,f338])).
fof(f3431,plain,(
  apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | empty_carrier(sK23) | ~one_sorted_str(sK23) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~spl28_12),
  inference(duplicate_literal_removal,[],[f3430])).
fof(f3430,plain,(
  apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),sK23,the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | empty_carrier(sK23) | ~one_sorted_str(sK23) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | empty_carrier(sK24) | ~one_sorted_str(sK23) | empty_carrier(sK23) | ~spl28_12),
  inference(resolution,[],[f1079,f267])).
fof(f1079,plain,(
  ( ! [X15] : (~net_str(netstr_restr_to_element(sK23,sK24,sK25),X15) | apply(the_mapping(X15,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),X15,the_mapping(X15,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | empty_carrier(X15) | ~one_sorted_str(X15)) ) | ~spl28_12),
  inference(avatar_component_clause,[],[f1078])).
fof(f1078,plain,(
  spl28_12 <=> ! [X15] : (apply(the_mapping(X15,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),X15,the_mapping(X15,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~net_str(netstr_restr_to_element(sK23,sK24,sK25),X15) | empty_carrier(X15) | ~one_sorted_str(X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_12])])).
fof(f89848,plain,(
  apply_netmap(sK23,sK24,sK26) = apply(the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | (~spl28_4 | ~spl28_16 | ~spl28_18)),
  inference(backward_demodulation,[],[f89827,f8342])).
fof(f8342,plain,(
  apply_netmap(sK23,sK24,sK26) = apply(relation_dom_restriction(the_mapping(sK23,sK24),the_carrier(netstr_restr_to_element(sK23,sK24,sK25))),sK26) | (~spl28_4 | ~spl28_16 | ~spl28_18)),
  inference(forward_demodulation,[],[f8341,f2476])).
fof(f2476,plain,(
  apply_netmap(sK23,sK24,sK26) = apply(the_mapping(sK23,sK24),sK26)),
  inference(forward_demodulation,[],[f2475,f854])).
fof(f854,plain,(
  apply_netmap(sK23,sK24,sK26) = apply_on_structs(sK24,sK23,the_mapping(sK23,sK24),sK26)),
  inference(subsumption_resolution,[],[f853,f338])).
fof(f853,plain,(
  apply_netmap(sK23,sK24,sK26) = apply_on_structs(sK24,sK23,the_mapping(sK23,sK24),sK26) | empty_carrier(sK23)),
  inference(subsumption_resolution,[],[f852,f339])).
fof(f852,plain,(
  apply_netmap(sK23,sK24,sK26) = apply_on_structs(sK24,sK23,the_mapping(sK23,sK24),sK26) | ~one_sorted_str(sK23) | empty_carrier(sK23)),
  inference(resolution,[],[f601,f342])).
fof(f601,plain,(
  ( ! [X1] : (~net_str(sK24,X1) | apply_netmap(X1,sK24,sK26) = apply_on_structs(sK24,X1,the_mapping(X1,sK24),sK26) | ~one_sorted_str(X1) | empty_carrier(X1)) )),
  inference(subsumption_resolution,[],[f588,f340])).
fof(f588,plain,(
  ( ! [X1] : (apply_netmap(X1,sK24,sK26) = apply_on_structs(sK24,X1,the_mapping(X1,sK24),sK26) | ~net_str(sK24,X1) | empty_carrier(sK24) | ~one_sorted_str(X1) | empty_carrier(X1)) )),
  inference(resolution,[],[f257,f344])).
fof(f344,plain,(
  element(sK26,the_carrier(sK24))),
  inference(cnf_transformation,[],[f223])).
fof(f257,plain,(
  ( ! [X2,X0,X1] : (~element(X2,the_carrier(X1)) | apply_netmap(X0,X1,X2) = apply_on_structs(X1,X0,the_mapping(X0,X1),X2) | ~net_str(X1,X0) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : (! [X1] : (! [X2] : (apply_netmap(X0,X1,X2) = apply_on_structs(X1,X0,the_mapping(X0,X1),X2) | ~element(X2,the_carrier(X1))) | ~net_str(X1,X0) | empty_carrier(X1)) | ~one_sorted_str(X0) | empty_carrier(X0))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : (! [X1] : (! [X2] : (apply_netmap(X0,X1,X2) = apply_on_structs(X1,X0,the_mapping(X0,X1),X2) | ~element(X2,the_carrier(X1))) | (~net_str(X1,X0) | empty_carrier(X1))) | (~one_sorted_str(X0) | empty_carrier(X0)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0] : ((one_sorted_str(X0) & ~empty_carrier(X0)) => ! [X1] : ((net_str(X1,X0) & ~empty_carrier(X1)) => ! [X2] : (element(X2,the_carrier(X1)) => apply_netmap(X0,X1,X2) = apply_on_structs(X1,X0,the_mapping(X0,X1),X2))))),
  file('SEU385+1.tptp',unknown)).
fof(f2475,plain,(
  apply_on_structs(sK24,sK23,the_mapping(sK23,sK24),sK26) = apply(the_mapping(sK23,sK24),sK26)),
  inference(subsumption_resolution,[],[f2474,f338])).
fof(f2474,plain,(
  empty_carrier(sK23) | apply_on_structs(sK24,sK23,the_mapping(sK23,sK24),sK26) = apply(the_mapping(sK23,sK24),sK26)),
  inference(subsumption_resolution,[],[f2473,f339])).
fof(f2473,plain,(
  ~one_sorted_str(sK23) | empty_carrier(sK23) | apply_on_structs(sK24,sK23,the_mapping(sK23,sK24),sK26) = apply(the_mapping(sK23,sK24),sK26)),
  inference(resolution,[],[f1082,f342])).
fof(f1082,plain,(
  ( ! [X16] : (~net_str(sK24,X16) | ~one_sorted_str(X16) | empty_carrier(X16) | apply(the_mapping(X16,sK24),sK26) = apply_on_structs(sK24,X16,the_mapping(X16,sK24),sK26)) )),
  inference(subsumption_resolution,[],[f1081,f340])).
fof(f1081,plain,(
  ( ! [X16] : (apply(the_mapping(X16,sK24),sK26) = apply_on_structs(sK24,X16,the_mapping(X16,sK24),sK26) | ~one_sorted_str(X16) | empty_carrier(X16) | empty_carrier(sK24) | ~net_str(sK24,X16)) )),
  inference(subsumption_resolution,[],[f1059,f389])).
fof(f389,plain,(
  one_sorted_str(sK24)),
  inference(resolution,[],[f388,f270])).
fof(f270,plain,(
  ( ! [X0] : (~rel_str(X0) | one_sorted_str(X0)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : (one_sorted_str(X0) | ~rel_str(X0))),
  inference(ennf_transformation,[],[f28])).
fof(f28,axiom,(
  ! [X0] : (rel_str(X0) => one_sorted_str(X0))),
  file('SEU385+1.tptp',unknown)).
fof(f388,plain,(
  rel_str(sK24)),
  inference(subsumption_resolution,[],[f383,f339])).
fof(f383,plain,(
  rel_str(sK24) | ~one_sorted_str(sK23)),
  inference(resolution,[],[f271,f342])).
fof(f271,plain,(
  ( ! [X0,X1] : (~net_str(X1,X0) | rel_str(X1) | ~one_sorted_str(X0)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : (! [X1] : (rel_str(X1) | ~net_str(X1,X0)) | ~one_sorted_str(X0))),
  inference(ennf_transformation,[],[f30])).
fof(f30,axiom,(
  ! [X0] : (one_sorted_str(X0) => ! [X1] : (net_str(X1,X0) => rel_str(X1)))),
  file('SEU385+1.tptp',unknown)).
fof(f1059,plain,(
  ( ! [X16] : (apply(the_mapping(X16,sK24),sK26) = apply_on_structs(sK24,X16,the_mapping(X16,sK24),sK26) | ~one_sorted_str(X16) | empty_carrier(X16) | ~one_sorted_str(sK24) | empty_carrier(sK24) | ~net_str(sK24,X16)) )),
  inference(resolution,[],[f659,f344])).
fof(f659,plain,(
  ( ! [X4,X5,X3] : (~element(X3,the_carrier(X4)) | apply(the_mapping(X5,X4),X3) = apply_on_structs(X4,X5,the_mapping(X5,X4),X3) | ~one_sorted_str(X5) | empty_carrier(X5) | ~one_sorted_str(X4) | empty_carrier(X4) | ~net_str(X4,X5)) )),
  inference(subsumption_resolution,[],[f658,f274])).
fof(f274,plain,(
  ( ! [X0,X1] : (function(the_mapping(X0,X1)) | ~net_str(X1,X0) | ~one_sorted_str(X0)) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ((relation_of2_as_subset(the_mapping(X0,X1),the_carrier(X1),the_carrier(X0)) & quasi_total(the_mapping(X0,X1),the_carrier(X1),the_carrier(X0)) & function(the_mapping(X0,X1))) | ~net_str(X1,X0) | ~one_sorted_str(X0))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((relation_of2_as_subset(the_mapping(X0,X1),the_carrier(X1),the_carrier(X0)) & quasi_total(the_mapping(X0,X1),the_carrier(X1),the_carrier(X0)) & function(the_mapping(X0,X1))) | (~net_str(X1,X0) | ~one_sorted_str(X0)))),
  inference(ennf_transformation,[],[f36])).
fof(f36,axiom,(
  ! [X0,X1] : ((net_str(X1,X0) & one_sorted_str(X0)) => (relation_of2_as_subset(the_mapping(X0,X1),the_carrier(X1),the_carrier(X0)) & quasi_total(the_mapping(X0,X1),the_carrier(X1),the_carrier(X0)) & function(the_mapping(X0,X1))))),
  file('SEU385+1.tptp',unknown)).
fof(f658,plain,(
  ( ! [X4,X5,X3] : (~element(X3,the_carrier(X4)) | apply(the_mapping(X5,X4),X3) = apply_on_structs(X4,X5,the_mapping(X5,X4),X3) | ~function(the_mapping(X5,X4)) | ~one_sorted_str(X5) | empty_carrier(X5) | ~one_sorted_str(X4) | empty_carrier(X4) | ~net_str(X4,X5)) )),
  inference(subsumption_resolution,[],[f655,f505])).
fof(f505,plain,(
  ( ! [X2,X3] : (relation_of2(the_mapping(X3,X2),the_carrier(X2),the_carrier(X3)) | ~one_sorted_str(X3) | ~net_str(X2,X3)) )),
  inference(resolution,[],[f276,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (~relation_of2_as_subset(X2,X0,X1) | relation_of2(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((relation_of2_as_subset(X2,X0,X1) | ~relation_of2(X2,X0,X1)) & (relation_of2(X2,X0,X1) | ~relation_of2_as_subset(X2,X0,X1)))),
  inference(nnf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : (relation_of2_as_subset(X2,X0,X1) <=> relation_of2(X2,X0,X1))),
  file('SEU385+1.tptp',unknown)).
fof(f276,plain,(
  ( ! [X0,X1] : (relation_of2_as_subset(the_mapping(X0,X1),the_carrier(X1),the_carrier(X0)) | ~net_str(X1,X0) | ~one_sorted_str(X0)) )),
  inference(cnf_transformation,[],[f127])).
fof(f655,plain,(
  ( ! [X4,X5,X3] : (~element(X3,the_carrier(X4)) | ~relation_of2(the_mapping(X5,X4),the_carrier(X4),the_carrier(X5)) | apply(the_mapping(X5,X4),X3) = apply_on_structs(X4,X5,the_mapping(X5,X4),X3) | ~function(the_mapping(X5,X4)) | ~one_sorted_str(X5) | empty_carrier(X5) | ~one_sorted_str(X4) | empty_carrier(X4) | ~net_str(X4,X5)) )),
  inference(duplicate_literal_removal,[],[f654])).
fof(f654,plain,(
  ( ! [X4,X5,X3] : (~element(X3,the_carrier(X4)) | ~relation_of2(the_mapping(X5,X4),the_carrier(X4),the_carrier(X5)) | apply(the_mapping(X5,X4),X3) = apply_on_structs(X4,X5,the_mapping(X5,X4),X3) | ~function(the_mapping(X5,X4)) | ~one_sorted_str(X5) | empty_carrier(X5) | ~one_sorted_str(X4) | empty_carrier(X4) | ~net_str(X4,X5) | ~one_sorted_str(X5)) )),
  inference(resolution,[],[f332,f275])).
fof(f275,plain,(
  ( ! [X0,X1] : (quasi_total(the_mapping(X0,X1),the_carrier(X1),the_carrier(X0)) | ~net_str(X1,X0) | ~one_sorted_str(X0)) )),
  inference(cnf_transformation,[],[f127])).
fof(f332,plain,(
  ( ! [X2,X0,X3,X1] : (~quasi_total(X2,the_carrier(X0),the_carrier(X1)) | ~element(X3,the_carrier(X0)) | ~relation_of2(X2,the_carrier(X0),the_carrier(X1)) | apply_on_structs(X0,X1,X2,X3) = apply(X2,X3) | ~function(X2) | ~one_sorted_str(X1) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (apply_on_structs(X0,X1,X2,X3) = apply(X2,X3) | ~element(X3,the_carrier(X0)) | ~relation_of2(X2,the_carrier(X0),the_carrier(X1)) | ~quasi_total(X2,the_carrier(X0),the_carrier(X1)) | ~function(X2) | ~one_sorted_str(X1) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : (apply_on_structs(X0,X1,X2,X3) = apply(X2,X3) | (~element(X3,the_carrier(X0)) | ~relation_of2(X2,the_carrier(X0),the_carrier(X1)) | ~quasi_total(X2,the_carrier(X0),the_carrier(X1)) | ~function(X2) | ~one_sorted_str(X1) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)))),
  inference(ennf_transformation,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3] : ((element(X3,the_carrier(X0)) & relation_of2(X2,the_carrier(X0),the_carrier(X1)) & quasi_total(X2,the_carrier(X0),the_carrier(X1)) & function(X2) & one_sorted_str(X1) & ~empty_carrier(X1) & one_sorted_str(X0) & ~empty_carrier(X0)) => apply_on_structs(X0,X1,X2,X3) = apply(X2,X3))),
  file('SEU385+1.tptp',unknown)).
fof(f8341,plain,(
  apply(the_mapping(sK23,sK24),sK26) = apply(relation_dom_restriction(the_mapping(sK23,sK24),the_carrier(netstr_restr_to_element(sK23,sK24,sK25))),sK26) | (~spl28_4 | ~spl28_16 | ~spl28_18)),
  inference(subsumption_resolution,[],[f8306,f1449])).
fof(f1449,plain,(
  function(the_mapping(sK23,sK24)) | ~spl28_16),
  inference(avatar_component_clause,[],[f1448])).
fof(f1448,plain,(
  spl28_16 <=> function(the_mapping(sK23,sK24))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_16])])).
fof(f8306,plain,(
  ~function(the_mapping(sK23,sK24)) | apply(the_mapping(sK23,sK24),sK26) = apply(relation_dom_restriction(the_mapping(sK23,sK24),the_carrier(netstr_restr_to_element(sK23,sK24,sK25))),sK26) | (~spl28_4 | ~spl28_18)),
  inference(resolution,[],[f8276,f1989])).
fof(f1989,plain,(
  relation(the_mapping(sK23,sK24)) | ~spl28_18),
  inference(resolution,[],[f1908,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (~relation_of2_as_subset(X0,X1,X2) | relation(X0)) )),
  inference(resolution,[],[f272,f232])).
fof(f232,plain,(
  ( ! [X2,X0,X1] : (~element(X2,powerset(cartesian_product2(X0,X1))) | relation(X2)) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : (relation(X2) | ~element(X2,powerset(cartesian_product2(X0,X1))))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2] : (element(X2,powerset(cartesian_product2(X0,X1))) => relation(X2))),
  file('SEU385+1.tptp',unknown)).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (element(X2,powerset(cartesian_product2(X0,X1))) | ~relation_of2_as_subset(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : (element(X2,powerset(cartesian_product2(X0,X1))) | ~relation_of2_as_subset(X2,X0,X1))),
  inference(ennf_transformation,[],[f33])).
fof(f33,axiom,(
  ! [X0,X1,X2] : (relation_of2_as_subset(X2,X0,X1) => element(X2,powerset(cartesian_product2(X0,X1))))),
  file('SEU385+1.tptp',unknown)).
fof(f1908,plain,(
  relation_of2_as_subset(the_mapping(sK23,sK24),the_carrier(sK24),the_carrier(sK23)) | ~spl28_18),
  inference(resolution,[],[f1457,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (~relation_of2(X2,X0,X1) | relation_of2_as_subset(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f217])).
fof(f1457,plain,(
  relation_of2(the_mapping(sK23,sK24),the_carrier(sK24),the_carrier(sK23)) | ~spl28_18),
  inference(avatar_component_clause,[],[f1456])).
fof(f1456,plain,(
  spl28_18 <=> relation_of2(the_mapping(sK23,sK24),the_carrier(sK24),the_carrier(sK23))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_18])])).
fof(f8276,plain,(
  ( ! [X0] : (~relation(X0) | ~function(X0) | apply(X0,sK26) = apply(relation_dom_restriction(X0,the_carrier(netstr_restr_to_element(sK23,sK24,sK25))),sK26)) ) | ~spl28_4),
  inference(resolution,[],[f432,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~in(X1,X0) | apply(relation_dom_restriction(X2,X0),X1) = apply(X2,X1) | ~function(X2) | ~relation(X2)) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : (apply(relation_dom_restriction(X2,X0),X1) = apply(X2,X1) | ~in(X1,X0) | ~function(X2) | ~relation(X2))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((apply(relation_dom_restriction(X2,X0),X1) = apply(X2,X1) | ~in(X1,X0)) | (~function(X2) | ~relation(X2)))),
  inference(ennf_transformation,[],[f78])).
fof(f78,axiom,(
  ! [X0,X1,X2] : ((function(X2) & relation(X2)) => (in(X1,X0) => apply(relation_dom_restriction(X2,X0),X1) = apply(X2,X1)))),
  file('SEU385+1.tptp',unknown)).
fof(f432,plain,(
  in(sK26,the_carrier(netstr_restr_to_element(sK23,sK24,sK25))) | ~spl28_4),
  inference(avatar_component_clause,[],[f430])).
fof(f430,plain,(
  spl28_4 <=> in(sK26,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_4])])).
fof(f89827,plain,(
  the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)) = relation_dom_restriction(the_mapping(sK23,sK24),the_carrier(netstr_restr_to_element(sK23,sK24,sK25))) | (~spl28_16 | ~spl28_18)),
  inference(superposition,[],[f19782,f1909])).
fof(f1909,plain,(
  ( ! [X0] : (partfun_dom_restriction(the_carrier(sK24),the_carrier(sK23),the_mapping(sK23,sK24),X0) = relation_dom_restriction(the_mapping(sK23,sK24),X0)) ) | (~spl28_16 | ~spl28_18)),
  inference(subsumption_resolution,[],[f1907,f1449])).
fof(f1907,plain,(
  ( ! [X0] : (partfun_dom_restriction(the_carrier(sK24),the_carrier(sK23),the_mapping(sK23,sK24),X0) = relation_dom_restriction(the_mapping(sK23,sK24),X0) | ~function(the_mapping(sK23,sK24))) ) | ~spl28_18),
  inference(resolution,[],[f1457,f333])).
fof(f333,plain,(
  ( ! [X2,X0,X3,X1] : (~relation_of2(X2,X0,X1) | partfun_dom_restriction(X0,X1,X2,X3) = relation_dom_restriction(X2,X3) | ~function(X2)) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (partfun_dom_restriction(X0,X1,X2,X3) = relation_dom_restriction(X2,X3) | ~relation_of2(X2,X0,X1) | ~function(X2))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : (partfun_dom_restriction(X0,X1,X2,X3) = relation_dom_restriction(X2,X3) | (~relation_of2(X2,X0,X1) | ~function(X2)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,axiom,(
  ! [X0,X1,X2,X3] : ((relation_of2(X2,X0,X1) & function(X2)) => partfun_dom_restriction(X0,X1,X2,X3) = relation_dom_restriction(X2,X3))),
  file('SEU385+1.tptp',unknown)).
fof(f19782,plain,(
  the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)) = partfun_dom_restriction(the_carrier(sK24),the_carrier(sK23),the_mapping(sK23,sK24),the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))),
  inference(subsumption_resolution,[],[f19781,f339])).
fof(f19781,plain,(
  ~one_sorted_str(sK23) | the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)) = partfun_dom_restriction(the_carrier(sK24),the_carrier(sK23),the_mapping(sK23,sK24),the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))),
  inference(subsumption_resolution,[],[f19780,f338])).
fof(f19780,plain,(
  empty_carrier(sK23) | ~one_sorted_str(sK23) | the_mapping(sK23,netstr_restr_to_element(sK23,sK24,sK25)) = partfun_dom_restriction(the_carrier(sK24),the_carrier(sK23),the_mapping(sK23,sK24),the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))),
  inference(resolution,[],[f3579,f342])).
fof(f3579,plain,(
  ( ! [X19] : (~net_str(sK24,X19) | empty_carrier(X19) | ~one_sorted_str(X19) | the_mapping(X19,netstr_restr_to_element(X19,sK24,sK25)) = partfun_dom_restriction(the_carrier(sK24),the_carrier(X19),the_mapping(X19,sK24),the_carrier(netstr_restr_to_element(X19,sK24,sK25)))) )),
  inference(subsumption_resolution,[],[f3553,f340])).
fof(f3553,plain,(
  ( ! [X19] : (empty_carrier(sK24) | ~one_sorted_str(X19) | empty_carrier(X19) | ~net_str(sK24,X19) | the_mapping(X19,netstr_restr_to_element(X19,sK24,sK25)) = partfun_dom_restriction(the_carrier(sK24),the_carrier(X19),the_mapping(X19,sK24),the_carrier(netstr_restr_to_element(X19,sK24,sK25)))) )),
  inference(resolution,[],[f885,f343])).
fof(f885,plain,(
  ( ! [X2,X0,X1] : (~element(X2,the_carrier(X0)) | empty_carrier(X0) | ~one_sorted_str(X1) | empty_carrier(X1) | ~net_str(X0,X1) | the_mapping(X1,netstr_restr_to_element(X1,X0,X2)) = partfun_dom_restriction(the_carrier(X0),the_carrier(X1),the_mapping(X1,X0),the_carrier(netstr_restr_to_element(X1,X0,X2)))) )),
  inference(resolution,[],[f582,f251])).
fof(f251,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3) | the_mapping(X2,X0) = partfun_dom_restriction(the_carrier(X1),the_carrier(X2),the_mapping(X2,X1),the_carrier(X0))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3] : ((sP0(X0,X1,X2,X3) | the_mapping(X2,X0) != partfun_dom_restriction(the_carrier(X1),the_carrier(X2),the_mapping(X2,X1),the_carrier(X0)) | the_InternalRel(X0) != relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X0)) | ((! [X5] : (~related(X1,X3,X5) | sK2(X0,X1,X3) != X5 | ~element(X5,the_carrier(X1))) | ~in(sK2(X0,X1,X3),the_carrier(X0))) & ((related(X1,X3,sK3(X0,X1,X3)) & sK2(X0,X1,X3) = sK3(X0,X1,X3) & element(sK3(X0,X1,X3),the_carrier(X1))) | in(sK2(X0,X1,X3),the_carrier(X0))))) & ((the_mapping(X2,X0) = partfun_dom_restriction(the_carrier(X1),the_carrier(X2),the_mapping(X2,X1),the_carrier(X0)) & the_InternalRel(X0) = relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X0)) & ! [X7] : ((in(X7,the_carrier(X0)) | ! [X8] : (~related(X1,X3,X8) | X7 != X8 | ~element(X8,the_carrier(X1)))) & ((related(X1,X3,sK4(X1,X3,X7)) & sK4(X1,X3,X7) = X7 & element(sK4(X1,X3,X7),the_carrier(X1))) | ~in(X7,the_carrier(X0))))) | ~sP0(X0,X1,X2,X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3,sK4])],[f176,f179,f178,f177])).
fof(f177,plain,(
  ! [X3,X1,X0] : (? [X4] : ((! [X5] : (~related(X1,X3,X5) | X4 != X5 | ~element(X5,the_carrier(X1))) | ~in(X4,the_carrier(X0))) & (? [X6] : (related(X1,X3,X6) & X4 = X6 & element(X6,the_carrier(X1))) | in(X4,the_carrier(X0)))) => ((! [X5] : (~related(X1,X3,X5) | sK2(X0,X1,X3) != X5 | ~element(X5,the_carrier(X1))) | ~in(sK2(X0,X1,X3),the_carrier(X0))) & (? [X6] : (related(X1,X3,X6) & sK2(X0,X1,X3) = X6 & element(X6,the_carrier(X1))) | in(sK2(X0,X1,X3),the_carrier(X0)))))),
  introduced(choice_axiom,[])).
fof(f178,plain,(
  ! [X3,X1,X0] : (? [X6] : (related(X1,X3,X6) & sK2(X0,X1,X3) = X6 & element(X6,the_carrier(X1))) => (related(X1,X3,sK3(X0,X1,X3)) & sK2(X0,X1,X3) = sK3(X0,X1,X3) & element(sK3(X0,X1,X3),the_carrier(X1))))),
  introduced(choice_axiom,[])).
fof(f179,plain,(
  ! [X7,X3,X1] : (? [X9] : (related(X1,X3,X9) & X7 = X9 & element(X9,the_carrier(X1))) => (related(X1,X3,sK4(X1,X3,X7)) & sK4(X1,X3,X7) = X7 & element(sK4(X1,X3,X7),the_carrier(X1))))),
  introduced(choice_axiom,[])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : ((sP0(X0,X1,X2,X3) | the_mapping(X2,X0) != partfun_dom_restriction(the_carrier(X1),the_carrier(X2),the_mapping(X2,X1),the_carrier(X0)) | the_InternalRel(X0) != relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X0)) | ? [X4] : ((! [X5] : (~related(X1,X3,X5) | X4 != X5 | ~element(X5,the_carrier(X1))) | ~in(X4,the_carrier(X0))) & (? [X6] : (related(X1,X3,X6) & X4 = X6 & element(X6,the_carrier(X1))) | in(X4,the_carrier(X0))))) & ((the_mapping(X2,X0) = partfun_dom_restriction(the_carrier(X1),the_carrier(X2),the_mapping(X2,X1),the_carrier(X0)) & the_InternalRel(X0) = relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X0)) & ! [X7] : ((in(X7,the_carrier(X0)) | ! [X8] : (~related(X1,X3,X8) | X7 != X8 | ~element(X8,the_carrier(X1)))) & (? [X9] : (related(X1,X3,X9) & X7 = X9 & element(X9,the_carrier(X1))) | ~in(X7,the_carrier(X0))))) | ~sP0(X0,X1,X2,X3)))),
  inference(rectify,[],[f175])).
fof(f175,plain,(
  ! [X3,X1,X0,X2] : ((sP0(X3,X1,X0,X2) | the_mapping(X0,X3) != partfun_dom_restriction(the_carrier(X1),the_carrier(X0),the_mapping(X0,X1),the_carrier(X3)) | the_InternalRel(X3) != relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X3)) | ? [X4] : ((! [X5] : (~related(X1,X2,X5) | X4 != X5 | ~element(X5,the_carrier(X1))) | ~in(X4,the_carrier(X3))) & (? [X5] : (related(X1,X2,X5) & X4 = X5 & element(X5,the_carrier(X1))) | in(X4,the_carrier(X3))))) & ((the_mapping(X0,X3) = partfun_dom_restriction(the_carrier(X1),the_carrier(X0),the_mapping(X0,X1),the_carrier(X3)) & the_InternalRel(X3) = relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X3)) & ! [X4] : ((in(X4,the_carrier(X3)) | ! [X5] : (~related(X1,X2,X5) | X4 != X5 | ~element(X5,the_carrier(X1)))) & (? [X5] : (related(X1,X2,X5) & X4 = X5 & element(X5,the_carrier(X1))) | ~in(X4,the_carrier(X3))))) | ~sP0(X3,X1,X0,X2)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X3,X1,X0,X2] : ((sP0(X3,X1,X0,X2) | (the_mapping(X0,X3) != partfun_dom_restriction(the_carrier(X1),the_carrier(X0),the_mapping(X0,X1),the_carrier(X3)) | the_InternalRel(X3) != relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X3)) | ? [X4] : ((! [X5] : (~related(X1,X2,X5) | X4 != X5 | ~element(X5,the_carrier(X1))) | ~in(X4,the_carrier(X3))) & (? [X5] : (related(X1,X2,X5) & X4 = X5 & element(X5,the_carrier(X1))) | in(X4,the_carrier(X3)))))) & ((the_mapping(X0,X3) = partfun_dom_restriction(the_carrier(X1),the_carrier(X0),the_mapping(X0,X1),the_carrier(X3)) & the_InternalRel(X3) = relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X3)) & ! [X4] : ((in(X4,the_carrier(X3)) | ! [X5] : (~related(X1,X2,X5) | X4 != X5 | ~element(X5,the_carrier(X1)))) & (? [X5] : (related(X1,X2,X5) & X4 = X5 & element(X5,the_carrier(X1))) | ~in(X4,the_carrier(X3))))) | ~sP0(X3,X1,X0,X2)))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X3,X1,X0,X2] : (sP0(X3,X1,X0,X2) <=> (the_mapping(X0,X3) = partfun_dom_restriction(the_carrier(X1),the_carrier(X0),the_mapping(X0,X1),the_carrier(X3)) & the_InternalRel(X3) = relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X3)) & ! [X4] : (in(X4,the_carrier(X3)) <=> ? [X5] : (related(X1,X2,X5) & X4 = X5 & element(X5,the_carrier(X1))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f582,plain,(
  ( ! [X6,X4,X5] : (sP0(netstr_restr_to_element(X4,X5,X6),X5,X4,X6) | ~net_str(X5,X4) | empty_carrier(X5) | ~one_sorted_str(X4) | empty_carrier(X4) | ~element(X6,the_carrier(X5))) )),
  inference(subsumption_resolution,[],[f581,f266])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (strict_net_str(netstr_restr_to_element(X0,X1,X2),X0) | ~element(X2,the_carrier(X1)) | ~net_str(X1,X0) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)) )),
  inference(cnf_transformation,[],[f118])).
fof(f581,plain,(
  ( ! [X6,X4,X5] : (~strict_net_str(netstr_restr_to_element(X4,X5,X6),X4) | ~element(X6,the_carrier(X5)) | ~net_str(X5,X4) | empty_carrier(X5) | ~one_sorted_str(X4) | empty_carrier(X4) | sP0(netstr_restr_to_element(X4,X5,X6),X5,X4,X6)) )),
  inference(subsumption_resolution,[],[f580,f267])).
fof(f580,plain,(
  ( ! [X6,X4,X5] : (~net_str(netstr_restr_to_element(X4,X5,X6),X4) | ~strict_net_str(netstr_restr_to_element(X4,X5,X6),X4) | ~element(X6,the_carrier(X5)) | ~net_str(X5,X4) | empty_carrier(X5) | ~one_sorted_str(X4) | empty_carrier(X4) | sP0(netstr_restr_to_element(X4,X5,X6),X5,X4,X6)) )),
  inference(resolution,[],[f256,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,netstr_restr_to_element(X1,X2,X0)) | sP0(netstr_restr_to_element(X1,X2,X0),X2,X1,X0)) )),
  inference(equality_resolution,[],[f244])).
fof(f244,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X3,X2,X1,X0) | netstr_restr_to_element(X1,X2,X0) != X3 | ~sP1(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (((netstr_restr_to_element(X1,X2,X0) = X3 | ~sP0(X3,X2,X1,X0)) & (sP0(X3,X2,X1,X0) | netstr_restr_to_element(X1,X2,X0) != X3)) | ~sP1(X0,X1,X2,X3))),
  inference(rectify,[],[f172])).
fof(f172,plain,(
  ! [X2,X0,X1,X3] : (((netstr_restr_to_element(X0,X1,X2) = X3 | ~sP0(X3,X1,X0,X2)) & (sP0(X3,X1,X0,X2) | netstr_restr_to_element(X0,X1,X2) != X3)) | ~sP1(X2,X0,X1,X3))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X2,X0,X1,X3] : ((netstr_restr_to_element(X0,X1,X2) = X3 <=> sP0(X3,X1,X0,X2)) | ~sP1(X2,X0,X1,X3))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f256,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X2,X0,X1,X3) | ~net_str(X3,X0) | ~strict_net_str(X3,X0) | ~element(X2,the_carrier(X1)) | ~net_str(X1,X0) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (sP1(X2,X0,X1,X3) | ~net_str(X3,X0) | ~strict_net_str(X3,X0)) | ~element(X2,the_carrier(X1))) | ~net_str(X1,X0) | empty_carrier(X1)) | ~one_sorted_str(X0) | empty_carrier(X0))),
  inference(definition_folding,[],[f104,f170,f169])).
fof(f104,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((netstr_restr_to_element(X0,X1,X2) = X3 <=> (the_mapping(X0,X3) = partfun_dom_restriction(the_carrier(X1),the_carrier(X0),the_mapping(X0,X1),the_carrier(X3)) & the_InternalRel(X3) = relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X3)) & ! [X4] : (in(X4,the_carrier(X3)) <=> ? [X5] : (related(X1,X2,X5) & X4 = X5 & element(X5,the_carrier(X1)))))) | ~net_str(X3,X0) | ~strict_net_str(X3,X0)) | ~element(X2,the_carrier(X1))) | ~net_str(X1,X0) | empty_carrier(X1)) | ~one_sorted_str(X0) | empty_carrier(X0))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((netstr_restr_to_element(X0,X1,X2) = X3 <=> (the_mapping(X0,X3) = partfun_dom_restriction(the_carrier(X1),the_carrier(X0),the_mapping(X0,X1),the_carrier(X3)) & the_InternalRel(X3) = relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X3)) & ! [X4] : (in(X4,the_carrier(X3)) <=> ? [X5] : (related(X1,X2,X5) & X4 = X5 & element(X5,the_carrier(X1)))))) | (~net_str(X3,X0) | ~strict_net_str(X3,X0))) | ~element(X2,the_carrier(X1))) | (~net_str(X1,X0) | empty_carrier(X1))) | (~one_sorted_str(X0) | empty_carrier(X0)))),
  inference(ennf_transformation,[],[f13])).
fof(f13,axiom,(
  ! [X0] : ((one_sorted_str(X0) & ~empty_carrier(X0)) => ! [X1] : ((net_str(X1,X0) & ~empty_carrier(X1)) => ! [X2] : (element(X2,the_carrier(X1)) => ! [X3] : ((net_str(X3,X0) & strict_net_str(X3,X0)) => (netstr_restr_to_element(X0,X1,X2) = X3 <=> (the_mapping(X0,X3) = partfun_dom_restriction(the_carrier(X1),the_carrier(X0),the_mapping(X0,X1),the_carrier(X3)) & the_InternalRel(X3) = relation_restriction_as_relation_of(the_InternalRel(X1),the_carrier(X3)) & ! [X4] : (in(X4,the_carrier(X3)) <=> ? [X5] : (related(X1,X2,X5) & X4 = X5 & element(X5,the_carrier(X1))))))))))),
  file('SEU385+1.tptp',unknown)).
fof(f8195,plain,(
  ~spl28_5 | spl28_6 | ~spl28_11),
  inference(avatar_contradiction_clause,[],[f8194])).
fof(f8194,plain,(
  $false | (~spl28_5 | spl28_6 | ~spl28_11)),
  inference(subsumption_resolution,[],[f8193,f595])).
fof(f595,plain,(
  ~empty_carrier(netstr_restr_to_element(sK23,sK24,sK25)) | spl28_6),
  inference(avatar_component_clause,[],[f594])).
fof(f594,plain,(
  spl28_6 <=> empty_carrier(netstr_restr_to_element(sK23,sK24,sK25))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_6])])).
fof(f8193,plain,(
  empty_carrier(netstr_restr_to_element(sK23,sK24,sK25)) | (~spl28_5 | ~spl28_11)),
  inference(subsumption_resolution,[],[f8179,f1075])).
fof(f1075,plain,(
  one_sorted_str(netstr_restr_to_element(sK23,sK24,sK25)) | ~spl28_11),
  inference(avatar_component_clause,[],[f1074])).
fof(f1074,plain,(
  spl28_11 <=> one_sorted_str(netstr_restr_to_element(sK23,sK24,sK25))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_11])])).
fof(f8179,plain,(
  ~one_sorted_str(netstr_restr_to_element(sK23,sK24,sK25)) | empty_carrier(netstr_restr_to_element(sK23,sK24,sK25)) | ~spl28_5),
  inference(resolution,[],[f436,f292])).
fof(f292,plain,(
  ( ! [X0] : (~empty(the_carrier(X0)) | ~one_sorted_str(X0) | empty_carrier(X0)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : (~empty(the_carrier(X0)) | ~one_sorted_str(X0) | empty_carrier(X0))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : (~empty(the_carrier(X0)) | (~one_sorted_str(X0) | empty_carrier(X0)))),
  inference(ennf_transformation,[],[f46])).
fof(f46,axiom,(
  ! [X0] : ((one_sorted_str(X0) & ~empty_carrier(X0)) => ~empty(the_carrier(X0)))),
  file('SEU385+1.tptp',unknown)).
fof(f436,plain,(
  empty(the_carrier(netstr_restr_to_element(sK23,sK24,sK25))) | ~spl28_5),
  inference(avatar_component_clause,[],[f434])).
fof(f434,plain,(
  spl28_5 <=> empty(the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_5])])).
fof(f2443,plain,(
  spl28_11),
  inference(avatar_contradiction_clause,[],[f2442])).
fof(f2442,plain,(
  $false | spl28_11),
  inference(subsumption_resolution,[],[f2441,f342])).
fof(f2441,plain,(
  ~net_str(sK24,sK23) | spl28_11),
  inference(subsumption_resolution,[],[f2440,f343])).
fof(f2440,plain,(
  ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | spl28_11),
  inference(subsumption_resolution,[],[f2439,f338])).
fof(f2439,plain,(
  empty_carrier(sK23) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | spl28_11),
  inference(subsumption_resolution,[],[f2438,f339])).
fof(f2438,plain,(
  ~one_sorted_str(sK23) | empty_carrier(sK23) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | spl28_11),
  inference(subsumption_resolution,[],[f2436,f340])).
fof(f2436,plain,(
  empty_carrier(sK24) | ~one_sorted_str(sK23) | empty_carrier(sK23) | ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | spl28_11),
  inference(resolution,[],[f851,f1076])).
fof(f1076,plain,(
  ~one_sorted_str(netstr_restr_to_element(sK23,sK24,sK25)) | spl28_11),
  inference(avatar_component_clause,[],[f1074])).
fof(f851,plain,(
  ( ! [X2,X0,X1] : (one_sorted_str(netstr_restr_to_element(X1,X0,X2)) | empty_carrier(X0) | ~one_sorted_str(X1) | empty_carrier(X1) | ~element(X2,the_carrier(X0)) | ~net_str(X0,X1)) )),
  inference(resolution,[],[f562,f270])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (rel_str(netstr_restr_to_element(X2,X1,X0)) | ~net_str(X1,X2) | empty_carrier(X1) | ~one_sorted_str(X2) | empty_carrier(X2) | ~element(X0,the_carrier(X1))) )),
  inference(duplicate_literal_removal,[],[f561])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (~element(X0,the_carrier(X1)) | ~net_str(X1,X2) | empty_carrier(X1) | ~one_sorted_str(X2) | empty_carrier(X2) | rel_str(netstr_restr_to_element(X2,X1,X0)) | ~one_sorted_str(X2)) )),
  inference(resolution,[],[f267,f271])).
fof(f1528,plain,(
  spl28_18),
  inference(avatar_contradiction_clause,[],[f1527])).
fof(f1527,plain,(
  $false | spl28_18),
  inference(subsumption_resolution,[],[f1526,f342])).
fof(f1526,plain,(
  ~net_str(sK24,sK23) | spl28_18),
  inference(subsumption_resolution,[],[f1525,f339])).
fof(f1525,plain,(
  ~one_sorted_str(sK23) | ~net_str(sK24,sK23) | spl28_18),
  inference(resolution,[],[f1458,f505])).
fof(f1458,plain,(
  ~relation_of2(the_mapping(sK23,sK24),the_carrier(sK24),the_carrier(sK23)) | spl28_18),
  inference(avatar_component_clause,[],[f1456])).
fof(f1477,plain,(
  spl28_16),
  inference(avatar_contradiction_clause,[],[f1476])).
fof(f1476,plain,(
  $false | spl28_16),
  inference(subsumption_resolution,[],[f1475,f339])).
fof(f1475,plain,(
  ~one_sorted_str(sK23) | spl28_16),
  inference(subsumption_resolution,[],[f1474,f342])).
fof(f1474,plain,(
  ~net_str(sK24,sK23) | ~one_sorted_str(sK23) | spl28_16),
  inference(resolution,[],[f1450,f274])).
fof(f1450,plain,(
  ~function(the_mapping(sK23,sK24)) | spl28_16),
  inference(avatar_component_clause,[],[f1448])).
fof(f1080,plain,(
  ~spl28_11 | spl28_12 | spl28_6),
  inference(avatar_split_clause,[],[f1072,f594,f1078,f1074])).
fof(f1072,plain,(
  ( ! [X15] : (apply(the_mapping(X15,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),X15,the_mapping(X15,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~one_sorted_str(X15) | empty_carrier(X15) | ~one_sorted_str(netstr_restr_to_element(sK23,sK24,sK25)) | ~net_str(netstr_restr_to_element(sK23,sK24,sK25),X15)) ) | spl28_6),
  inference(subsumption_resolution,[],[f1058,f595])).
fof(f1058,plain,(
  ( ! [X15] : (apply(the_mapping(X15,netstr_restr_to_element(sK23,sK24,sK25)),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),X15,the_mapping(X15,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~one_sorted_str(X15) | empty_carrier(X15) | ~one_sorted_str(netstr_restr_to_element(sK23,sK24,sK25)) | empty_carrier(netstr_restr_to_element(sK23,sK24,sK25)) | ~net_str(netstr_restr_to_element(sK23,sK24,sK25),X15)) )),
  inference(resolution,[],[f659,f362])).
fof(f362,plain,(
  element(sK26,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))),
  inference(forward_demodulation,[],[f345,f346])).
fof(f345,plain,(
  element(sK27,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))),
  inference(cnf_transformation,[],[f223])).
fof(f610,plain,(
  ~spl28_6),
  inference(avatar_contradiction_clause,[],[f609])).
fof(f609,plain,(
  $false | ~spl28_6),
  inference(subsumption_resolution,[],[f608,f338])).
fof(f608,plain,(
  empty_carrier(sK23) | ~spl28_6),
  inference(subsumption_resolution,[],[f607,f339])).
fof(f607,plain,(
  ~one_sorted_str(sK23) | empty_carrier(sK23) | ~spl28_6),
  inference(subsumption_resolution,[],[f606,f340])).
fof(f606,plain,(
  empty_carrier(sK24) | ~one_sorted_str(sK23) | empty_carrier(sK23) | ~spl28_6),
  inference(subsumption_resolution,[],[f605,f341])).
fof(f341,plain,(
  directed_relstr(sK24)),
  inference(cnf_transformation,[],[f223])).
fof(f605,plain,(
  ~directed_relstr(sK24) | empty_carrier(sK24) | ~one_sorted_str(sK23) | empty_carrier(sK23) | ~spl28_6),
  inference(subsumption_resolution,[],[f604,f342])).
fof(f604,plain,(
  ~net_str(sK24,sK23) | ~directed_relstr(sK24) | empty_carrier(sK24) | ~one_sorted_str(sK23) | empty_carrier(sK23) | ~spl28_6),
  inference(subsumption_resolution,[],[f603,f343])).
fof(f603,plain,(
  ~element(sK25,the_carrier(sK24)) | ~net_str(sK24,sK23) | ~directed_relstr(sK24) | empty_carrier(sK24) | ~one_sorted_str(sK23) | empty_carrier(sK23) | ~spl28_6),
  inference(resolution,[],[f596,f294])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (~empty_carrier(netstr_restr_to_element(X0,X1,X2)) | ~element(X2,the_carrier(X1)) | ~net_str(X1,X0) | ~directed_relstr(X1) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ((strict_net_str(netstr_restr_to_element(X0,X1,X2),X0) & ~empty_carrier(netstr_restr_to_element(X0,X1,X2))) | ~element(X2,the_carrier(X1)) | ~net_str(X1,X0) | ~directed_relstr(X1) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((strict_net_str(netstr_restr_to_element(X0,X1,X2),X0) & ~empty_carrier(netstr_restr_to_element(X0,X1,X2))) | (~element(X2,the_carrier(X1)) | ~net_str(X1,X0) | ~directed_relstr(X1) | empty_carrier(X1) | ~one_sorted_str(X0) | empty_carrier(X0)))),
  inference(ennf_transformation,[],[f48])).
fof(f48,axiom,(
  ! [X0,X1,X2] : ((element(X2,the_carrier(X1)) & net_str(X1,X0) & directed_relstr(X1) & ~empty_carrier(X1) & one_sorted_str(X0) & ~empty_carrier(X0)) => (strict_net_str(netstr_restr_to_element(X0,X1,X2),X0) & ~empty_carrier(netstr_restr_to_element(X0,X1,X2))))),
  file('SEU385+1.tptp',unknown)).
fof(f596,plain,(
  empty_carrier(netstr_restr_to_element(sK23,sK24,sK25)) | ~spl28_6),
  inference(avatar_component_clause,[],[f594])).
fof(f600,plain,(
  spl28_6 | spl28_7),
  inference(avatar_split_clause,[],[f587,f598,f594])).
fof(f587,plain,(
  ( ! [X0] : (apply_netmap(X0,netstr_restr_to_element(sK23,sK24,sK25),sK26) = apply_on_structs(netstr_restr_to_element(sK23,sK24,sK25),X0,the_mapping(X0,netstr_restr_to_element(sK23,sK24,sK25)),sK26) | ~net_str(netstr_restr_to_element(sK23,sK24,sK25),X0) | empty_carrier(netstr_restr_to_element(sK23,sK24,sK25)) | ~one_sorted_str(X0) | empty_carrier(X0)) )),
  inference(resolution,[],[f257,f362])).
fof(f437,plain,(
  spl28_4 | spl28_5),
  inference(avatar_split_clause,[],[f410,f434,f430])).
fof(f410,plain,(
  empty(the_carrier(netstr_restr_to_element(sK23,sK24,sK25))) | in(sK26,the_carrier(netstr_restr_to_element(sK23,sK24,sK25)))),
  inference(resolution,[],[f349,f362])).
fof(f349,plain,(
  ( ! [X0,X1] : (~element(X0,X1) | empty(X1) | in(X0,X1)) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : (in(X0,X1) | empty(X1) | ~element(X0,X1))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ((in(X0,X1) | empty(X1)) | ~element(X0,X1))),
  inference(ennf_transformation,[],[f73])).
fof(f73,axiom,(
  ! [X0,X1] : (element(X0,X1) => (in(X0,X1) | empty(X1)))),
  file('SEU385+1.tptp',unknown)).
% SZS output end Proof for SEU385+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 106053
% Time elapsed: 2.809 s
% ------------------------------
% ------------------------------
