% Refutation found. Thanks to Tanya!
% SZS status Theorem for SEU230+2
% SZS output start Proof for SEU230+2
fof(f2260,plain,(
  $false),
  inference(subsumption_resolution,[],[f2252,f1404])).
fof(f1404,plain,(
  ( ! [X3] : (in(X3,unordered_pair(X3,X3))) )),
  inference(equality_resolution,[],[f1403])).
fof(f1403,plain,(
  ( ! [X3,X1] : (in(X3,X1) | unordered_pair(X3,X3) != X1) )),
  inference(equality_resolution,[],[f1258])).
fof(f1258,plain,(
  ( ! [X0,X3,X1] : (in(X3,X1) | X0 != X3 | unordered_pair(X0,X0) != X1) )),
  inference(definition_unfolding,[],[f829,f1189])).
fof(f1189,plain,(
  ( ! [X0] : (singleton(X0) = unordered_pair(X0,X0)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,axiom,(
  ! [X0] : singleton(X0) = unordered_pair(X0,X0)),
  file('SEU230+2.tptp',unknown)).
fof(f829,plain,(
  ( ! [X0,X3,X1] : (in(X3,X1) | X0 != X3 | singleton(X0) != X1) )),
  inference(cnf_transformation,[],[f555])).
fof(f555,plain,(
  ! [X0,X1] : ((singleton(X0) = X1 | ((sK24(X0,X1) != X0 | ~in(sK24(X0,X1),X1)) & (sK24(X0,X1) = X0 | in(sK24(X0,X1),X1)))) & (! [X3] : ((in(X3,X1) | X0 != X3) & (X0 = X3 | ~in(X3,X1))) | singleton(X0) != X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24])],[f553,f554])).
fof(f554,plain,(
  ! [X1,X0] : (? [X2] : ((X0 != X2 | ~in(X2,X1)) & (X0 = X2 | in(X2,X1))) => ((sK24(X0,X1) != X0 | ~in(sK24(X0,X1),X1)) & (sK24(X0,X1) = X0 | in(sK24(X0,X1),X1))))),
  introduced(choice_axiom,[])).
fof(f553,plain,(
  ! [X0,X1] : ((singleton(X0) = X1 | ? [X2] : ((X0 != X2 | ~in(X2,X1)) & (X0 = X2 | in(X2,X1)))) & (! [X3] : ((in(X3,X1) | X0 != X3) & (X0 = X3 | ~in(X3,X1))) | singleton(X0) != X1))),
  inference(rectify,[],[f552])).
fof(f552,plain,(
  ! [X0,X1] : ((singleton(X0) = X1 | ? [X2] : ((X0 != X2 | ~in(X2,X1)) & (X0 = X2 | in(X2,X1)))) & (! [X2] : ((in(X2,X1) | X0 != X2) & (X0 = X2 | ~in(X2,X1))) | singleton(X0) != X1))),
  inference(nnf_transformation,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1] : (singleton(X0) = X1 <=> ! [X2] : (in(X2,X1) <=> X0 = X2))),
  file('SEU230+2.tptp',unknown)).
fof(f2252,plain,(
  ~in(sK75,unordered_pair(sK75,sK75))),
  inference(resolution,[],[f1416,f1320])).
fof(f1320,plain,(
  ~in(sK75,set_union2(sK75,unordered_pair(sK75,sK75)))),
  inference(definition_unfolding,[],[f1047,f1226])).
fof(f1226,plain,(
  ( ! [X0] : (succ(X0) = set_union2(X0,unordered_pair(X0,X0))) )),
  inference(definition_unfolding,[],[f816,f1189])).
fof(f816,plain,(
  ( ! [X0] : (succ(X0) = set_union2(X0,singleton(X0))) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,axiom,(
  ! [X0] : succ(X0) = set_union2(X0,singleton(X0))),
  file('SEU230+2.tptp',unknown)).
fof(f1047,plain,(
  ~in(sK75,succ(sK75))),
  inference(cnf_transformation,[],[f685])).
fof(f685,plain,(
  ~in(sK75,succ(sK75))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK75])],[f366,f684])).
fof(f684,plain,(
  ? [X0] : ~in(X0,succ(X0)) => ~in(sK75,succ(sK75))),
  introduced(choice_axiom,[])).
fof(f366,plain,(
  ? [X0] : ~in(X0,succ(X0))),
  inference(ennf_transformation,[],[f144])).
fof(f144,negated_conjecture,(
  ~! [X0] : in(X0,succ(X0))),
  inference(negated_conjecture,[],[f143])).
fof(f143,conjecture,(
  ! [X0] : in(X0,succ(X0))),
  file('SEU230+2.tptp',unknown)).
fof(f1416,plain,(
  ( ! [X4,X0,X1] : (in(X4,set_union2(X0,X1)) | ~in(X4,X1)) )),
  inference(equality_resolution,[],[f854])).
fof(f854,plain,(
  ( ! [X4,X2,X0,X1] : (in(X4,X2) | ~in(X4,X1) | set_union2(X0,X1) != X2) )),
  inference(cnf_transformation,[],[f578])).
fof(f578,plain,(
  ! [X0,X1,X2] : ((set_union2(X0,X1) = X2 | (((~in(sK30(X0,X1,X2),X1) & ~in(sK30(X0,X1,X2),X0)) | ~in(sK30(X0,X1,X2),X2)) & (in(sK30(X0,X1,X2),X1) | in(sK30(X0,X1,X2),X0) | in(sK30(X0,X1,X2),X2)))) & (! [X4] : ((in(X4,X2) | (~in(X4,X1) & ~in(X4,X0))) & (in(X4,X1) | in(X4,X0) | ~in(X4,X2))) | set_union2(X0,X1) != X2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK30])],[f576,f577])).
fof(f577,plain,(
  ! [X2,X1,X0] : (? [X3] : (((~in(X3,X1) & ~in(X3,X0)) | ~in(X3,X2)) & (in(X3,X1) | in(X3,X0) | in(X3,X2))) => (((~in(sK30(X0,X1,X2),X1) & ~in(sK30(X0,X1,X2),X0)) | ~in(sK30(X0,X1,X2),X2)) & (in(sK30(X0,X1,X2),X1) | in(sK30(X0,X1,X2),X0) | in(sK30(X0,X1,X2),X2))))),
  introduced(choice_axiom,[])).
fof(f576,plain,(
  ! [X0,X1,X2] : ((set_union2(X0,X1) = X2 | ? [X3] : (((~in(X3,X1) & ~in(X3,X0)) | ~in(X3,X2)) & (in(X3,X1) | in(X3,X0) | in(X3,X2)))) & (! [X4] : ((in(X4,X2) | (~in(X4,X1) & ~in(X4,X0))) & (in(X4,X1) | in(X4,X0) | ~in(X4,X2))) | set_union2(X0,X1) != X2))),
  inference(rectify,[],[f575])).
fof(f575,plain,(
  ! [X0,X1,X2] : ((set_union2(X0,X1) = X2 | ? [X3] : (((~in(X3,X1) & ~in(X3,X0)) | ~in(X3,X2)) & (in(X3,X1) | in(X3,X0) | in(X3,X2)))) & (! [X3] : ((in(X3,X2) | (~in(X3,X1) & ~in(X3,X0))) & (in(X3,X1) | in(X3,X0) | ~in(X3,X2))) | set_union2(X0,X1) != X2))),
  inference(flattening,[],[f574])).
fof(f574,plain,(
  ! [X0,X1,X2] : ((set_union2(X0,X1) = X2 | ? [X3] : (((~in(X3,X1) & ~in(X3,X0)) | ~in(X3,X2)) & ((in(X3,X1) | in(X3,X0)) | in(X3,X2)))) & (! [X3] : ((in(X3,X2) | (~in(X3,X1) & ~in(X3,X0))) & ((in(X3,X1) | in(X3,X0)) | ~in(X3,X2))) | set_union2(X0,X1) != X2))),
  inference(nnf_transformation,[],[f27])).
fof(f27,axiom,(
  ! [X0,X1,X2] : (set_union2(X0,X1) = X2 <=> ! [X3] : (in(X3,X2) <=> (in(X3,X1) | in(X3,X0))))),
  file('SEU230+2.tptp',unknown)).
% SZS output end Proof for SEU230+2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6140
% Time elapsed: 0.024 s
% ------------------------------
% ------------------------------
