% Refutation found. Thanks to Tanya!
% SZS status Theorem for SET684+3
% SZS output start Proof for SET684+3
fof(f2414,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f212,f217,f222,f612,f640,f802,f805,f1548,f2413])).
fof(f2413,plain,(
  ~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23),
  inference(avatar_contradiction_clause,[],[f2412])).
fof(f2412,plain,(
  $false | (~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23)),
  inference(subsumption_resolution,[],[f2410,f1965])).
fof(f1965,plain,(
  ~member(sK0(sK16,sK17,sK14,sK15),sK12) | (~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1964,f170])).
fof(f170,plain,(
  ( ! [X0] : (ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,axiom,(
  ! [X0] : ilf_type(X0,set_type)),
  file('SET684+3.tptp',unknown)).
fof(f1964,plain,(
  ~member(sK0(sK16,sK17,sK14,sK15),sK12) | ~ilf_type(sK0(sK16,sK17,sK14,sK15),set_type) | (~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1963,f173])).
fof(f173,plain,(
  ~empty(sK12)),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  (((((((! [X7] : (~member(ordered_pair(X7,sK17),sK15) | ~member(ordered_pair(sK16,X7),sK14) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15))) & ((member(ordered_pair(sK18,sK17),sK15) & member(ordered_pair(sK16,sK18),sK14) & ilf_type(sK18,member_type(sK12))) | member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15))) & ilf_type(sK17,member_type(sK13))) & ilf_type(sK16,member_type(sK11))) & ilf_type(sK15,relation_type(sK12,sK13))) & ilf_type(sK14,relation_type(sK11,sK12))) & ilf_type(sK13,set_type) & ~empty(sK13)) & ilf_type(sK12,set_type) & ~empty(sK12)) & ilf_type(sK11,set_type) & ~empty(sK11)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13,sK14,sK15,sK16,sK17,sK18])],[f107,f115,f114,f113,f112,f111,f110,f109,f108])).
fof(f108,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(X1))) | ~member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),X3) & ilf_type(X8,member_type(X1))) | member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(X0))) & ilf_type(X4,relation_type(X1,X2))) & ilf_type(X3,relation_type(X0,X1))) & ilf_type(X2,set_type) & ~empty(X2)) & ilf_type(X1,set_type) & ~empty(X1)) & ilf_type(X0,set_type) & ~empty(X0)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(X1))) | ~member(ordered_pair(X5,X6),compose5(sK11,X1,X2,X3,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),X3) & ilf_type(X8,member_type(X1))) | member(ordered_pair(X5,X6),compose5(sK11,X1,X2,X3,X4))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(sK11))) & ilf_type(X4,relation_type(X1,X2))) & ilf_type(X3,relation_type(sK11,X1))) & ilf_type(X2,set_type) & ~empty(X2)) & ilf_type(X1,set_type) & ~empty(X1)) & ilf_type(sK11,set_type) & ~empty(sK11))),
  introduced(choice_axiom,[])).
fof(f109,plain,(
  ? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(X1))) | ~member(ordered_pair(X5,X6),compose5(sK11,X1,X2,X3,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),X3) & ilf_type(X8,member_type(X1))) | member(ordered_pair(X5,X6),compose5(sK11,X1,X2,X3,X4))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(sK11))) & ilf_type(X4,relation_type(X1,X2))) & ilf_type(X3,relation_type(sK11,X1))) & ilf_type(X2,set_type) & ~empty(X2)) & ilf_type(X1,set_type) & ~empty(X1)) => (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(X5,X6),compose5(sK11,sK12,X2,X3,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),X3) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(X5,X6),compose5(sK11,sK12,X2,X3,X4))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(sK11))) & ilf_type(X4,relation_type(sK12,X2))) & ilf_type(X3,relation_type(sK11,sK12))) & ilf_type(X2,set_type) & ~empty(X2)) & ilf_type(sK12,set_type) & ~empty(sK12))),
  introduced(choice_axiom,[])).
fof(f110,plain,(
  ? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(X5,X6),compose5(sK11,sK12,X2,X3,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),X3) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(X5,X6),compose5(sK11,sK12,X2,X3,X4))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(sK11))) & ilf_type(X4,relation_type(sK12,X2))) & ilf_type(X3,relation_type(sK11,sK12))) & ilf_type(X2,set_type) & ~empty(X2)) => (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,X3,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),X3) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,X3,X4))) & ilf_type(X6,member_type(sK13))) & ilf_type(X5,member_type(sK11))) & ilf_type(X4,relation_type(sK12,sK13))) & ilf_type(X3,relation_type(sK11,sK12))) & ilf_type(sK13,set_type) & ~empty(sK13))),
  introduced(choice_axiom,[])).
fof(f111,plain,(
  ? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,X3,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),X3) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,X3,X4))) & ilf_type(X6,member_type(sK13))) & ilf_type(X5,member_type(sK11))) & ilf_type(X4,relation_type(sK12,sK13))) & ilf_type(X3,relation_type(sK11,sK12))) => (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),sK14) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,sK14,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),sK14) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,sK14,X4))) & ilf_type(X6,member_type(sK13))) & ilf_type(X5,member_type(sK11))) & ilf_type(X4,relation_type(sK12,sK13))) & ilf_type(sK14,relation_type(sK11,sK12)))),
  introduced(choice_axiom,[])).
fof(f112,plain,(
  ? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),sK14) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,sK14,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),sK14) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,sK14,X4))) & ilf_type(X6,member_type(sK13))) & ilf_type(X5,member_type(sK11))) & ilf_type(X4,relation_type(sK12,sK13))) => (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),sK15) | ~member(ordered_pair(X5,X7),sK14) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,sK14,sK15))) & (? [X8] : (member(ordered_pair(X8,X6),sK15) & member(ordered_pair(X5,X8),sK14) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,sK14,sK15))) & ilf_type(X6,member_type(sK13))) & ilf_type(X5,member_type(sK11))) & ilf_type(sK15,relation_type(sK12,sK13)))),
  introduced(choice_axiom,[])).
fof(f113,plain,(
  ? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),sK15) | ~member(ordered_pair(X5,X7),sK14) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,sK14,sK15))) & (? [X8] : (member(ordered_pair(X8,X6),sK15) & member(ordered_pair(X5,X8),sK14) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(X5,X6),compose5(sK11,sK12,sK13,sK14,sK15))) & ilf_type(X6,member_type(sK13))) & ilf_type(X5,member_type(sK11))) => (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),sK15) | ~member(ordered_pair(sK16,X7),sK14) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(sK16,X6),compose5(sK11,sK12,sK13,sK14,sK15))) & (? [X8] : (member(ordered_pair(X8,X6),sK15) & member(ordered_pair(sK16,X8),sK14) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(sK16,X6),compose5(sK11,sK12,sK13,sK14,sK15))) & ilf_type(X6,member_type(sK13))) & ilf_type(sK16,member_type(sK11)))),
  introduced(choice_axiom,[])).
fof(f114,plain,(
  ? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),sK15) | ~member(ordered_pair(sK16,X7),sK14) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(sK16,X6),compose5(sK11,sK12,sK13,sK14,sK15))) & (? [X8] : (member(ordered_pair(X8,X6),sK15) & member(ordered_pair(sK16,X8),sK14) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(sK16,X6),compose5(sK11,sK12,sK13,sK14,sK15))) & ilf_type(X6,member_type(sK13))) => ((! [X7] : (~member(ordered_pair(X7,sK17),sK15) | ~member(ordered_pair(sK16,X7),sK14) | ~ilf_type(X7,member_type(sK12))) | ~member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15))) & (? [X8] : (member(ordered_pair(X8,sK17),sK15) & member(ordered_pair(sK16,X8),sK14) & ilf_type(X8,member_type(sK12))) | member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15))) & ilf_type(sK17,member_type(sK13)))),
  introduced(choice_axiom,[])).
fof(f115,plain,(
  ? [X8] : (member(ordered_pair(X8,sK17),sK15) & member(ordered_pair(sK16,X8),sK14) & ilf_type(X8,member_type(sK12))) => (member(ordered_pair(sK18,sK17),sK15) & member(ordered_pair(sK16,sK18),sK14) & ilf_type(sK18,member_type(sK12)))),
  introduced(choice_axiom,[])).
fof(f107,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(X1))) | ~member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4))) & (? [X8] : (member(ordered_pair(X8,X6),X4) & member(ordered_pair(X5,X8),X3) & ilf_type(X8,member_type(X1))) | member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(X0))) & ilf_type(X4,relation_type(X1,X2))) & ilf_type(X3,relation_type(X0,X1))) & ilf_type(X2,set_type) & ~empty(X2)) & ilf_type(X1,set_type) & ~empty(X1)) & ilf_type(X0,set_type) & ~empty(X0))),
  inference(rectify,[],[f106])).
fof(f106,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(X1))) | ~member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4))) & (? [X7] : (member(ordered_pair(X7,X6),X4) & member(ordered_pair(X5,X7),X3) & ilf_type(X7,member_type(X1))) | member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(X0))) & ilf_type(X4,relation_type(X1,X2))) & ilf_type(X3,relation_type(X0,X1))) & ilf_type(X2,set_type) & ~empty(X2)) & ilf_type(X1,set_type) & ~empty(X1)) & ilf_type(X0,set_type) & ~empty(X0))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : (((! [X7] : (~member(ordered_pair(X7,X6),X4) | ~member(ordered_pair(X5,X7),X3) | ~ilf_type(X7,member_type(X1))) | ~member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4))) & (? [X7] : (member(ordered_pair(X7,X6),X4) & member(ordered_pair(X5,X7),X3) & ilf_type(X7,member_type(X1))) | member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4)))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(X0))) & ilf_type(X4,relation_type(X1,X2))) & ilf_type(X3,relation_type(X0,X1))) & ilf_type(X2,set_type) & ~empty(X2)) & ilf_type(X1,set_type) & ~empty(X1)) & ilf_type(X0,set_type) & ~empty(X0))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4)) <~> ? [X7] : (member(ordered_pair(X7,X6),X4) & member(ordered_pair(X5,X7),X3) & ilf_type(X7,member_type(X1)))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(X0))) & ilf_type(X4,relation_type(X1,X2))) & ilf_type(X3,relation_type(X0,X1))) & ilf_type(X2,set_type) & ~empty(X2)) & ilf_type(X1,set_type) & ~empty(X1)) & ilf_type(X0,set_type) & ~empty(X0))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : ((member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4)) <~> ? [X7] : (member(ordered_pair(X7,X6),X4) & member(ordered_pair(X5,X7),X3) & ilf_type(X7,member_type(X1)))) & ilf_type(X6,member_type(X2))) & ilf_type(X5,member_type(X0))) & ilf_type(X4,relation_type(X1,X2))) & ilf_type(X3,relation_type(X0,X1))) & (ilf_type(X2,set_type) & ~empty(X2))) & (ilf_type(X1,set_type) & ~empty(X1))) & (ilf_type(X0,set_type) & ~empty(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,negated_conjecture,(
  ~! [X0] : ((ilf_type(X0,set_type) & ~empty(X0)) => ! [X1] : ((ilf_type(X1,set_type) & ~empty(X1)) => ! [X2] : ((ilf_type(X2,set_type) & ~empty(X2)) => ! [X3] : (ilf_type(X3,relation_type(X0,X1)) => ! [X4] : (ilf_type(X4,relation_type(X1,X2)) => ! [X5] : (ilf_type(X5,member_type(X0)) => ! [X6] : (ilf_type(X6,member_type(X2)) => (member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4)) <=> ? [X7] : (member(ordered_pair(X7,X6),X4) & member(ordered_pair(X5,X7),X3) & ilf_type(X7,member_type(X1)))))))))))),
  inference(negated_conjecture,[],[f30])).
fof(f30,conjecture,(
  ! [X0] : ((ilf_type(X0,set_type) & ~empty(X0)) => ! [X1] : ((ilf_type(X1,set_type) & ~empty(X1)) => ! [X2] : ((ilf_type(X2,set_type) & ~empty(X2)) => ! [X3] : (ilf_type(X3,relation_type(X0,X1)) => ! [X4] : (ilf_type(X4,relation_type(X1,X2)) => ! [X5] : (ilf_type(X5,member_type(X0)) => ! [X6] : (ilf_type(X6,member_type(X2)) => (member(ordered_pair(X5,X6),compose5(X0,X1,X2,X3,X4)) <=> ? [X7] : (member(ordered_pair(X7,X6),X4) & member(ordered_pair(X5,X7),X3) & ilf_type(X7,member_type(X1)))))))))))),
  file('SET684+3.tptp',unknown)).
fof(f1963,plain,(
  ~member(sK0(sK16,sK17,sK14,sK15),sK12) | empty(sK12) | ~ilf_type(sK0(sK16,sK17,sK14,sK15),set_type) | (~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1962,f170])).
fof(f1962,plain,(
  ~member(sK0(sK16,sK17,sK14,sK15),sK12) | ~ilf_type(sK12,set_type) | empty(sK12) | ~ilf_type(sK0(sK16,sK17,sK14,sK15),set_type) | (~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23)),
  inference(resolution,[],[f1961,f130])).
fof(f130,plain,(
  ( ! [X0,X1] : (ilf_type(X0,member_type(X1)) | ~member(X0,X1) | ~ilf_type(X1,set_type) | empty(X1) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (! [X1] : (((ilf_type(X0,member_type(X1)) | ~member(X0,X1)) & (member(X0,X1) | ~ilf_type(X0,member_type(X1)))) | ~ilf_type(X1,set_type) | empty(X1)) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (! [X1] : ((ilf_type(X0,member_type(X1)) <=> member(X0,X1)) | ~ilf_type(X1,set_type) | empty(X1)) | ~ilf_type(X0,set_type))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : (! [X1] : ((ilf_type(X0,member_type(X1)) <=> member(X0,X1)) | (~ilf_type(X1,set_type) | empty(X1))) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : ((ilf_type(X1,set_type) & ~empty(X1)) => (ilf_type(X0,member_type(X1)) <=> member(X0,X1))))),
  file('SET684+3.tptp',unknown)).
fof(f1961,plain,(
  ~ilf_type(sK0(sK16,sK17,sK14,sK15),member_type(sK12)) | (~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1960,f170])).
fof(f1960,plain,(
  ~ilf_type(sK0(sK16,sK17,sK14,sK15),member_type(sK12)) | ~ilf_type(sK16,set_type) | (~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1959,f170])).
fof(f1959,plain,(
  ~ilf_type(sK0(sK16,sK17,sK14,sK15),member_type(sK12)) | ~ilf_type(sK17,set_type) | ~ilf_type(sK16,set_type) | (~spl19_4 | ~spl19_19 | ~spl19_20 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1958,f607])).
fof(f607,plain,(
  ilf_type(sK14,binary_relation_type) | ~spl19_19),
  inference(avatar_component_clause,[],[f606])).
fof(f606,plain,(
  spl19_19 <=> ilf_type(sK14,binary_relation_type)),
  introduced(avatar_definition,[new_symbols(naming,[spl19_19])])).
fof(f1958,plain,(
  ~ilf_type(sK0(sK16,sK17,sK14,sK15),member_type(sK12)) | ~ilf_type(sK14,binary_relation_type) | ~ilf_type(sK17,set_type) | ~ilf_type(sK16,set_type) | (~spl19_4 | ~spl19_20 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1957,f797])).
fof(f797,plain,(
  ilf_type(sK15,binary_relation_type) | ~spl19_23),
  inference(avatar_component_clause,[],[f796])).
fof(f796,plain,(
  spl19_23 <=> ilf_type(sK15,binary_relation_type)),
  introduced(avatar_definition,[new_symbols(naming,[spl19_23])])).
fof(f1957,plain,(
  ~ilf_type(sK0(sK16,sK17,sK14,sK15),member_type(sK12)) | ~ilf_type(sK15,binary_relation_type) | ~ilf_type(sK14,binary_relation_type) | ~ilf_type(sK17,set_type) | ~ilf_type(sK16,set_type) | (~spl19_4 | ~spl19_20)),
  inference(subsumption_resolution,[],[f1956,f1564])).
fof(f1564,plain,(
  member(ordered_pair(sK16,sK17),compose(sK14,sK15)) | ~spl19_4),
  inference(forward_demodulation,[],[f207,f1322])).
fof(f1322,plain,(
  compose5(sK11,sK12,sK13,sK14,sK15) = compose(sK14,sK15)),
  inference(resolution,[],[f753,f177])).
fof(f177,plain,(
  ilf_type(sK14,relation_type(sK11,sK12))),
  inference(cnf_transformation,[],[f116])).
fof(f753,plain,(
  ( ! [X10,X11] : (~ilf_type(X11,relation_type(X10,sK12)) | compose5(X10,sK12,sK13,X11,sK15) = compose(X11,sK15)) )),
  inference(subsumption_resolution,[],[f752,f170])).
fof(f752,plain,(
  ( ! [X10,X11] : (compose5(X10,sK12,sK13,X11,sK15) = compose(X11,sK15) | ~ilf_type(X11,relation_type(X10,sK12)) | ~ilf_type(X10,set_type)) )),
  inference(subsumption_resolution,[],[f751,f170])).
fof(f751,plain,(
  ( ! [X10,X11] : (compose5(X10,sK12,sK13,X11,sK15) = compose(X11,sK15) | ~ilf_type(X11,relation_type(X10,sK12)) | ~ilf_type(sK12,set_type) | ~ilf_type(X10,set_type)) )),
  inference(subsumption_resolution,[],[f740,f170])).
fof(f740,plain,(
  ( ! [X10,X11] : (compose5(X10,sK12,sK13,X11,sK15) = compose(X11,sK15) | ~ilf_type(X11,relation_type(X10,sK12)) | ~ilf_type(sK13,set_type) | ~ilf_type(sK12,set_type) | ~ilf_type(X10,set_type)) )),
  inference(resolution,[],[f168,f178])).
fof(f178,plain,(
  ilf_type(sK15,relation_type(sK12,sK13))),
  inference(cnf_transformation,[],[f116])).
fof(f168,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~ilf_type(X4,relation_type(X1,X2)) | compose5(X0,X1,X2,X3,X4) = compose(X3,X4) | ~ilf_type(X3,relation_type(X0,X1)) | ~ilf_type(X2,set_type) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (compose5(X0,X1,X2,X3,X4) = compose(X3,X4) | ~ilf_type(X4,relation_type(X1,X2))) | ~ilf_type(X3,relation_type(X0,X1))) | ~ilf_type(X2,set_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f27])).
fof(f27,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => ! [X2] : (ilf_type(X2,set_type) => ! [X3] : (ilf_type(X3,relation_type(X0,X1)) => ! [X4] : (ilf_type(X4,relation_type(X1,X2)) => compose5(X0,X1,X2,X3,X4) = compose(X3,X4))))))),
  file('SET684+3.tptp',unknown)).
fof(f207,plain,(
  member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15)) | ~spl19_4),
  inference(avatar_component_clause,[],[f206])).
fof(f206,plain,(
  spl19_4 <=> member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15))),
  introduced(avatar_definition,[new_symbols(naming,[spl19_4])])).
fof(f1956,plain,(
  ~member(ordered_pair(sK16,sK17),compose(sK14,sK15)) | ~ilf_type(sK0(sK16,sK17,sK14,sK15),member_type(sK12)) | ~ilf_type(sK15,binary_relation_type) | ~ilf_type(sK14,binary_relation_type) | ~ilf_type(sK17,set_type) | ~ilf_type(sK16,set_type) | ~spl19_20),
  inference(duplicate_literal_removal,[],[f1955])).
fof(f1955,plain,(
  ~member(ordered_pair(sK16,sK17),compose(sK14,sK15)) | ~ilf_type(sK0(sK16,sK17,sK14,sK15),member_type(sK12)) | ~ilf_type(sK15,binary_relation_type) | ~member(ordered_pair(sK16,sK17),compose(sK14,sK15)) | ~ilf_type(sK15,binary_relation_type) | ~ilf_type(sK14,binary_relation_type) | ~ilf_type(sK17,set_type) | ~ilf_type(sK16,set_type) | ~spl19_20),
  inference(resolution,[],[f611,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X3,X1] : (member(ordered_pair(sK0(X0,X1,X2,X3),X1),X3) | ~member(ordered_pair(X0,X1),compose(X2,X3)) | ~ilf_type(X3,binary_relation_type) | ~ilf_type(X2,binary_relation_type) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((member(ordered_pair(X0,X1),compose(X2,X3)) | ! [X4] : (~member(ordered_pair(X4,X1),X3) | ~member(ordered_pair(X0,X4),X2) | ~ilf_type(X4,set_type))) & ((member(ordered_pair(sK0(X0,X1,X2,X3),X1),X3) & member(ordered_pair(X0,sK0(X0,X1,X2,X3)),X2) & ilf_type(sK0(X0,X1,X2,X3),set_type)) | ~member(ordered_pair(X0,X1),compose(X2,X3)))) | ~ilf_type(X3,binary_relation_type)) | ~ilf_type(X2,binary_relation_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f70,f71])).
fof(f71,plain,(
  ! [X3,X2,X1,X0] : (? [X5] : (member(ordered_pair(X5,X1),X3) & member(ordered_pair(X0,X5),X2) & ilf_type(X5,set_type)) => (member(ordered_pair(sK0(X0,X1,X2,X3),X1),X3) & member(ordered_pair(X0,sK0(X0,X1,X2,X3)),X2) & ilf_type(sK0(X0,X1,X2,X3),set_type)))),
  introduced(choice_axiom,[])).
fof(f70,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((member(ordered_pair(X0,X1),compose(X2,X3)) | ! [X4] : (~member(ordered_pair(X4,X1),X3) | ~member(ordered_pair(X0,X4),X2) | ~ilf_type(X4,set_type))) & (? [X5] : (member(ordered_pair(X5,X1),X3) & member(ordered_pair(X0,X5),X2) & ilf_type(X5,set_type)) | ~member(ordered_pair(X0,X1),compose(X2,X3)))) | ~ilf_type(X3,binary_relation_type)) | ~ilf_type(X2,binary_relation_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(rectify,[],[f69])).
fof(f69,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((member(ordered_pair(X0,X1),compose(X2,X3)) | ! [X4] : (~member(ordered_pair(X4,X1),X3) | ~member(ordered_pair(X0,X4),X2) | ~ilf_type(X4,set_type))) & (? [X4] : (member(ordered_pair(X4,X1),X3) & member(ordered_pair(X0,X4),X2) & ilf_type(X4,set_type)) | ~member(ordered_pair(X0,X1),compose(X2,X3)))) | ~ilf_type(X3,binary_relation_type)) | ~ilf_type(X2,binary_relation_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((member(ordered_pair(X0,X1),compose(X2,X3)) <=> ? [X4] : (member(ordered_pair(X4,X1),X3) & member(ordered_pair(X0,X4),X2) & ilf_type(X4,set_type))) | ~ilf_type(X3,binary_relation_type)) | ~ilf_type(X2,binary_relation_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => ! [X2] : (ilf_type(X2,binary_relation_type) => ! [X3] : (ilf_type(X3,binary_relation_type) => (member(ordered_pair(X0,X1),compose(X2,X3)) <=> ? [X4] : (member(ordered_pair(X4,X1),X3) & member(ordered_pair(X0,X4),X2) & ilf_type(X4,set_type)))))))),
  file('SET684+3.tptp',unknown)).
fof(f611,plain,(
  ( ! [X0,X1] : (~member(ordered_pair(sK0(sK16,X0,sK14,X1),sK17),sK15) | ~member(ordered_pair(sK16,X0),compose(sK14,X1)) | ~ilf_type(sK0(sK16,X0,sK14,X1),member_type(sK12)) | ~ilf_type(X1,binary_relation_type)) ) | ~spl19_20),
  inference(avatar_component_clause,[],[f610])).
fof(f610,plain,(
  spl19_20 <=> ! [X1,X0] : (~member(ordered_pair(sK16,X0),compose(sK14,X1)) | ~member(ordered_pair(sK0(sK16,X0,sK14,X1),sK17),sK15) | ~ilf_type(sK0(sK16,X0,sK14,X1),member_type(sK12)) | ~ilf_type(X1,binary_relation_type))),
  introduced(avatar_definition,[new_symbols(naming,[spl19_20])])).
fof(f2410,plain,(
  member(sK0(sK16,sK17,sK14,sK15),sK12) | (~spl19_4 | ~spl19_19 | ~spl19_23)),
  inference(resolution,[],[f1925,f177])).
fof(f1925,plain,(
  ( ! [X4,X5] : (~ilf_type(sK14,relation_type(X5,X4)) | member(sK0(sK16,sK17,sK14,sK15),X4)) ) | (~spl19_4 | ~spl19_19 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1924,f607])).
fof(f1924,plain,(
  ( ! [X4,X5] : (~ilf_type(sK14,binary_relation_type) | member(sK0(sK16,sK17,sK14,sK15),X4) | ~ilf_type(sK14,relation_type(X5,X4))) ) | (~spl19_4 | ~spl19_23)),
  inference(subsumption_resolution,[],[f1915,f797])).
fof(f1915,plain,(
  ( ! [X4,X5] : (~ilf_type(sK15,binary_relation_type) | ~ilf_type(sK14,binary_relation_type) | member(sK0(sK16,sK17,sK14,sK15),X4) | ~ilf_type(sK14,relation_type(X5,X4))) ) | ~spl19_4),
  inference(resolution,[],[f617,f1564])).
fof(f617,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~member(ordered_pair(X2,X3),compose(X4,X5)) | ~ilf_type(X5,binary_relation_type) | ~ilf_type(X4,binary_relation_type) | member(sK0(X2,X3,X4,X5),X6) | ~ilf_type(X4,relation_type(X7,X6))) )),
  inference(subsumption_resolution,[],[f616,f170])).
fof(f616,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~member(ordered_pair(X2,X3),compose(X4,X5)) | ~ilf_type(X5,binary_relation_type) | ~ilf_type(X4,binary_relation_type) | member(sK0(X2,X3,X4,X5),X6) | ~ilf_type(X4,relation_type(X7,X6)) | ~ilf_type(X7,set_type)) )),
  inference(subsumption_resolution,[],[f615,f170])).
fof(f615,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~member(ordered_pair(X2,X3),compose(X4,X5)) | ~ilf_type(X5,binary_relation_type) | ~ilf_type(X4,binary_relation_type) | member(sK0(X2,X3,X4,X5),X6) | ~ilf_type(X4,relation_type(X7,X6)) | ~ilf_type(X6,set_type) | ~ilf_type(X7,set_type)) )),
  inference(subsumption_resolution,[],[f614,f170])).
fof(f614,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~member(ordered_pair(X2,X3),compose(X4,X5)) | ~ilf_type(X5,binary_relation_type) | ~ilf_type(X4,binary_relation_type) | member(sK0(X2,X3,X4,X5),X6) | ~ilf_type(X4,relation_type(X7,X6)) | ~ilf_type(sK0(X2,X3,X4,X5),set_type) | ~ilf_type(X6,set_type) | ~ilf_type(X7,set_type)) )),
  inference(subsumption_resolution,[],[f613,f170])).
fof(f613,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~member(ordered_pair(X2,X3),compose(X4,X5)) | ~ilf_type(X5,binary_relation_type) | ~ilf_type(X4,binary_relation_type) | ~ilf_type(X2,set_type) | member(sK0(X2,X3,X4,X5),X6) | ~ilf_type(X4,relation_type(X7,X6)) | ~ilf_type(sK0(X2,X3,X4,X5),set_type) | ~ilf_type(X6,set_type) | ~ilf_type(X7,set_type)) )),
  inference(subsumption_resolution,[],[f602,f170])).
fof(f602,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~member(ordered_pair(X2,X3),compose(X4,X5)) | ~ilf_type(X5,binary_relation_type) | ~ilf_type(X4,binary_relation_type) | ~ilf_type(X3,set_type) | ~ilf_type(X2,set_type) | member(sK0(X2,X3,X4,X5),X6) | ~ilf_type(X4,relation_type(X7,X6)) | ~ilf_type(sK0(X2,X3,X4,X5),set_type) | ~ilf_type(X6,set_type) | ~ilf_type(X7,set_type)) )),
  inference(duplicate_literal_removal,[],[f591])).
fof(f591,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~member(ordered_pair(X2,X3),compose(X4,X5)) | ~ilf_type(X5,binary_relation_type) | ~ilf_type(X4,binary_relation_type) | ~ilf_type(X3,set_type) | ~ilf_type(X2,set_type) | member(sK0(X2,X3,X4,X5),X6) | ~ilf_type(X4,relation_type(X7,X6)) | ~ilf_type(sK0(X2,X3,X4,X5),set_type) | ~ilf_type(X2,set_type) | ~ilf_type(X6,set_type) | ~ilf_type(X7,set_type)) )),
  inference(resolution,[],[f118,f122])).
fof(f122,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~member(ordered_pair(X2,X3),X4) | member(X3,X1) | ~ilf_type(X4,relation_type(X0,X1)) | ~ilf_type(X3,set_type) | ~ilf_type(X2,set_type) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((member(X3,X1) & member(X2,X0)) | ~member(ordered_pair(X2,X3),X4) | ~ilf_type(X4,relation_type(X0,X1))) | ~ilf_type(X3,set_type)) | ~ilf_type(X2,set_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((member(X3,X1) & member(X2,X0)) | ~member(ordered_pair(X2,X3),X4)) | ~ilf_type(X4,relation_type(X0,X1))) | ~ilf_type(X3,set_type)) | ~ilf_type(X2,set_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => ! [X2] : (ilf_type(X2,set_type) => ! [X3] : (ilf_type(X3,set_type) => ! [X4] : (ilf_type(X4,relation_type(X0,X1)) => (member(ordered_pair(X2,X3),X4) => (member(X3,X1) & member(X2,X0))))))))),
  file('SET684+3.tptp',unknown)).
fof(f118,plain,(
  ( ! [X2,X0,X3,X1] : (member(ordered_pair(X0,sK0(X0,X1,X2,X3)),X2) | ~member(ordered_pair(X0,X1),compose(X2,X3)) | ~ilf_type(X3,binary_relation_type) | ~ilf_type(X2,binary_relation_type) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f72])).
fof(f1548,plain,(
  spl19_4 | ~spl19_7 | ~spl19_19 | ~spl19_24),
  inference(avatar_contradiction_clause,[],[f1547])).
fof(f1547,plain,(
  $false | (spl19_4 | ~spl19_7 | ~spl19_19 | ~spl19_24)),
  inference(subsumption_resolution,[],[f1546,f221])).
fof(f221,plain,(
  member(ordered_pair(sK16,sK18),sK14) | ~spl19_7),
  inference(avatar_component_clause,[],[f219])).
fof(f219,plain,(
  spl19_7 <=> member(ordered_pair(sK16,sK18),sK14)),
  introduced(avatar_definition,[new_symbols(naming,[spl19_7])])).
fof(f1546,plain,(
  ~member(ordered_pair(sK16,sK18),sK14) | (spl19_4 | ~spl19_19 | ~spl19_24)),
  inference(subsumption_resolution,[],[f1534,f607])).
fof(f1534,plain,(
  ~ilf_type(sK14,binary_relation_type) | ~member(ordered_pair(sK16,sK18),sK14) | (spl19_4 | ~spl19_24)),
  inference(resolution,[],[f801,f1440])).
fof(f1440,plain,(
  ~member(ordered_pair(sK16,sK17),compose(sK14,sK15)) | spl19_4),
  inference(forward_demodulation,[],[f208,f1322])).
fof(f208,plain,(
  ~member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15)) | spl19_4),
  inference(avatar_component_clause,[],[f206])).
fof(f801,plain,(
  ( ! [X24,X23] : (member(ordered_pair(X23,sK17),compose(X24,sK15)) | ~ilf_type(X24,binary_relation_type) | ~member(ordered_pair(X23,sK18),X24)) ) | ~spl19_24),
  inference(avatar_component_clause,[],[f800])).
fof(f800,plain,(
  spl19_24 <=> ! [X23,X24] : (member(ordered_pair(X23,sK17),compose(X24,sK15)) | ~ilf_type(X24,binary_relation_type) | ~member(ordered_pair(X23,sK18),X24))),
  introduced(avatar_definition,[new_symbols(naming,[spl19_24])])).
fof(f805,plain,(
  spl19_23),
  inference(avatar_contradiction_clause,[],[f804])).
fof(f804,plain,(
  $false | spl19_23),
  inference(subsumption_resolution,[],[f803,f405])).
fof(f405,plain,(
  relation_like(sK15)),
  inference(resolution,[],[f401,f178])).
fof(f401,plain,(
  ( ! [X4,X5,X3] : (~ilf_type(X3,relation_type(X4,X5)) | relation_like(X3)) )),
  inference(subsumption_resolution,[],[f400,f170])).
fof(f400,plain,(
  ( ! [X4,X5,X3] : (~ilf_type(X3,relation_type(X4,X5)) | ~ilf_type(X4,set_type) | relation_like(X3)) )),
  inference(subsumption_resolution,[],[f398,f170])).
fof(f398,plain,(
  ( ! [X4,X5,X3] : (~ilf_type(X3,relation_type(X4,X5)) | ~ilf_type(X5,set_type) | ~ilf_type(X4,set_type) | relation_like(X3)) )),
  inference(duplicate_literal_removal,[],[f397])).
fof(f397,plain,(
  ( ! [X4,X5,X3] : (~ilf_type(X3,relation_type(X4,X5)) | ~ilf_type(X5,set_type) | ~ilf_type(X4,set_type) | relation_like(X3) | ~ilf_type(X5,set_type) | ~ilf_type(X4,set_type)) )),
  inference(resolution,[],[f127,f167])).
fof(f167,plain,(
  ( ! [X2,X0,X1] : (~ilf_type(X2,subset_type(cross_product(X0,X1))) | relation_like(X2) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (! [X1] : (! [X2] : (relation_like(X2) | ~ilf_type(X2,subset_type(cross_product(X0,X1)))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f26])).
fof(f26,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => ! [X2] : (ilf_type(X2,subset_type(cross_product(X0,X1))) => relation_like(X2))))),
  file('SET684+3.tptp',unknown)).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (ilf_type(X2,subset_type(cross_product(X0,X1))) | ~ilf_type(X2,relation_type(X0,X1)) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (! [X1] : ((! [X2] : (ilf_type(X2,subset_type(cross_product(X0,X1))) | ~ilf_type(X2,relation_type(X0,X1))) & ! [X3] : (ilf_type(X3,relation_type(X0,X1)) | ~ilf_type(X3,subset_type(cross_product(X0,X1))))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => (! [X2] : (ilf_type(X2,relation_type(X0,X1)) => ilf_type(X2,subset_type(cross_product(X0,X1)))) & ! [X3] : (ilf_type(X3,subset_type(cross_product(X0,X1))) => ilf_type(X3,relation_type(X0,X1))))))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => (! [X3] : (ilf_type(X3,relation_type(X0,X1)) => ilf_type(X3,subset_type(cross_product(X0,X1)))) & ! [X2] : (ilf_type(X2,subset_type(cross_product(X0,X1))) => ilf_type(X2,relation_type(X0,X1))))))),
  file('SET684+3.tptp',unknown)).
fof(f803,plain,(
  ~relation_like(sK15) | spl19_23),
  inference(resolution,[],[f798,f236])).
fof(f236,plain,(
  ( ! [X0] : (ilf_type(X0,binary_relation_type) | ~relation_like(X0)) )),
  inference(resolution,[],[f193,f170])).
fof(f193,plain,(
  ( ! [X0] : (~ilf_type(X0,set_type) | ilf_type(X0,binary_relation_type) | ~relation_like(X0)) )),
  inference(duplicate_literal_removal,[],[f142])).
fof(f142,plain,(
  ( ! [X0] : (ilf_type(X0,binary_relation_type) | ~ilf_type(X0,set_type) | ~relation_like(X0) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (((ilf_type(X0,binary_relation_type) | ~ilf_type(X0,set_type) | ~relation_like(X0)) & ((ilf_type(X0,set_type) & relation_like(X0)) | ~ilf_type(X0,binary_relation_type))) | ~ilf_type(X0,set_type))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : (((ilf_type(X0,binary_relation_type) | (~ilf_type(X0,set_type) | ~relation_like(X0))) & ((ilf_type(X0,set_type) & relation_like(X0)) | ~ilf_type(X0,binary_relation_type))) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((ilf_type(X0,binary_relation_type) <=> (ilf_type(X0,set_type) & relation_like(X0))) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => (ilf_type(X0,binary_relation_type) <=> (ilf_type(X0,set_type) & relation_like(X0))))),
  file('SET684+3.tptp',unknown)).
fof(f798,plain,(
  ~ilf_type(sK15,binary_relation_type) | spl19_23),
  inference(avatar_component_clause,[],[f796])).
fof(f802,plain,(
  ~spl19_23 | spl19_24 | ~spl19_6),
  inference(avatar_split_clause,[],[f794,f214,f800,f796])).
fof(f214,plain,(
  spl19_6 <=> member(ordered_pair(sK18,sK17),sK15)),
  introduced(avatar_definition,[new_symbols(naming,[spl19_6])])).
fof(f794,plain,(
  ( ! [X24,X23] : (member(ordered_pair(X23,sK17),compose(X24,sK15)) | ~member(ordered_pair(X23,sK18),X24) | ~ilf_type(sK15,binary_relation_type) | ~ilf_type(X24,binary_relation_type)) ) | ~spl19_6),
  inference(subsumption_resolution,[],[f793,f170])).
fof(f793,plain,(
  ( ! [X24,X23] : (member(ordered_pair(X23,sK17),compose(X24,sK15)) | ~member(ordered_pair(X23,sK18),X24) | ~ilf_type(sK15,binary_relation_type) | ~ilf_type(X24,binary_relation_type) | ~ilf_type(X23,set_type)) ) | ~spl19_6),
  inference(subsumption_resolution,[],[f792,f170])).
fof(f792,plain,(
  ( ! [X24,X23] : (member(ordered_pair(X23,sK17),compose(X24,sK15)) | ~member(ordered_pair(X23,sK18),X24) | ~ilf_type(sK15,binary_relation_type) | ~ilf_type(X24,binary_relation_type) | ~ilf_type(sK17,set_type) | ~ilf_type(X23,set_type)) ) | ~spl19_6),
  inference(subsumption_resolution,[],[f771,f170])).
fof(f771,plain,(
  ( ! [X24,X23] : (member(ordered_pair(X23,sK17),compose(X24,sK15)) | ~member(ordered_pair(X23,sK18),X24) | ~ilf_type(sK18,set_type) | ~ilf_type(sK15,binary_relation_type) | ~ilf_type(X24,binary_relation_type) | ~ilf_type(sK17,set_type) | ~ilf_type(X23,set_type)) ) | ~spl19_6),
  inference(resolution,[],[f120,f216])).
fof(f216,plain,(
  member(ordered_pair(sK18,sK17),sK15) | ~spl19_6),
  inference(avatar_component_clause,[],[f214])).
fof(f120,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~member(ordered_pair(X4,X1),X3) | member(ordered_pair(X0,X1),compose(X2,X3)) | ~member(ordered_pair(X0,X4),X2) | ~ilf_type(X4,set_type) | ~ilf_type(X3,binary_relation_type) | ~ilf_type(X2,binary_relation_type) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f72])).
fof(f640,plain,(
  spl19_19),
  inference(avatar_contradiction_clause,[],[f639])).
fof(f639,plain,(
  $false | spl19_19),
  inference(subsumption_resolution,[],[f638,f404])).
fof(f404,plain,(
  relation_like(sK14)),
  inference(resolution,[],[f401,f177])).
fof(f638,plain,(
  ~relation_like(sK14) | spl19_19),
  inference(resolution,[],[f608,f236])).
fof(f608,plain,(
  ~ilf_type(sK14,binary_relation_type) | spl19_19),
  inference(avatar_component_clause,[],[f606])).
fof(f612,plain,(
  ~spl19_19 | spl19_20 | ~spl19_5),
  inference(avatar_split_clause,[],[f604,f210,f610,f606])).
fof(f210,plain,(
  spl19_5 <=> ! [X7] : (~member(ordered_pair(X7,sK17),sK15) | ~ilf_type(X7,member_type(sK12)) | ~member(ordered_pair(sK16,X7),sK14))),
  introduced(avatar_definition,[new_symbols(naming,[spl19_5])])).
fof(f604,plain,(
  ( ! [X0,X1] : (~member(ordered_pair(sK16,X0),compose(sK14,X1)) | ~ilf_type(X1,binary_relation_type) | ~ilf_type(sK14,binary_relation_type) | ~ilf_type(sK0(sK16,X0,sK14,X1),member_type(sK12)) | ~member(ordered_pair(sK0(sK16,X0,sK14,X1),sK17),sK15)) ) | ~spl19_5),
  inference(subsumption_resolution,[],[f603,f170])).
fof(f603,plain,(
  ( ! [X0,X1] : (~member(ordered_pair(sK16,X0),compose(sK14,X1)) | ~ilf_type(X1,binary_relation_type) | ~ilf_type(sK14,binary_relation_type) | ~ilf_type(sK16,set_type) | ~ilf_type(sK0(sK16,X0,sK14,X1),member_type(sK12)) | ~member(ordered_pair(sK0(sK16,X0,sK14,X1),sK17),sK15)) ) | ~spl19_5),
  inference(subsumption_resolution,[],[f590,f170])).
fof(f590,plain,(
  ( ! [X0,X1] : (~member(ordered_pair(sK16,X0),compose(sK14,X1)) | ~ilf_type(X1,binary_relation_type) | ~ilf_type(sK14,binary_relation_type) | ~ilf_type(X0,set_type) | ~ilf_type(sK16,set_type) | ~ilf_type(sK0(sK16,X0,sK14,X1),member_type(sK12)) | ~member(ordered_pair(sK0(sK16,X0,sK14,X1),sK17),sK15)) ) | ~spl19_5),
  inference(resolution,[],[f118,f211])).
fof(f211,plain,(
  ( ! [X7] : (~member(ordered_pair(sK16,X7),sK14) | ~ilf_type(X7,member_type(sK12)) | ~member(ordered_pair(X7,sK17),sK15)) ) | ~spl19_5),
  inference(avatar_component_clause,[],[f210])).
fof(f222,plain,(
  spl19_4 | spl19_7),
  inference(avatar_split_clause,[],[f182,f219,f206])).
fof(f182,plain,(
  member(ordered_pair(sK16,sK18),sK14) | member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15))),
  inference(cnf_transformation,[],[f116])).
fof(f217,plain,(
  spl19_4 | spl19_6),
  inference(avatar_split_clause,[],[f183,f214,f206])).
fof(f183,plain,(
  member(ordered_pair(sK18,sK17),sK15) | member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15))),
  inference(cnf_transformation,[],[f116])).
fof(f212,plain,(
  ~spl19_4 | spl19_5),
  inference(avatar_split_clause,[],[f184,f210,f206])).
fof(f184,plain,(
  ( ! [X7] : (~member(ordered_pair(X7,sK17),sK15) | ~member(ordered_pair(sK16,X7),sK14) | ~ilf_type(X7,member_type(sK12)) | ~member(ordered_pair(sK16,sK17),compose5(sK11,sK12,sK13,sK14,sK15))) )),
  inference(cnf_transformation,[],[f116])).
% SZS output end Proof for SET684+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6908
% Time elapsed: 0.037 s
% ------------------------------
% ------------------------------
