% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWV731-1
% SZS output start Proof for SWV731-1
fof(f333138,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f987,f1006,f74525,f76826,f77019,f164949,f168507,f201970,f201974,f289320])).
fof(f289320,plain,(
  ~spl0_1184 | spl0_1225),
  inference(avatar_contradiction_clause,[],[f289319])).
fof(f289319,plain,(
  $false | (~spl0_1184 | spl0_1225)),
  inference(subsumption_resolution,[],[f77007,f76821])).
fof(f76821,plain,(
  ( ! [X0] : (c_Public_Okeymode_OEncryption = c_Message_Osynth(X0)) ) | ~spl0_1184),
  inference(avatar_component_clause,[],[f76820])).
fof(f76820,plain,(
  spl0_1184 <=> ! [X0] : c_Public_Okeymode_OEncryption = c_Message_Osynth(X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1184])])).
fof(f77007,plain,(
  c_Public_Okeymode_OEncryption != c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool))) | spl0_1225),
  inference(avatar_component_clause,[],[f77006])).
fof(f77006,plain,(
  spl0_1225 <=> c_Public_Okeymode_OEncryption = c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1225])])).
fof(f201974,plain,(
  spl0_1184 | ~spl0_7 | ~spl0_10 | ~spl0_716 | ~spl0_729 | ~spl0_1185),
  inference(avatar_split_clause,[],[f201973,f76823,f74585,f74523,f985,f971,f76820])).
fof(f971,plain,(
  spl0_7 <=> class_Lattices_Oupper__semilattice(tc_fun(tc_Message_Omsg,tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f985,plain,(
  spl0_10 <=> ! [X23] : c_Message_Osynth(X23) = c_Lattices_Oupper__semilattice__class_Osup(X23,c_Message_Osynth(X23),tc_fun(tc_Message_Omsg,tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f74523,plain,(
  spl0_716 <=> ! [X1] : c_Public_Okeymode_OSignature != c_Message_Osynth(X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_716])])).
fof(f74585,plain,(
  spl0_729 <=> c_Public_Okeymode_OSignature = c_Message_Oanalz(c_Public_Okeymode_OSignature)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_729])])).
fof(f76823,plain,(
  spl0_1185 <=> c_Public_Okeymode_OSignature = c_Message_Osynth(c_Public_Okeymode_OSignature)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1185])])).
fof(f201973,plain,(
  ( ! [X0] : (c_Public_Okeymode_OEncryption = c_Message_Osynth(X0)) ) | (~spl0_7 | ~spl0_10 | ~spl0_716 | ~spl0_729 | ~spl0_1185)),
  inference(subsumption_resolution,[],[f166095,f74524])).
fof(f74524,plain,(
  ( ! [X1] : (c_Public_Okeymode_OSignature != c_Message_Osynth(X1)) ) | ~spl0_716),
  inference(avatar_component_clause,[],[f74523])).
fof(f166095,plain,(
  ( ! [X0] : (c_Public_Okeymode_OSignature = c_Message_Osynth(c_Message_Oanalz(X0)) | c_Public_Okeymode_OEncryption = c_Message_Osynth(X0)) ) | (~spl0_7 | ~spl0_10 | ~spl0_729 | ~spl0_1185)),
  inference(forward_demodulation,[],[f125716,f76825])).
fof(f76825,plain,(
  c_Public_Okeymode_OSignature = c_Message_Osynth(c_Public_Okeymode_OSignature) | ~spl0_1185),
  inference(avatar_component_clause,[],[f76823])).
fof(f125716,plain,(
  ( ! [X0] : (c_Message_Osynth(c_Message_Oanalz(X0)) = c_Message_Osynth(c_Public_Okeymode_OSignature) | c_Public_Okeymode_OEncryption = c_Message_Osynth(X0)) ) | (~spl0_7 | ~spl0_10 | ~spl0_729)),
  inference(backward_demodulation,[],[f74587,f67494])).
fof(f67494,plain,(
  ( ! [X0] : (c_Message_Osynth(c_Message_Oanalz(X0)) = c_Message_Osynth(c_Message_Oanalz(c_Public_Okeymode_OSignature)) | c_Public_Okeymode_OEncryption = c_Message_Osynth(X0)) ) | (~spl0_7 | ~spl0_10)),
  inference(superposition,[],[f54194,f65103])).
fof(f65103,plain,(
  ( ! [X11] : (c_Public_Okeymode_OSignature = X11 | c_Public_Okeymode_OEncryption = X11) )),
  inference(resolution,[],[f53105,f3970])).
fof(f3970,plain,(
  ( ! [X4,X2,X3] : (~hBOOL(hAPP(c_Set_Oinsert(X3,c_Orderings_Obot__class_Obot(tc_fun(X4,tc_bool)),X4),X2)) | X2 = X3) )),
  inference(resolution,[],[f244,f584])).
fof(f584,axiom,(
  ( ! [X10,X23,X11] : (hBOOL(hAPP(hAPP(c_in(X10),X11),X23)) | ~hBOOL(hAPP(X23,X11))) )),
  file('SWV731-1.tptp',unknown)).
fof(f244,axiom,(
  ( ! [X10,X31,X25] : (~hBOOL(hAPP(hAPP(c_in(X10),X25),c_Set_Oinsert(X31,c_Orderings_Obot__class_Obot(tc_fun(X10,tc_bool)),X10))) | X25 = X31) )),
  file('SWV731-1.tptp',unknown)).
fof(f53105,plain,(
  ( ! [X2,X0,X1] : (hBOOL(hAPP(c_Set_Oinsert(c_Public_Okeymode_OSignature,X1,X2),X0)) | c_Public_Okeymode_OEncryption = X0) )),
  inference(resolution,[],[f24481,f470])).
fof(f470,axiom,(
  ( ! [X2,X12,X10,X11] : (~hBOOL(hAPP(c_Set_Oinsert(X12,X2,X10),X11)) | X11 = X12 | hBOOL(hAPP(X2,X11))) )),
  file('SWV731-1.tptp',unknown)).
fof(f24481,plain,(
  ( ! [X4,X5,X3] : (hBOOL(hAPP(c_Set_Oinsert(c_Public_Okeymode_OEncryption,c_Set_Oinsert(c_Public_Okeymode_OSignature,X3,X4),X4),X5))) )),
  inference(forward_demodulation,[],[f24473,f458])).
fof(f458,axiom,(
  ( ! [X2,X12,X10,X11] : (c_Set_Oinsert(X11,c_Set_Oinsert(X12,X2,X10),X10) = c_Set_Oinsert(X12,c_Set_Oinsert(X11,X2,X10),X10)) )),
  file('SWV731-1.tptp',unknown)).
fof(f24473,plain,(
  ( ! [X4,X5,X3] : (hBOOL(hAPP(c_Set_Oinsert(c_Public_Okeymode_OSignature,c_Set_Oinsert(c_Public_Okeymode_OEncryption,X3,X4),X4),X5))) )),
  inference(resolution,[],[f933,f3389])).
fof(f3389,plain,(
  ( ! [X54,X56,X55,X53] : (hBOOL(hAPP(c_Set_Oinsert(X54,c_Set_Oinsert(X53,X55,X56),X56),X53))) )),
  inference(superposition,[],[f484,f458])).
fof(f484,axiom,(
  ( ! [X2,X10,X11] : (hBOOL(hAPP(c_Set_Oinsert(X11,X2,X10),X11))) )),
  file('SWV731-1.tptp',unknown)).
fof(f933,plain,(
  ( ! [X2,X0,X1] : (~hBOOL(hAPP(c_Set_Oinsert(c_Public_Okeymode_OSignature,X0,X1),c_Public_Okeymode_OEncryption)) | hBOOL(hAPP(c_Set_Oinsert(c_Public_Okeymode_OSignature,X0,X1),X2))) )),
  inference(resolution,[],[f15,f484])).
fof(f15,axiom,(
  ( ! [X8,X7] : (~hBOOL(hAPP(X7,c_Public_Okeymode_OSignature)) | ~hBOOL(hAPP(X7,c_Public_Okeymode_OEncryption)) | hBOOL(hAPP(X7,X8))) )),
  file('SWV731-1.tptp',unknown)).
fof(f54194,plain,(
  ( ! [X6] : (c_Message_Osynth(c_Message_Oanalz(X6)) = c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X6)))) ) | (~spl0_7 | ~spl0_10)),
  inference(backward_demodulation,[],[f54193,f23487])).
fof(f23487,plain,(
  ( ! [X6] : (c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X6))) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X6))),c_Message_Osynth(c_Message_Oanalz(X6)),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(resolution,[],[f23475,f89])).
fof(f89,axiom,(
  ( ! [X4,X2,X10] : (~c_lessequals(X4,X2,tc_fun(X10,tc_bool)) | c_Lattices_Oupper__semilattice__class_Osup(X2,X4,tc_fun(X10,tc_bool)) = X2) )),
  file('SWV731-1.tptp',unknown)).
fof(f23475,plain,(
  ( ! [X25] : (c_lessequals(c_Message_Osynth(c_Message_Oanalz(X25)),c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X25))),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(backward_demodulation,[],[f23448,f6211])).
fof(f6211,plain,(
  ( ! [X25] : (c_lessequals(c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(X25),c_Message_Osynth(c_Message_Oanalz(X25)),tc_fun(tc_Message_Omsg,tc_bool)),c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X25))),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(forward_demodulation,[],[f6210,f123])).
fof(f123,axiom,(
  ( ! [X4,X2,X10] : (c_Lattices_Oupper__semilattice__class_Osup(X2,X4,tc_fun(X10,tc_bool)) = c_Lattices_Oupper__semilattice__class_Osup(X4,X2,tc_fun(X10,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f6210,plain,(
  ( ! [X25] : (c_lessequals(c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(c_Message_Oanalz(X25)),c_Message_Osynth(X25),tc_fun(tc_Message_Omsg,tc_bool)),c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X25))),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(forward_demodulation,[],[f6192,f446])).
fof(f446,axiom,(
  ( ! [X1] : (c_Message_Osynth(X1) = c_Message_Osynth(c_Message_Osynth(X1))) )),
  file('SWV731-1.tptp',unknown)).
fof(f6192,plain,(
  ( ! [X25] : (c_lessequals(c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(c_Message_Oanalz(X25)),c_Message_Osynth(c_Message_Osynth(X25)),tc_fun(tc_Message_Omsg,tc_bool)),c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X25))),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(superposition,[],[f69,f291])).
fof(f291,axiom,(
  ( ! [X1] : (c_Message_Oanalz(c_Message_Osynth(X1)) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(X1),c_Message_Osynth(X1),tc_fun(tc_Message_Omsg,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f69,axiom,(
  ( ! [X1,X16] : (c_lessequals(c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(X16),c_Message_Osynth(X1),tc_fun(tc_Message_Omsg,tc_bool)),c_Message_Osynth(c_Lattices_Oupper__semilattice__class_Osup(X16,X1,tc_fun(tc_Message_Omsg,tc_bool))),tc_fun(tc_Message_Omsg,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f23448,plain,(
  ( ! [X5] : (c_Message_Osynth(c_Message_Oanalz(X5)) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(X5),c_Message_Osynth(c_Message_Oanalz(X5)),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(resolution,[],[f23351,f88])).
fof(f88,axiom,(
  ( ! [X4,X2,X10] : (~c_lessequals(X2,X4,tc_fun(X10,tc_bool)) | c_Lattices_Oupper__semilattice__class_Osup(X2,X4,tc_fun(X10,tc_bool)) = X4) )),
  file('SWV731-1.tptp',unknown)).
fof(f23351,plain,(
  ( ! [X47] : (c_lessequals(c_Message_Osynth(X47),c_Message_Osynth(c_Message_Oanalz(X47)),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(superposition,[],[f6150,f2084])).
fof(f2084,plain,(
  ( ! [X18] : (c_Message_Oanalz(X18) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(X18),X18,tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(resolution,[],[f89,f68])).
fof(f68,axiom,(
  ( ! [X1] : (c_lessequals(X1,c_Message_Oanalz(X1),tc_fun(tc_Message_Omsg,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f6150,plain,(
  ( ! [X0,X1] : (c_lessequals(c_Message_Osynth(X0),c_Message_Osynth(c_Lattices_Oupper__semilattice__class_Osup(X1,X0,tc_fun(tc_Message_Omsg,tc_bool))),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(resolution,[],[f69,f212])).
fof(f212,axiom,(
  ( ! [X4,X2,X10,X9] : (~c_lessequals(c_Lattices_Oupper__semilattice__class_Osup(X2,X4,tc_fun(X10,tc_bool)),X9,tc_fun(X10,tc_bool)) | c_lessequals(X4,X9,tc_fun(X10,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f54193,plain,(
  ( ! [X6] : (c_Message_Osynth(c_Message_Oanalz(X6)) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X6))),c_Message_Osynth(c_Message_Oanalz(X6)),tc_fun(tc_Message_Omsg,tc_bool))) ) | (~spl0_7 | ~spl0_10)),
  inference(forward_demodulation,[],[f54142,f446])).
fof(f54142,plain,(
  ( ! [X6] : (c_Message_Osynth(c_Message_Osynth(c_Message_Oanalz(X6))) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(c_Message_Oanalz(c_Message_Osynth(X6))),c_Message_Osynth(c_Message_Osynth(c_Message_Oanalz(X6))),tc_fun(tc_Message_Omsg,tc_bool))) ) | (~spl0_7 | ~spl0_10)),
  inference(resolution,[],[f29378,f2025])).
fof(f2025,plain,(
  ( ! [X4,X5] : (~c_lessequals(X4,X5,tc_fun(tc_Message_Omsg,tc_bool)) | c_Message_Osynth(X5) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(X4),c_Message_Osynth(X5),tc_fun(tc_Message_Omsg,tc_bool))) ) | ~spl0_7),
  inference(subsumption_resolution,[],[f2020,f972])).
fof(f972,plain,(
  class_Lattices_Oupper__semilattice(tc_fun(tc_Message_Omsg,tc_bool)) | ~spl0_7),
  inference(avatar_component_clause,[],[f971])).
fof(f2020,plain,(
  ( ! [X4,X5] : (~c_lessequals(X4,X5,tc_fun(tc_Message_Omsg,tc_bool)) | c_Message_Osynth(X5) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(X4),c_Message_Osynth(X5),tc_fun(tc_Message_Omsg,tc_bool)) | ~class_Lattices_Oupper__semilattice(tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(resolution,[],[f43,f87])).
fof(f87,axiom,(
  ( ! [X12,X10,X11] : (~c_lessequals(X11,X12,X10) | c_Lattices_Oupper__semilattice__class_Osup(X11,X12,X10) = X12 | ~class_Lattices_Oupper__semilattice(X10)) )),
  file('SWV731-1.tptp',unknown)).
fof(f43,axiom,(
  ( ! [X1,X16] : (c_lessequals(c_Message_Osynth(X16),c_Message_Osynth(X1),tc_fun(tc_Message_Omsg,tc_bool)) | ~c_lessequals(X16,X1,tc_fun(tc_Message_Omsg,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f29378,plain,(
  ( ! [X23] : (c_lessequals(c_Message_Oanalz(c_Message_Osynth(X23)),c_Message_Osynth(c_Message_Oanalz(X23)),tc_fun(tc_Message_Omsg,tc_bool))) ) | (~spl0_7 | ~spl0_10)),
  inference(forward_demodulation,[],[f29278,f1143])).
fof(f1143,plain,(
  ( ! [X1] : (c_Message_Osynth(c_Message_Oanalz(X1)) = c_Message_Oanalz(c_Message_Osynth(c_Message_Oanalz(X1)))) ) | ~spl0_10),
  inference(forward_demodulation,[],[f1130,f986])).
fof(f986,plain,(
  ( ! [X23] : (c_Message_Osynth(X23) = c_Lattices_Oupper__semilattice__class_Osup(X23,c_Message_Osynth(X23),tc_fun(tc_Message_Omsg,tc_bool))) ) | ~spl0_10),
  inference(avatar_component_clause,[],[f985])).
fof(f1130,plain,(
  ( ! [X1] : (c_Message_Oanalz(c_Message_Osynth(c_Message_Oanalz(X1))) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(X1),c_Message_Osynth(c_Message_Oanalz(X1)),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(superposition,[],[f291,f497])).
fof(f497,axiom,(
  ( ! [X1] : (c_Message_Oanalz(X1) = c_Message_Oanalz(c_Message_Oanalz(X1))) )),
  file('SWV731-1.tptp',unknown)).
fof(f29278,plain,(
  ( ! [X23] : (c_lessequals(c_Message_Oanalz(c_Message_Osynth(X23)),c_Message_Oanalz(c_Message_Osynth(c_Message_Oanalz(X23))),tc_fun(tc_Message_Omsg,tc_bool))) ) | (~spl0_7 | ~spl0_10)),
  inference(superposition,[],[f1598,f23478])).
fof(f23478,plain,(
  ( ! [X27] : (c_Message_Osynth(c_Message_Oanalz(X27)) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(c_Message_Osynth(X27)),c_Message_Osynth(c_Message_Oanalz(X27)),tc_fun(tc_Message_Omsg,tc_bool))) ) | ~spl0_10),
  inference(forward_demodulation,[],[f23474,f1143])).
fof(f23474,plain,(
  ( ! [X27] : (c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(c_Message_Osynth(X27)),c_Message_Osynth(c_Message_Oanalz(X27)),tc_fun(tc_Message_Omsg,tc_bool)) = c_Message_Oanalz(c_Message_Osynth(c_Message_Oanalz(X27)))) )),
  inference(backward_demodulation,[],[f23448,f8811])).
fof(f8811,plain,(
  ( ! [X27] : (c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(c_Message_Osynth(X27)),c_Message_Osynth(c_Message_Oanalz(X27)),tc_fun(tc_Message_Omsg,tc_bool)) = c_Message_Oanalz(c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(X27),c_Message_Osynth(c_Message_Oanalz(X27)),tc_fun(tc_Message_Omsg,tc_bool)))) )),
  inference(forward_demodulation,[],[f8810,f123])).
fof(f8810,plain,(
  ( ! [X27] : (c_Message_Oanalz(c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(c_Message_Oanalz(X27)),c_Message_Osynth(X27),tc_fun(tc_Message_Omsg,tc_bool))) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(c_Message_Osynth(X27)),c_Message_Osynth(c_Message_Oanalz(X27)),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(forward_demodulation,[],[f8734,f497])).
fof(f8734,plain,(
  ( ! [X27] : (c_Message_Oanalz(c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(c_Message_Oanalz(X27)),c_Message_Osynth(X27),tc_fun(tc_Message_Omsg,tc_bool))) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(c_Message_Oanalz(c_Message_Osynth(X27))),c_Message_Osynth(c_Message_Oanalz(X27)),tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(superposition,[],[f292,f291])).
fof(f292,axiom,(
  ( ! [X1,X16] : (c_Message_Oanalz(c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(X16),X1,tc_fun(tc_Message_Omsg,tc_bool))) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(c_Lattices_Oupper__semilattice__class_Osup(X16,X1,tc_fun(tc_Message_Omsg,tc_bool))),c_Message_Osynth(X16),tc_fun(tc_Message_Omsg,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f1598,plain,(
  ( ! [X33,X34] : (c_lessequals(X33,c_Message_Oanalz(c_Lattices_Oupper__semilattice__class_Osup(X33,X34,tc_fun(tc_Message_Omsg,tc_bool))),tc_fun(tc_Message_Omsg,tc_bool))) ) | ~spl0_7),
  inference(subsumption_resolution,[],[f1566,f972])).
fof(f1566,plain,(
  ( ! [X33,X34] : (c_lessequals(X33,c_Message_Oanalz(c_Lattices_Oupper__semilattice__class_Osup(X33,X34,tc_fun(tc_Message_Omsg,tc_bool))),tc_fun(tc_Message_Omsg,tc_bool)) | ~class_Lattices_Oupper__semilattice(tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(resolution,[],[f206,f68])).
fof(f206,axiom,(
  ( ! [X12,X10,X13,X11] : (~c_lessequals(c_Lattices_Oupper__semilattice__class_Osup(X11,X12,X10),X13,X10) | c_lessequals(X11,X13,X10) | ~class_Lattices_Oupper__semilattice(X10)) )),
  file('SWV731-1.tptp',unknown)).
fof(f74587,plain,(
  c_Public_Okeymode_OSignature = c_Message_Oanalz(c_Public_Okeymode_OSignature) | ~spl0_729),
  inference(avatar_component_clause,[],[f74585])).
fof(f201970,plain,(
  spl0_1254 | ~spl0_10 | ~spl0_716),
  inference(avatar_split_clause,[],[f201969,f74523,f985,f77173])).
fof(f77173,plain,(
  spl0_1254 <=> ! [X0] : c_Public_Okeymode_OEncryption = c_Message_Osynth(c_Message_Oanalz(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1254])])).
fof(f201969,plain,(
  ( ! [X27] : (c_Public_Okeymode_OEncryption = c_Message_Osynth(c_Message_Oanalz(X27))) ) | (~spl0_10 | ~spl0_716)),
  inference(backward_demodulation,[],[f201968,f23478])).
fof(f201968,plain,(
  ( ! [X0] : (c_Public_Okeymode_OEncryption = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(c_Message_Osynth(X0)),c_Message_Osynth(c_Message_Oanalz(X0)),tc_fun(tc_Message_Omsg,tc_bool))) ) | (~spl0_10 | ~spl0_716)),
  inference(subsumption_resolution,[],[f68076,f74524])).
fof(f68076,plain,(
  ( ! [X0] : (c_Public_Okeymode_OSignature = c_Message_Osynth(c_Message_Oanalz(X0)) | c_Public_Okeymode_OEncryption = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Oanalz(c_Message_Osynth(X0)),c_Message_Osynth(c_Message_Oanalz(X0)),tc_fun(tc_Message_Omsg,tc_bool))) ) | ~spl0_10),
  inference(superposition,[],[f23478,f65103])).
fof(f168507,plain,(
  ~spl0_1225 | ~spl0_1254),
  inference(avatar_contradiction_clause,[],[f168506])).
fof(f168506,plain,(
  $false | (~spl0_1225 | ~spl0_1254)),
  inference(subsumption_resolution,[],[f168505,f166770])).
fof(f166770,plain,(
  ( ! [X32] : (hBOOL(hAPP(c_Public_Okeymode_OEncryption,c_Message_Omsg_ONonce(X32)))) ) | ~spl0_1254),
  inference(backward_demodulation,[],[f77174,f15927])).
fof(f15927,plain,(
  ( ! [X33,X32] : (hBOOL(hAPP(c_Message_Osynth(c_Message_Oanalz(c_Set_Oinsert(c_Message_Omsg_ONonce(X32),X33,tc_Message_Omsg))),c_Message_Omsg_ONonce(X32)))) )),
  inference(superposition,[],[f15156,f286])).
fof(f286,axiom,(
  ( ! [X48,X1] : (c_Message_Oanalz(c_Set_Oinsert(c_Message_Omsg_ONonce(X48),X1,tc_Message_Omsg)) = c_Set_Oinsert(c_Message_Omsg_ONonce(X48),c_Message_Oanalz(X1),tc_Message_Omsg)) )),
  file('SWV731-1.tptp',unknown)).
fof(f15156,plain,(
  ( ! [X39,X38] : (hBOOL(hAPP(c_Message_Osynth(c_Set_Oinsert(X38,X39,tc_Message_Omsg)),X38))) )),
  inference(resolution,[],[f3895,f585])).
fof(f585,axiom,(
  ( ! [X10,X23,X11] : (~hBOOL(hAPP(hAPP(c_in(X10),X11),X23)) | hBOOL(hAPP(X23,X11))) )),
  file('SWV731-1.tptp',unknown)).
fof(f3895,plain,(
  ( ! [X4,X5] : (hBOOL(hAPP(hAPP(c_in(tc_Message_Omsg),X4),c_Message_Osynth(c_Set_Oinsert(X4,X5,tc_Message_Omsg))))) )),
  inference(resolution,[],[f242,f293])).
fof(f293,axiom,(
  ( ! [X0,X1] : (c_lessequals(c_Set_Oinsert(X0,c_Message_Osynth(X1),tc_Message_Omsg),c_Message_Osynth(c_Set_Oinsert(X0,X1,tc_Message_Omsg)),tc_fun(tc_Message_Omsg,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f242,axiom,(
  ( ! [X4,X2,X10,X11] : (~c_lessequals(c_Set_Oinsert(X11,X2,X10),X4,tc_fun(X10,tc_bool)) | hBOOL(hAPP(hAPP(c_in(X10),X11),X4))) )),
  file('SWV731-1.tptp',unknown)).
fof(f77174,plain,(
  ( ! [X0] : (c_Public_Okeymode_OEncryption = c_Message_Osynth(c_Message_Oanalz(X0))) ) | ~spl0_1254),
  inference(avatar_component_clause,[],[f77173])).
fof(f168505,plain,(
  ( ! [X23] : (~hBOOL(hAPP(c_Public_Okeymode_OEncryption,c_Message_Omsg_ONonce(X23)))) ) | ~spl0_1225),
  inference(forward_demodulation,[],[f28457,f77008])).
fof(f77008,plain,(
  c_Public_Okeymode_OEncryption = c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool))) | ~spl0_1225),
  inference(avatar_component_clause,[],[f77006])).
fof(f28457,plain,(
  ( ! [X23] : (~hBOOL(hAPP(c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool))),c_Message_Omsg_ONonce(X23)))) )),
  inference(resolution,[],[f4825,f143])).
fof(f143,axiom,(
  ( ! [X10,X11] : (~hBOOL(hAPP(hAPP(c_in(X10),X11),c_Orderings_Obot__class_Obot(tc_fun(X10,tc_bool))))) )),
  file('SWV731-1.tptp',unknown)).
fof(f4825,plain,(
  ( ! [X2,X3] : (hBOOL(hAPP(hAPP(c_in(tc_Message_Omsg),c_Message_Omsg_ONonce(X2)),X3)) | ~hBOOL(hAPP(c_Message_Osynth(X3),c_Message_Omsg_ONonce(X2)))) )),
  inference(resolution,[],[f268,f584])).
fof(f268,axiom,(
  ( ! [X47,X1] : (~hBOOL(hAPP(hAPP(c_in(tc_Message_Omsg),c_Message_Omsg_ONonce(X47)),c_Message_Osynth(X1))) | hBOOL(hAPP(hAPP(c_in(tc_Message_Omsg),c_Message_Omsg_ONonce(X47)),X1))) )),
  file('SWV731-1.tptp',unknown)).
fof(f164949,plain,(
  ~spl0_711),
  inference(avatar_contradiction_clause,[],[f164948])).
fof(f164948,plain,(
  $false | ~spl0_711),
  inference(subsumption_resolution,[],[f164947,f160430])).
fof(f160430,plain,(
  ( ! [X26,X27] : (hBOOL(hAPP(c_Public_Okeymode_OSignature,hAPP(X26,X27)))) ) | ~spl0_711),
  inference(backward_demodulation,[],[f158906,f15917])).
fof(f15917,plain,(
  ( ! [X28,X26,X29,X27] : (hBOOL(hAPP(c_Message_Osynth(c_Set_Oimage(X26,c_Set_Oinsert(X27,X28,X29),X29,tc_Message_Omsg)),hAPP(X26,X27)))) )),
  inference(superposition,[],[f15156,f410])).
fof(f410,axiom,(
  ( ! [X4,X14,X10,X31,X15] : (c_Set_Oinsert(hAPP(X14,X31),c_Set_Oimage(X14,X4,X15,X10),X10) = c_Set_Oimage(X14,c_Set_Oinsert(X31,X4,X15),X15,X10)) )),
  file('SWV731-1.tptp',unknown)).
fof(f158906,plain,(
  ( ! [X0] : (c_Public_Okeymode_OSignature = c_Message_Osynth(X0)) ) | ~spl0_711),
  inference(subsumption_resolution,[],[f67690,f153529])).
fof(f153529,plain,(
  ( ! [X3] : (c_Public_Okeymode_OEncryption != c_Message_Osynth(X3)) ) | ~spl0_711),
  inference(forward_demodulation,[],[f2002,f74502])).
fof(f74502,plain,(
  ( ! [X0] : (c_Public_Okeymode_OEncryption = c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool))) ) | ~spl0_711),
  inference(avatar_component_clause,[],[f74501])).
fof(f74501,plain,(
  spl0_711 <=> ! [X0] : c_Public_Okeymode_OEncryption = c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_711])])).
fof(f2002,plain,(
  ( ! [X3] : (c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool)) != c_Message_Osynth(X3)) )),
  inference(superposition,[],[f46,f1859])).
fof(f1859,plain,(
  ( ! [X8,X9] : (c_Message_Osynth(X9) = c_Set_Oinsert(c_Message_Omsg_OAgent(X8),c_Message_Osynth(X9),tc_Message_Omsg)) )),
  inference(resolution,[],[f422,f516])).
fof(f516,axiom,(
  ( ! [X67,X1] : (hBOOL(hAPP(hAPP(c_in(tc_Message_Omsg),c_Message_Omsg_OAgent(X67)),c_Message_Osynth(X1)))) )),
  file('SWV731-1.tptp',unknown)).
fof(f422,axiom,(
  ( ! [X2,X10,X31] : (~hBOOL(hAPP(hAPP(c_in(X10),X31),X2)) | c_Set_Oinsert(X31,X2,X10) = X2) )),
  file('SWV731-1.tptp',unknown)).
fof(f46,axiom,(
  ( ! [X2,X10,X31] : (c_Orderings_Obot__class_Obot(tc_fun(X10,tc_bool)) != c_Set_Oinsert(X31,X2,X10)) )),
  file('SWV731-1.tptp',unknown)).
fof(f67690,plain,(
  ( ! [X0] : (c_Public_Okeymode_OSignature = c_Message_Osynth(X0) | c_Public_Okeymode_OEncryption = c_Message_Osynth(c_Message_Osynth(X0))) )),
  inference(superposition,[],[f446,f65103])).
fof(f164947,plain,(
  ( ! [X36] : (~hBOOL(hAPP(c_Public_Okeymode_OSignature,hAPP(c_Message_Omsg_OKey,X36)))) ) | ~spl0_711),
  inference(forward_demodulation,[],[f44004,f158906])).
fof(f44004,plain,(
  ( ! [X36] : (~hBOOL(hAPP(c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool))),hAPP(c_Message_Omsg_OKey,X36)))) )),
  inference(resolution,[],[f6728,f143])).
fof(f6728,plain,(
  ( ! [X2,X3] : (hBOOL(hAPP(hAPP(c_in(tc_Message_Omsg),hAPP(c_Message_Omsg_OKey,X2)),X3)) | ~hBOOL(hAPP(c_Message_Osynth(X3),hAPP(c_Message_Omsg_OKey,X2)))) )),
  inference(resolution,[],[f509,f584])).
fof(f509,axiom,(
  ( ! [X33,X1] : (~hBOOL(hAPP(hAPP(c_in(tc_Message_Omsg),hAPP(c_Message_Omsg_OKey,X33)),c_Message_Osynth(X1))) | hBOOL(hAPP(hAPP(c_in(tc_Message_Omsg),hAPP(c_Message_Omsg_OKey,X33)),X1))) )),
  file('SWV731-1.tptp',unknown)).
fof(f77019,plain,(
  spl0_1225 | spl0_729),
  inference(avatar_split_clause,[],[f67506,f74585,f77006])).
fof(f67506,plain,(
  c_Public_Okeymode_OSignature = c_Message_Oanalz(c_Public_Okeymode_OSignature) | c_Public_Okeymode_OEncryption = c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool)))),
  inference(superposition,[],[f1144,f65103])).
fof(f1144,plain,(
  c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool))) = c_Message_Oanalz(c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool))))),
  inference(forward_demodulation,[],[f1131,f64])).
fof(f64,axiom,(
  ( ! [X4,X10] : (c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Obot__class_Obot(tc_fun(X10,tc_bool)),X4,tc_fun(X10,tc_bool)) = X4) )),
  file('SWV731-1.tptp',unknown)).
fof(f1131,plain,(
  c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool)),c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool))),tc_fun(tc_Message_Omsg,tc_bool)) = c_Message_Oanalz(c_Message_Osynth(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool))))),
  inference(superposition,[],[f291,f8])).
fof(f8,axiom,(
  c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool)) = c_Message_Oanalz(c_Orderings_Obot__class_Obot(tc_fun(tc_Message_Omsg,tc_bool)))),
  file('SWV731-1.tptp',unknown)).
fof(f76826,plain,(
  spl0_1184 | spl0_1185),
  inference(avatar_split_clause,[],[f67373,f76823,f76820])).
fof(f67373,plain,(
  ( ! [X0] : (c_Public_Okeymode_OSignature = c_Message_Osynth(c_Public_Okeymode_OSignature) | c_Public_Okeymode_OEncryption = c_Message_Osynth(X0)) )),
  inference(superposition,[],[f446,f65103])).
fof(f74525,plain,(
  spl0_711 | spl0_716),
  inference(avatar_split_clause,[],[f65828,f74523,f74501])).
fof(f65828,plain,(
  ( ! [X0,X1] : (c_Public_Okeymode_OSignature != c_Message_Osynth(X1) | c_Public_Okeymode_OEncryption = c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool))) )),
  inference(superposition,[],[f3502,f65103])).
fof(f3502,plain,(
  ( ! [X15,X16] : (c_Orderings_Obot__class_Obot(tc_fun(X16,tc_bool)) != c_Message_Osynth(X15)) )),
  inference(superposition,[],[f46,f2634])).
fof(f2634,plain,(
  ( ! [X28,X29,X27] : (c_Message_Osynth(X28) = c_Set_Oinsert(c_Message_Omsg_OAgent(X27),c_Message_Osynth(X28),X29)) )),
  inference(resolution,[],[f1856,f924])).
fof(f924,plain,(
  ( ! [X6,X5] : (hBOOL(hAPP(c_Message_Osynth(X5),c_Message_Omsg_OAgent(X6)))) )),
  inference(resolution,[],[f585,f516])).
fof(f1856,plain,(
  ( ! [X2,X0,X1] : (~hBOOL(hAPP(X1,X0)) | c_Set_Oinsert(X0,X1,X2) = X1) )),
  inference(resolution,[],[f422,f584])).
fof(f1006,plain,(
  spl0_7),
  inference(avatar_contradiction_clause,[],[f1005])).
fof(f1005,plain,(
  $false | spl0_7),
  inference(subsumption_resolution,[],[f1004,f628])).
fof(f628,axiom,(
  class_Lattices_Olattice(tc_bool)),
  file('SWV731-1.tptp',unknown)).
fof(f1004,plain,(
  ~class_Lattices_Olattice(tc_bool) | spl0_7),
  inference(resolution,[],[f973,f607])).
fof(f607,axiom,(
  ( ! [X80,X79] : (class_Lattices_Oupper__semilattice(tc_fun(X79,X80)) | ~class_Lattices_Olattice(X80)) )),
  file('SWV731-1.tptp',unknown)).
fof(f973,plain,(
  ~class_Lattices_Oupper__semilattice(tc_fun(tc_Message_Omsg,tc_bool)) | spl0_7),
  inference(avatar_component_clause,[],[f971])).
fof(f987,plain,(
  ~spl0_7 | spl0_10),
  inference(avatar_split_clause,[],[f983,f985,f971])).
fof(f983,plain,(
  ( ! [X23] : (c_Message_Osynth(X23) = c_Lattices_Oupper__semilattice__class_Osup(X23,c_Message_Osynth(X23),tc_fun(tc_Message_Omsg,tc_bool)) | ~class_Lattices_Oupper__semilattice(tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(forward_demodulation,[],[f949,f123])).
fof(f949,plain,(
  ( ! [X23] : (c_Message_Osynth(X23) = c_Lattices_Oupper__semilattice__class_Osup(c_Message_Osynth(X23),X23,tc_fun(tc_Message_Omsg,tc_bool)) | ~class_Lattices_Oupper__semilattice(tc_fun(tc_Message_Omsg,tc_bool))) )),
  inference(resolution,[],[f85,f204])).
fof(f204,axiom,(
  ( ! [X1] : (c_lessequals(X1,c_Message_Osynth(X1),tc_fun(tc_Message_Omsg,tc_bool))) )),
  file('SWV731-1.tptp',unknown)).
fof(f85,axiom,(
  ( ! [X12,X10,X11] : (~c_lessequals(X12,X11,X10) | c_Lattices_Oupper__semilattice__class_Osup(X11,X12,X10) = X11 | ~class_Lattices_Oupper__semilattice(X10)) )),
  file('SWV731-1.tptp',unknown)).
% SZS output end Proof for SWV731-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 188696
% Time elapsed: 5.112 s
% ------------------------------
% ------------------------------
