% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SYN231-1
% SZS output start Proof for SYN231-1
fof(f1639,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f544,f697,f1042,f1433,f1492,f1638])).
fof(f1638,plain,(
  ~spl0_39 | ~spl0_156),
  inference(avatar_contradiction_clause,[],[f1637])).
fof(f1637,plain,(
  $false | (~spl0_39 | ~spl0_156)),
  inference(subsumption_resolution,[],[f1635,f1588])).
fof(f1588,plain,(
  ~n2(d) | ~spl0_156),
  inference(resolution,[],[f1041,f369])).
fof(f369,axiom,(
  ( ! [X0] : (~l3(d,X0)) )),
  file('SYN231-1.tptp',unknown)).
fof(f1041,plain,(
  ( ! [X6] : (l3(X6,X6) | ~n2(X6)) ) | ~spl0_156),
  inference(avatar_component_clause,[],[f1040])).
fof(f1040,plain,(
  spl0_156 <=> ! [X6] : (~n2(X6) | l3(X6,X6))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_156])])).
fof(f1635,plain,(
  n2(d) | ~spl0_39),
  inference(resolution,[],[f1622,f1])).
fof(f1,axiom,(
  s0(d)),
  file('SYN231-1.tptp',unknown)).
fof(f1622,plain,(
  ( ! [X0] : (~s0(X0) | n2(X0)) ) | ~spl0_39),
  inference(resolution,[],[f543,f175])).
fof(f175,axiom,(
  ( ! [X6,X10,X11] : (~p1(X11,X6,X10) | n2(X10)) )),
  file('SYN231-1.tptp',unknown)).
fof(f543,plain,(
  ( ! [X9] : (p1(X9,X9,X9) | ~s0(X9)) ) | ~spl0_39),
  inference(avatar_component_clause,[],[f542])).
fof(f542,plain,(
  spl0_39 <=> ! [X9] : (~s0(X9) | p1(X9,X9,X9))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_39])])).
fof(f1492,plain,(
  spl0_73),
  inference(avatar_split_clause,[],[f1463,f694])).
fof(f694,plain,(
  spl0_73 <=> k1(b)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_73])])).
fof(f1463,plain,(
  k1(b)),
  inference(resolution,[],[f39,f7])).
fof(f7,axiom,(
  n0(d,b)),
  file('SYN231-1.tptp',unknown)).
fof(f39,axiom,(
  ( ! [X2,X3] : (~n0(X3,X2) | k1(X2)) )),
  file('SYN231-1.tptp',unknown)).
fof(f1433,plain,(
  ~spl0_38),
  inference(avatar_contradiction_clause,[],[f1430])).
fof(f1430,plain,(
  $false | ~spl0_38),
  inference(resolution,[],[f540,f1])).
fof(f540,plain,(
  ( ! [X4] : (~s0(X4)) ) | ~spl0_38),
  inference(avatar_component_clause,[],[f539])).
fof(f539,plain,(
  spl0_38 <=> ! [X4] : ~s0(X4)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_38])])).
fof(f1042,plain,(
  spl0_156 | ~spl0_72),
  inference(avatar_split_clause,[],[f255,f690,f1040])).
fof(f690,plain,(
  spl0_72 <=> m2(b)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_72])])).
fof(f255,axiom,(
  ( ! [X6] : (~m2(b) | ~n2(X6) | l3(X6,X6)) )),
  file('SYN231-1.tptp',unknown)).
fof(f697,plain,(
  spl0_72 | ~spl0_73),
  inference(avatar_split_clause,[],[f174,f694,f690])).
fof(f174,axiom,(
  ~k1(b) | m2(b)),
  file('SYN231-1.tptp',unknown)).
fof(f544,plain,(
  spl0_38 | spl0_39),
  inference(avatar_split_clause,[],[f110,f542,f539])).
fof(f110,axiom,(
  ( ! [X4,X9] : (~s0(X9) | ~s0(X4) | p1(X9,X9,X9)) )),
  file('SYN231-1.tptp',unknown)).
% SZS output end Proof for SYN231-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5884
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
