% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET254-6
% SZS output start Proof for SET254-6
fof(f112536,plain,(
  $false),
  inference(subsumption_resolution,[],[f112445,f113])).
fof(f113,axiom,(
  subclass(x1,x2)),
  file('SET254-6.tptp',unknown)).
fof(f112445,plain,(
  ~subclass(x1,x2)),
  inference(resolution,[],[f24390,f177])).
fof(f177,plain,(
  ~subclass(intersection(x1,cross_product(y,z)),intersection(x2,cross_product(y,z)))),
  inference(definition_unfolding,[],[f114,f28,f28])).
fof(f28,axiom,(
  ( ! [X0,X5,X1] : (intersection(X5,cross_product(X0,X1)) = restrict(X5,X0,X1)) )),
  file('SET254-6.tptp',unknown)).
fof(f114,axiom,(
  ~subclass(restrict(x1,y,z),restrict(x2,y,z))),
  file('SET254-6.tptp',unknown)).
fof(f24390,plain,(
  ( ! [X30,X31,X32] : (subclass(intersection(X30,X31),intersection(X32,X31)) | ~subclass(X30,X32)) )),
  inference(duplicate_literal_removal,[],[f24347])).
fof(f24347,plain,(
  ( ! [X30,X31,X32] : (subclass(intersection(X30,X31),intersection(X32,X31)) | subclass(intersection(X30,X31),intersection(X32,X31)) | ~subclass(X30,X32)) )),
  inference(resolution,[],[f1944,f1205])).
fof(f1205,plain,(
  ( ! [X4,X2,X5,X3] : (member(not_subclass_element(intersection(X2,X3),X4),X5) | subclass(intersection(X2,X3),X4) | ~subclass(X2,X5)) )),
  inference(resolution,[],[f189,f1])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (~member(X2,X0) | member(X2,X1) | ~subclass(X0,X1)) )),
  file('SET254-6.tptp',unknown)).
fof(f189,plain,(
  ( ! [X6,X4,X5] : (member(not_subclass_element(intersection(X4,X5),X6),X4) | subclass(intersection(X4,X5),X6)) )),
  inference(resolution,[],[f21,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (member(not_subclass_element(X0,X1),X0) | subclass(X0,X1)) )),
  file('SET254-6.tptp',unknown)).
fof(f21,axiom,(
  ( ! [X4,X0,X1] : (~member(X4,intersection(X0,X1)) | member(X4,X0)) )),
  file('SET254-6.tptp',unknown)).
fof(f1944,plain,(
  ( ! [X6,X7,X5] : (~member(not_subclass_element(intersection(X5,X6),intersection(X7,X6)),X7) | subclass(intersection(X5,X6),intersection(X7,X6))) )),
  inference(duplicate_literal_removal,[],[f1928])).
fof(f1928,plain,(
  ( ! [X6,X7,X5] : (~member(not_subclass_element(intersection(X5,X6),intersection(X7,X6)),X7) | subclass(intersection(X5,X6),intersection(X7,X6)) | subclass(intersection(X5,X6),intersection(X7,X6))) )),
  inference(resolution,[],[f347,f192])).
fof(f192,plain,(
  ( ! [X6,X4,X5] : (member(not_subclass_element(intersection(X4,X5),X6),X5) | subclass(intersection(X4,X5),X6)) )),
  inference(resolution,[],[f22,f2])).
fof(f22,axiom,(
  ( ! [X4,X0,X1] : (~member(X4,intersection(X0,X1)) | member(X4,X1)) )),
  file('SET254-6.tptp',unknown)).
fof(f347,plain,(
  ( ! [X12,X10,X11] : (~member(not_subclass_element(X10,intersection(X11,X12)),X12) | ~member(not_subclass_element(X10,intersection(X11,X12)),X11) | subclass(X10,intersection(X11,X12))) )),
  inference(resolution,[],[f23,f3])).
fof(f3,axiom,(
  ( ! [X0,X1] : (~member(not_subclass_element(X0,X1),X1) | subclass(X0,X1)) )),
  file('SET254-6.tptp',unknown)).
fof(f23,axiom,(
  ( ! [X4,X0,X1] : (member(X4,intersection(X0,X1)) | ~member(X4,X1) | ~member(X4,X0)) )),
  file('SET254-6.tptp',unknown)).
% SZS output end Proof for SET254-6
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 134965
% Time elapsed: 7.193 s
% ------------------------------
% ------------------------------
