%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for KRS126+1.tptp

% SZS output start CNFRefutation for KRS126+1.tptp
fof(axiom_2, axiom, (! [X] : (cUnsatisfiable(X) => cd1xcomp(X)))).

fof(axiom_3, axiom, (! [X] : (cUnsatisfiable(X) => cd1(X)))).

fof(axiom_7, axiom, (! [X] : (cd1(X) <=> ? [Y0] : ra_Px2(X, Y0)))).

fof(axiom_8, axiom, (! [X] : (cd1xcomp(X) <=> ~ ? [Y] : ra_Px2(X, Y)))).

fof(axiom_11, axiom, (cUnsatisfiable(i2003_11_14_17_22_23554))).

fof(normalize_0_0, plain, (cUnsatisfiable(i2003_11_14_17_22_23554)),
    inference(canonicalize, [], [axiom_11])).

fof(normalize_0_1, plain, (! [X] : (~ cUnsatisfiable(X) | cd1(X))),
    inference(canonicalize, [], [axiom_3])).

fof(normalize_0_2, plain, (! [X] : (~ cUnsatisfiable(X) | cd1(X))),
    inference(specialize, [], [normalize_0_1])).

fof(normalize_0_3, plain,
    (! [X] : (~ cd1(X) <=> ! [Y0] : ~ ra_Px2(X, Y0))),
    inference(canonicalize, [], [axiom_7])).

fof(normalize_0_4, plain,
    (! [X] : (~ cd1(X) <=> ! [Y0] : ~ ra_Px2(X, Y0))),
    inference(specialize, [], [normalize_0_3])).

fof(normalize_0_5, plain,
    (! [X, Y0] :
       ((~ cd1(X) | ra_Px2(X, skolemFOFtoCNF_Y0_1(X))) &
        (~ ra_Px2(X, Y0) | cd1(X)))),
    inference(clausify, [], [normalize_0_4])).

fof(normalize_0_6, plain,
    (! [X] : (~ cd1(X) | ra_Px2(X, skolemFOFtoCNF_Y0_1(X)))),
    inference(conjunct, [], [normalize_0_5])).

fof(normalize_0_7, plain,
    (! [X] : (~ cd1xcomp(X) <=> ? [Y] : ra_Px2(X, Y))),
    inference(canonicalize, [], [axiom_8])).

fof(normalize_0_8, plain,
    (! [X] : (~ cd1xcomp(X) <=> ? [Y] : ra_Px2(X, Y))),
    inference(specialize, [], [normalize_0_7])).

fof(normalize_0_9, plain,
    (! [X, Y] :
       ((~ cd1xcomp(X) | ~ ra_Px2(X, Y)) &
        (cd1xcomp(X) | ra_Px2(X, skolemFOFtoCNF_Y_1(X))))),
    inference(clausify, [], [normalize_0_8])).

fof(normalize_0_10, plain, (! [X, Y] : (~ cd1xcomp(X) | ~ ra_Px2(X, Y))),
    inference(conjunct, [], [normalize_0_9])).

fof(normalize_0_11, plain, (! [X] : (~ cUnsatisfiable(X) | cd1xcomp(X))),
    inference(canonicalize, [], [axiom_2])).

fof(normalize_0_12, plain, (! [X] : (~ cUnsatisfiable(X) | cd1xcomp(X))),
    inference(specialize, [], [normalize_0_11])).

cnf(refute_0_0, plain, (cUnsatisfiable(i2003_11_14_17_22_23554)),
    inference(canonicalize, [], [normalize_0_0])).

cnf(refute_0_1, plain, (~ cUnsatisfiable(X) | cd1(X)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain,
    (~ cUnsatisfiable(i2003_11_14_17_22_23554) |
     cd1(i2003_11_14_17_22_23554)),
    inference(subst, [],
              [refute_0_1 : [bind(X, $fot(i2003_11_14_17_22_23554))]])).

cnf(refute_0_3, plain, (cd1(i2003_11_14_17_22_23554)),
    inference(resolve, [$cnf(cUnsatisfiable(i2003_11_14_17_22_23554))],
              [refute_0_0, refute_0_2])).

cnf(refute_0_4, plain, (~ cd1(X) | ra_Px2(X, skolemFOFtoCNF_Y0_1(X))),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_5, plain,
    (~ cd1(i2003_11_14_17_22_23554) |
     ra_Px2(i2003_11_14_17_22_23554,
       skolemFOFtoCNF_Y0_1(i2003_11_14_17_22_23554))),
    inference(subst, [],
              [refute_0_4 : [bind(X, $fot(i2003_11_14_17_22_23554))]])).

cnf(refute_0_6, plain,
    (ra_Px2(i2003_11_14_17_22_23554,
       skolemFOFtoCNF_Y0_1(i2003_11_14_17_22_23554))),
    inference(resolve, [$cnf(cd1(i2003_11_14_17_22_23554))],
              [refute_0_3, refute_0_5])).

cnf(refute_0_7, plain, (~ cd1xcomp(X) | ~ ra_Px2(X, Y)),
    inference(canonicalize, [], [normalize_0_10])).

cnf(refute_0_8, plain,
    (~ cd1xcomp(i2003_11_14_17_22_23554) |
     ~
       ra_Px2(i2003_11_14_17_22_23554,
         skolemFOFtoCNF_Y0_1(i2003_11_14_17_22_23554))),
    inference(subst, [],
              [refute_0_7 :
               [bind(X, $fot(i2003_11_14_17_22_23554)),
                bind(Y,
                     $fot(skolemFOFtoCNF_Y0_1(i2003_11_14_17_22_23554)))]])).

cnf(refute_0_9, plain, (~ cd1xcomp(i2003_11_14_17_22_23554)),
    inference(resolve,
              [$cnf(ra_Px2(i2003_11_14_17_22_23554,
                      skolemFOFtoCNF_Y0_1(i2003_11_14_17_22_23554)))],
              [refute_0_6, refute_0_8])).

cnf(refute_0_10, plain, (~ cUnsatisfiable(X) | cd1xcomp(X)),
    inference(canonicalize, [], [normalize_0_12])).

cnf(refute_0_11, plain,
    (~ cUnsatisfiable(i2003_11_14_17_22_23554) |
     cd1xcomp(i2003_11_14_17_22_23554)),
    inference(subst, [],
              [refute_0_10 : [bind(X, $fot(i2003_11_14_17_22_23554))]])).

cnf(refute_0_12, plain, (cd1xcomp(i2003_11_14_17_22_23554)),
    inference(resolve, [$cnf(cUnsatisfiable(i2003_11_14_17_22_23554))],
              [refute_0_0, refute_0_11])).

cnf(refute_0_13, plain, ($false),
    inference(resolve, [$cnf(cd1xcomp(i2003_11_14_17_22_23554))],
              [refute_0_12, refute_0_9])).
% SZS output end CNFRefutation for KRS126+1.tptp

