%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for KRS117+1.tptp

% SZS output start CNFRefutation for KRS117+1.tptp
fof(axiom_0, axiom, (! [X] : (cowlThing(X) & ~ cowlNothing(X)))).

fof(axiom_2, axiom,
    (! [X] :
       (cUnsatisfiable(X) <=>
        (ccxcomp(X) & ! [Y] : (rinvR(X, Y) => ca_Vx3(Y)) &
         ? [Y] : (rinvF(X, Y) & cd(Y)))))).

fof(axiom_3, axiom, (! [X] : (cc(X) <=> ~ ? [Y] : ra_Px1(X, Y)))).

fof(axiom_4, axiom, (! [X] : (ccxcomp(X) <=> ? [Y0] : ra_Px1(X, Y0)))).

fof(axiom_5, axiom,
    (! [X] : (cd(X) <=> (? [Y] : (rf(X, Y) & ccxcomp(Y)) & cc(X))))).

fof(axiom_6, axiom,
    (! [X] : (ca_Vx3(X) <=> ? [Y] : (rinvF(X, Y) & cd(Y))))).

fof(axiom_7, axiom,
    (! [X] :
       (cowlThing(X) =>
        ! [Y0, Y1] : ((rf(X, Y0) & rf(X, Y1)) => Y0 = Y1)))).

fof(axiom_8, axiom, (! [X, Y] : (rinvF(X, Y) <=> rf(Y, X)))).

fof(axiom_9, axiom, (! [X, Y] : (rinvR(X, Y) <=> rr(Y, X)))).

fof(axiom_11, axiom, (cUnsatisfiable(i2003_11_14_17_21_37349))).

fof(axiom_12, axiom, (! [X, Y] : (rf(X, Y) => rr(X, Y)))).

fof(normalize_0_0, plain, (! [X] : (~ cc(X) <=> ? [Y] : ra_Px1(X, Y))),
    inference(canonicalize, [], [axiom_3])).

fof(normalize_0_1, plain, (! [X] : (~ cc(X) <=> ? [Y] : ra_Px1(X, Y))),
    inference(specialize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (! [X, Y] :
       ((~ cc(X) | ~ ra_Px1(X, Y)) &
        (cc(X) | ra_Px1(X, skolemFOFtoCNF_Y_2(X))))),
    inference(clausify, [], [normalize_0_1])).

fof(normalize_0_3, plain, (! [X, Y] : (~ cc(X) | ~ ra_Px1(X, Y))),
    inference(conjunct, [], [normalize_0_2])).

fof(normalize_0_4, plain,
    (! [X] : (~ ccxcomp(X) <=> ! [Y0] : ~ ra_Px1(X, Y0))),
    inference(canonicalize, [], [axiom_4])).

fof(normalize_0_5, plain,
    (! [X] : (~ ccxcomp(X) <=> ! [Y0] : ~ ra_Px1(X, Y0))),
    inference(specialize, [], [normalize_0_4])).

fof(normalize_0_6, plain,
    (! [X, Y0] :
       ((~ ccxcomp(X) | ra_Px1(X, skolemFOFtoCNF_Y0(X))) &
        (~ ra_Px1(X, Y0) | ccxcomp(X)))),
    inference(clausify, [], [normalize_0_5])).

fof(normalize_0_7, plain,
    (! [X] : (~ ccxcomp(X) | ra_Px1(X, skolemFOFtoCNF_Y0(X)))),
    inference(conjunct, [], [normalize_0_6])).

fof(normalize_0_8, plain,
    (! [X] :
       (~ cd(X) <=> (~ cc(X) | ! [Y] : (~ ccxcomp(Y) | ~ rf(X, Y))))),
    inference(canonicalize, [], [axiom_5])).

fof(normalize_0_9, plain,
    (! [X] :
       (~ cd(X) <=> (~ cc(X) | ! [Y] : (~ ccxcomp(Y) | ~ rf(X, Y))))),
    inference(specialize, [], [normalize_0_8])).

fof(normalize_0_10, plain,
    (! [X, Y] :
       ((~ cd(X) | cc(X)) & (~ cd(X) | ccxcomp(skolemFOFtoCNF_Y_3(X))) &
        (~ cd(X) | rf(X, skolemFOFtoCNF_Y_3(X))) &
        (~ cc(X) | ~ ccxcomp(Y) | ~ rf(X, Y) | cd(X)))),
    inference(clausify, [], [normalize_0_9])).

fof(normalize_0_11, plain,
    (! [X] : (~ cd(X) | ccxcomp(skolemFOFtoCNF_Y_3(X)))),
    inference(conjunct, [], [normalize_0_10])).

fof(normalize_0_12, plain,
    (! [X] : (~ ca_Vx3(X) <=> ! [Y] : (~ cd(Y) | ~ rinvF(X, Y)))),
    inference(canonicalize, [], [axiom_6])).

fof(normalize_0_13, plain,
    (! [X] : (~ ca_Vx3(X) <=> ! [Y] : (~ cd(Y) | ~ rinvF(X, Y)))),
    inference(specialize, [], [normalize_0_12])).

fof(normalize_0_14, plain,
    (! [X, Y] :
       ((~ ca_Vx3(X) | cd(skolemFOFtoCNF_Y_4(X))) &
        (~ ca_Vx3(X) | rinvF(X, skolemFOFtoCNF_Y_4(X))) &
        (~ cd(Y) | ~ rinvF(X, Y) | ca_Vx3(X)))),
    inference(clausify, [], [normalize_0_13])).

fof(normalize_0_15, plain,
    (! [X] : (~ ca_Vx3(X) | cd(skolemFOFtoCNF_Y_4(X)))),
    inference(conjunct, [], [normalize_0_14])).

fof(normalize_0_16, plain, (! [X, Y] : (~ rinvR(X, Y) <=> ~ rr(Y, X))),
    inference(canonicalize, [], [axiom_9])).

fof(normalize_0_17, plain, (! [X, Y] : (~ rinvR(X, Y) <=> ~ rr(Y, X))),
    inference(specialize, [], [normalize_0_16])).

fof(normalize_0_18, plain,
    (! [X, Y] : ((~ rinvR(X, Y) | rr(Y, X)) & (~ rr(Y, X) | rinvR(X, Y)))),
    inference(clausify, [], [normalize_0_17])).

fof(normalize_0_19, plain, (! [X, Y] : (~ rr(Y, X) | rinvR(X, Y))),
    inference(conjunct, [], [normalize_0_18])).

fof(normalize_0_20, plain, (cUnsatisfiable(i2003_11_14_17_21_37349)),
    inference(canonicalize, [], [axiom_11])).

fof(normalize_0_21, plain,
    (! [X] :
       (~ cUnsatisfiable(X) <=>
        (~ ccxcomp(X) | ? [Y] : (~ ca_Vx3(Y) & rinvR(X, Y)) |
         ! [Y] : (~ cd(Y) | ~ rinvF(X, Y))))),
    inference(canonicalize, [], [axiom_2])).

fof(normalize_0_22, plain,
    (! [X] :
       (~ cUnsatisfiable(X) <=>
        (~ ccxcomp(X) | ? [Y] : (~ ca_Vx3(Y) & rinvR(X, Y)) |
         ! [Y] : (~ cd(Y) | ~ rinvF(X, Y))))),
    inference(specialize, [], [normalize_0_21])).

fof(normalize_0_23, plain,
    (! [X, Y] :
       ((~ cUnsatisfiable(X) | ccxcomp(X)) &
        (~ cUnsatisfiable(X) | cd(skolemFOFtoCNF_Y_1(X))) &
        (~ cUnsatisfiable(X) | rinvF(X, skolemFOFtoCNF_Y_1(X))) &
        (~ cUnsatisfiable(X) | ~ rinvR(X, Y) | ca_Vx3(Y)) &
        (~ ca_Vx3(skolemFOFtoCNF_Y(X)) | ~ ccxcomp(X) | ~ cd(Y) |
         ~ rinvF(X, Y) | cUnsatisfiable(X)) &
        (~ ccxcomp(X) | ~ cd(Y) | ~ rinvF(X, Y) | cUnsatisfiable(X) |
         rinvR(X, skolemFOFtoCNF_Y(X))))),
    inference(clausify, [], [normalize_0_22])).

fof(normalize_0_24, plain,
    (! [X] : (~ cUnsatisfiable(X) | rinvF(X, skolemFOFtoCNF_Y_1(X)))),
    inference(conjunct, [], [normalize_0_23])).

fof(normalize_0_25, plain, (! [X, Y] : (~ rf(Y, X) <=> ~ rinvF(X, Y))),
    inference(canonicalize, [], [axiom_8])).

fof(normalize_0_26, plain, (! [X, Y] : (~ rf(Y, X) <=> ~ rinvF(X, Y))),
    inference(specialize, [], [normalize_0_25])).

fof(normalize_0_27, plain,
    (! [X, Y] : ((~ rf(Y, X) | rinvF(X, Y)) & (~ rinvF(X, Y) | rf(Y, X)))),
    inference(clausify, [], [normalize_0_26])).

fof(normalize_0_28, plain, (! [X, Y] : (~ rinvF(X, Y) | rf(Y, X))),
    inference(conjunct, [], [normalize_0_27])).

fof(normalize_0_29, plain, (! [X, Y] : (~ rf(X, Y) | rr(X, Y))),
    inference(canonicalize, [], [axiom_12])).

fof(normalize_0_30, plain, (! [X, Y] : (~ rf(X, Y) | rr(X, Y))),
    inference(specialize, [], [normalize_0_29])).

fof(normalize_0_31, plain,
    (! [X, Y] : (~ cUnsatisfiable(X) | ~ rinvR(X, Y) | ca_Vx3(Y))),
    inference(conjunct, [], [normalize_0_23])).

fof(normalize_0_32, plain,
    (! [X] : (~ cd(X) | rf(X, skolemFOFtoCNF_Y_3(X)))),
    inference(conjunct, [], [normalize_0_10])).

fof(normalize_0_33, plain,
    (! [X] : (~ ca_Vx3(X) | rinvF(X, skolemFOFtoCNF_Y_4(X)))),
    inference(conjunct, [], [normalize_0_14])).

fof(normalize_0_34, plain,
    (! [X] :
       (~ cowlThing(X) |
        ! [Y0, Y1] : (~ rf(X, Y0) | ~ rf(X, Y1) | Y0 = Y1))),
    inference(canonicalize, [], [axiom_7])).

fof(normalize_0_35, plain,
    (! [X] : ~ cowlNothing(X) & ! [X] : cowlThing(X)),
    inference(canonicalize, [], [axiom_0])).

fof(normalize_0_36, plain, (! [X] : cowlThing(X)),
    inference(conjunct, [], [normalize_0_35])).

fof(normalize_0_37, plain, (! [X] : cowlThing(X)),
    inference(specialize, [], [normalize_0_36])).

fof(normalize_0_38, plain,
    (! [X, Y0, Y1] : (~ rf(X, Y0) | ~ rf(X, Y1) | Y0 = Y1)),
    inference(simplify, [], [normalize_0_34, normalize_0_37])).

fof(normalize_0_39, plain,
    (! [X, Y0, Y1] : (~ rf(X, Y0) | ~ rf(X, Y1) | Y0 = Y1)),
    inference(specialize, [], [normalize_0_38])).

fof(normalize_0_40, plain, (! [X] : (~ cd(X) | cc(X))),
    inference(conjunct, [], [normalize_0_10])).

fof(normalize_0_41, plain,
    (! [X] : (~ cUnsatisfiable(X) | cd(skolemFOFtoCNF_Y_1(X)))),
    inference(conjunct, [], [normalize_0_23])).

cnf(refute_0_0, plain, (~ cc(X) | ~ ra_Px1(X, Y)),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_1, plain,
    (~
       cc(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))) |
     ~
       ra_Px1(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
         skolemFOFtoCNF_Y0(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))),
    inference(subst, [],
              [refute_0_0 :
               [bind(X,
                     $fot(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
                bind(Y,
                     $fot(skolemFOFtoCNF_Y0(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))))]])).

cnf(refute_0_2, plain, (~ ccxcomp(X) | ra_Px1(X, skolemFOFtoCNF_Y0(X))),
    inference(canonicalize, [], [normalize_0_7])).

cnf(refute_0_3, plain,
    (~
       ccxcomp(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))) |
     ra_Px1(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
       skolemFOFtoCNF_Y0(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))),
    inference(subst, [],
              [refute_0_2 :
               [bind(X,
                     $fot(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))]])).

cnf(refute_0_4, plain, (~ cd(X) | ccxcomp(skolemFOFtoCNF_Y_3(X))),
    inference(canonicalize, [], [normalize_0_11])).

cnf(refute_0_5, plain,
    (~
       cd(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))) |
     ccxcomp(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
    inference(subst, [],
              [refute_0_4 :
               [bind(X,
                     $fot(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))]])).

cnf(refute_0_6, plain, (~ ca_Vx3(X) | cd(skolemFOFtoCNF_Y_4(X))),
    inference(canonicalize, [], [normalize_0_15])).

cnf(refute_0_7, plain,
    (~ ca_Vx3(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)) |
     cd(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))),
    inference(subst, [],
              [refute_0_6 :
               [bind(X,
                     $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))]])).

cnf(refute_0_8, plain, (~ rr(Y, X) | rinvR(X, Y)),
    inference(canonicalize, [], [normalize_0_19])).

cnf(refute_0_9, plain,
    (~
       rr(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
         i2003_11_14_17_21_37349) |
     rinvR(i2003_11_14_17_21_37349,
       skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(subst, [],
              [refute_0_8 :
               [bind(X, $fot(i2003_11_14_17_21_37349)),
                bind(Y,
                     $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))]])).

cnf(refute_0_10, plain, (cUnsatisfiable(i2003_11_14_17_21_37349)),
    inference(canonicalize, [], [normalize_0_20])).

cnf(refute_0_11, plain,
    (~ cUnsatisfiable(X) | rinvF(X, skolemFOFtoCNF_Y_1(X))),
    inference(canonicalize, [], [normalize_0_24])).

cnf(refute_0_12, plain,
    (~ cUnsatisfiable(i2003_11_14_17_21_37349) |
     rinvF(i2003_11_14_17_21_37349,
       skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(subst, [],
              [refute_0_11 : [bind(X, $fot(i2003_11_14_17_21_37349))]])).

cnf(refute_0_13, plain,
    (rinvF(i2003_11_14_17_21_37349,
       skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(resolve, [$cnf(cUnsatisfiable(i2003_11_14_17_21_37349))],
              [refute_0_10, refute_0_12])).

cnf(refute_0_14, plain, (~ rinvF(X, Y) | rf(Y, X)),
    inference(canonicalize, [], [normalize_0_28])).

cnf(refute_0_15, plain,
    (~
       rinvF(i2003_11_14_17_21_37349,
         skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)) |
     rf(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
       i2003_11_14_17_21_37349)),
    inference(subst, [],
              [refute_0_14 :
               [bind(X, $fot(i2003_11_14_17_21_37349)),
                bind(Y,
                     $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))]])).

cnf(refute_0_16, plain,
    (rf(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
       i2003_11_14_17_21_37349)),
    inference(resolve,
              [$cnf(rinvF(i2003_11_14_17_21_37349,
                      skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))],
              [refute_0_13, refute_0_15])).

cnf(refute_0_17, plain, (~ rf(X, Y) | rr(X, Y)),
    inference(canonicalize, [], [normalize_0_30])).

cnf(refute_0_18, plain,
    (~
       rf(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
         i2003_11_14_17_21_37349) |
     rr(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
       i2003_11_14_17_21_37349)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
                bind(Y, $fot(i2003_11_14_17_21_37349))]])).

cnf(refute_0_19, plain,
    (rr(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
       i2003_11_14_17_21_37349)),
    inference(resolve,
              [$cnf(rf(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
                      i2003_11_14_17_21_37349))],
              [refute_0_16, refute_0_18])).

cnf(refute_0_20, plain,
    (rinvR(i2003_11_14_17_21_37349,
       skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(resolve,
              [$cnf(rr(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
                      i2003_11_14_17_21_37349))],
              [refute_0_19, refute_0_9])).

cnf(refute_0_21, plain, (~ cUnsatisfiable(X) | ~ rinvR(X, Y) | ca_Vx3(Y)),
    inference(canonicalize, [], [normalize_0_31])).

cnf(refute_0_22, plain,
    (~ cUnsatisfiable(i2003_11_14_17_21_37349) |
     ~
       rinvR(i2003_11_14_17_21_37349,
         skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)) |
     ca_Vx3(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(subst, [],
              [refute_0_21 :
               [bind(X, $fot(i2003_11_14_17_21_37349)),
                bind(Y,
                     $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))]])).

cnf(refute_0_23, plain,
    (~ cUnsatisfiable(i2003_11_14_17_21_37349) |
     ca_Vx3(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(resolve,
              [$cnf(rinvR(i2003_11_14_17_21_37349,
                      skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))],
              [refute_0_20, refute_0_22])).

cnf(refute_0_24, plain,
    (ca_Vx3(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(resolve, [$cnf(cUnsatisfiable(i2003_11_14_17_21_37349))],
              [refute_0_10, refute_0_23])).

cnf(refute_0_25, plain,
    (cd(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))),
    inference(resolve,
              [$cnf(ca_Vx3(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))],
              [refute_0_24, refute_0_7])).

cnf(refute_0_26, plain,
    (ccxcomp(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
    inference(resolve,
              [$cnf(cd(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))],
              [refute_0_25, refute_0_5])).

cnf(refute_0_27, plain,
    (ra_Px1(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
       skolemFOFtoCNF_Y0(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))),
    inference(resolve,
              [$cnf(ccxcomp(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))],
              [refute_0_26, refute_0_3])).

cnf(refute_0_28, plain,
    (~
       cc(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
    inference(resolve,
              [$cnf(ra_Px1(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
                      skolemFOFtoCNF_Y0(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))))],
              [refute_0_27, refute_0_1])).

cnf(refute_0_29, plain, (~ cd(X) | rf(X, skolemFOFtoCNF_Y_3(X))),
    inference(canonicalize, [], [normalize_0_32])).

cnf(refute_0_30, plain,
    (~
       cd(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))) |
     rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
       skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
    inference(subst, [],
              [refute_0_29 :
               [bind(X,
                     $fot(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))]])).

cnf(refute_0_31, plain,
    (rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
       skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
    inference(resolve,
              [$cnf(cd(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))],
              [refute_0_25, refute_0_30])).

cnf(refute_0_32, plain,
    (~
       rinvF(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
         skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))) |
     rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
       skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(subst, [],
              [refute_0_14 :
               [bind(X, $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
                bind(Y,
                     $fot(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))]])).

cnf(refute_0_33, plain, (~ ca_Vx3(X) | rinvF(X, skolemFOFtoCNF_Y_4(X))),
    inference(canonicalize, [], [normalize_0_33])).

cnf(refute_0_34, plain,
    (~ ca_Vx3(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)) |
     rinvF(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
       skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))),
    inference(subst, [],
              [refute_0_33 :
               [bind(X,
                     $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))]])).

cnf(refute_0_35, plain,
    (rinvF(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
       skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))),
    inference(resolve,
              [$cnf(ca_Vx3(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))],
              [refute_0_24, refute_0_34])).

cnf(refute_0_36, plain,
    (rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
       skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(resolve,
              [$cnf(rinvF(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
                      skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))],
              [refute_0_35, refute_0_32])).

cnf(refute_0_37, plain, (~ rf(X, Y0) | ~ rf(X, Y1) | Y0 = Y1),
    inference(canonicalize, [], [normalize_0_39])).

cnf(refute_0_38, plain,
    (~
       rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
         X_38) |
     ~
       rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
         skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)) |
     skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349) = X_38),
    inference(subst, [],
              [refute_0_37 :
               [bind(X,
                     $fot(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))),
                bind(Y0,
                     $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
                bind(Y1, $fot(X_38))]])).

cnf(refute_0_39, plain,
    (~
       rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
         X_38) | skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349) = X_38),
    inference(resolve,
              [$cnf(rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
                      skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))],
              [refute_0_36, refute_0_38])).

cnf(refute_0_40, plain,
    (~
       rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
         skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))) |
     skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349) =
     skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))),
    inference(subst, [],
              [refute_0_39 :
               [bind(X_38,
                     $fot(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))]])).

cnf(refute_0_41, plain,
    (skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349) =
     skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))),
    inference(resolve,
              [$cnf(rf(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
                      skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))],
              [refute_0_31, refute_0_40])).

cnf(refute_0_42, plain, (X0 = X0),
    introduced(tautology, [refl, [$fot(X0)]])).

cnf(refute_0_43, plain, (X0 != X0 | X0 != Y2 | Y2 = X0),
    introduced(tautology,
               [equality, [$cnf($equal(X0, X0)), [0], $fot(Y2)]])).

cnf(refute_0_44, plain, (X0 != Y2 | Y2 = X0),
    inference(resolve, [$cnf($equal(X0, X0))],
              [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain,
    (skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349) !=
     skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))) |
     skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))) =
     skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
    inference(subst, [],
              [refute_0_44 :
               [bind(X0,
                     $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
                bind(Y2,
                     $fot(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))]])).

cnf(refute_0_46, plain,
    (skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))) =
     skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349),
                      skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))],
              [refute_0_41, refute_0_45])).

cnf(refute_0_47, plain,
    (skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))) !=
     skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349) |
     ~ cc(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)) |
     cc(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
    introduced(tautology,
               [equality,
                [$cnf(~ cc(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
                 [0],
                 $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))]])).

cnf(refute_0_48, plain,
    (~ cc(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)) |
     cc(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))))),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
                      skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))],
              [refute_0_46, refute_0_47])).

cnf(refute_0_49, plain,
    (~ cc(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(resolve,
              [$cnf(cc(skolemFOFtoCNF_Y_3(skolemFOFtoCNF_Y_4(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))))],
              [refute_0_48, refute_0_28])).

cnf(refute_0_50, plain, (~ cd(X) | cc(X)),
    inference(canonicalize, [], [normalize_0_40])).

cnf(refute_0_51, plain,
    (~ cd(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)) |
     cc(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(subst, [],
              [refute_0_50 :
               [bind(X,
                     $fot(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))]])).

cnf(refute_0_52, plain, (~ cUnsatisfiable(X) | cd(skolemFOFtoCNF_Y_1(X))),
    inference(canonicalize, [], [normalize_0_41])).

cnf(refute_0_53, plain,
    (~ cUnsatisfiable(i2003_11_14_17_21_37349) |
     cd(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(subst, [],
              [refute_0_52 : [bind(X, $fot(i2003_11_14_17_21_37349))]])).

cnf(refute_0_54, plain, (cd(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(resolve, [$cnf(cUnsatisfiable(i2003_11_14_17_21_37349))],
              [refute_0_10, refute_0_53])).

cnf(refute_0_55, plain, (cc(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349))),
    inference(resolve,
              [$cnf(cd(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))],
              [refute_0_54, refute_0_51])).

cnf(refute_0_56, plain, ($false),
    inference(resolve,
              [$cnf(cc(skolemFOFtoCNF_Y_1(i2003_11_14_17_21_37349)))],
              [refute_0_55, refute_0_49])).
% SZS output end CNFRefutation for KRS117+1.tptp

