%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for CAT013-4.tptp

% SZS output start CNFRefutation for CAT013-4.tptp
cnf(codomain_has_elements, axiom,
    (~ there_exists(codomain(X)) | there_exists(X))).

cnf(domain_codomain_composition1, axiom,
    (~ there_exists(compose(X, Y)) | domain(X) = codomain(Y))).

cnf(compose_codomain, axiom, (compose(codomain(X), X) = X)).

cnf(assume_codomain_exists, hypothesis, (there_exists(codomain(a)))).

cnf(prove_domain_of_codomain_is_codomain, negated_conjecture,
    (domain(codomain(a)) != codomain(a))).

cnf(refute_0_0, plain, (~ there_exists(codomain(a)) | there_exists(a)),
    inference(subst, [], [codomain_has_elements : [bind(X, $fot(a))]])).

cnf(refute_0_1, plain, (there_exists(a)),
    inference(resolve, [$cnf(there_exists(codomain(a)))],
              [assume_codomain_exists, refute_0_0])).

cnf(refute_0_2, plain,
    (~ there_exists(compose(codomain(X_42), X_42)) |
     domain(codomain(X_42)) = codomain(X_42)),
    inference(subst, [],
              [domain_codomain_composition1 :
               [bind(X, $fot(codomain(X_42))), bind(Y, $fot(X_42))]])).

cnf(refute_0_3, plain, (compose(codomain(X_42), X_42) = X_42),
    inference(subst, [], [compose_codomain : [bind(X, $fot(X_42))]])).

cnf(refute_0_4, plain,
    (compose(codomain(X_42), X_42) != X_42 | ~ there_exists(X_42) |
     there_exists(compose(codomain(X_42), X_42))),
    introduced(tautology,
               [equality,
                [$cnf(~ there_exists(compose(codomain(X_42), X_42))), [0],
                 $fot(X_42)]])).

cnf(refute_0_5, plain,
    (~ there_exists(X_42) | there_exists(compose(codomain(X_42), X_42))),
    inference(resolve, [$cnf($equal(compose(codomain(X_42), X_42), X_42))],
              [refute_0_3, refute_0_4])).

cnf(refute_0_6, plain,
    (~ there_exists(X_42) | domain(codomain(X_42)) = codomain(X_42)),
    inference(resolve, [$cnf(there_exists(compose(codomain(X_42), X_42)))],
              [refute_0_5, refute_0_2])).

cnf(refute_0_7, plain,
    (~ there_exists(a) | domain(codomain(a)) = codomain(a)),
    inference(subst, [], [refute_0_6 : [bind(X_42, $fot(a))]])).

cnf(refute_0_8, plain, (domain(codomain(a)) = codomain(a)),
    inference(resolve, [$cnf(there_exists(a))], [refute_0_1, refute_0_7])).

cnf(refute_0_9, plain, ($false),
    inference(resolve, [$cnf($equal(domain(codomain(a)), codomain(a)))],
              [refute_0_8, prove_domain_of_codomain_is_codomain])).
% SZS output end CNFRefutation for CAT013-4.tptp

