%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for GRP123-6.003.tptp

% SZS output start CNFRefutation for GRP123-6.003.tptp
cnf(element_1, axiom, (group_element(e_1))).

cnf(element_2, axiom, (group_element(e_2))).

cnf(element_3, axiom, (group_element(e_3))).

cnf(e_1_is_not_e_2, axiom, (~ equalish(e_1, e_2))).

cnf(e_1_is_not_e_3, axiom, (~ equalish(e_1, e_3))).

cnf(e_2_is_not_e_1, axiom, (~ equalish(e_2, e_1))).

cnf(e_3_is_not_e_1, axiom, (~ equalish(e_3, e_1))).

cnf(product1_total_function1, axiom,
    (~ group_element(X) | ~ group_element(Y) | product1(X, Y, e_1) |
     product1(X, Y, e_2) | product1(X, Y, e_3))).

cnf(product1_right_cancellation, axiom,
    (~ product1(X, W, Y) | ~ product1(X, Z, Y) | equalish(W, Z))).

cnf(product1_left_cancellation, axiom,
    (~ product1(W, Y, X) | ~ product1(Z, Y, X) | equalish(W, Z))).

cnf(product1_idempotence, axiom, (product1(X, X, X))).

cnf(product2_total_function2, axiom,
    (~ product2(X, Y, W) | ~ product2(X, Y, Z) | equalish(W, Z))).

cnf(product2_idempotence, axiom, (product2(X, X, X))).

cnf(qg1a, negated_conjecture,
    (~ product1(X, Y, Z1) | ~ product1(Z1, Y, Z2) | product2(Z2, X, Y))).

cnf(refute_0_0, plain, (product1(X_27, X_27, X_27)),
    inference(subst, [], [product1_idempotence : [bind(X, $fot(X_27))]])).

cnf(refute_0_1, plain,
    (~ product1(X_27, X_27, X_27) | ~ product1(X_29, X_27, X_27) |
     equalish(X_27, X_29)),
    inference(subst, [],
              [product1_left_cancellation :
               [bind(W, $fot(X_27)), bind(X, $fot(X_27)),
                bind(Y, $fot(X_27)), bind(Z, $fot(X_29))]])).

cnf(refute_0_2, plain,
    (~ product1(X_29, X_27, X_27) | equalish(X_27, X_29)),
    inference(resolve, [$cnf(product1(X_27, X_27, X_27))],
              [refute_0_0, refute_0_1])).

cnf(refute_0_3, plain, (~ product1(e_2, e_1, e_1) | equalish(e_1, e_2)),
    inference(subst, [],
              [refute_0_2 :
               [bind(X_27, $fot(e_1)), bind(X_29, $fot(e_2))]])).

cnf(refute_0_4, plain, (product1(X_10, X_10, X_10)),
    inference(subst, [], [product1_idempotence : [bind(X, $fot(X_10))]])).

cnf(refute_0_5, plain,
    (~ product1(X_10, X_10, X_10) | ~ product1(X_10, X_11, X_10) |
     equalish(X_10, X_11)),
    inference(subst, [],
              [product1_right_cancellation :
               [bind(W, $fot(X_10)), bind(X, $fot(X_10)),
                bind(Y, $fot(X_10)), bind(Z, $fot(X_11))]])).

cnf(refute_0_6, plain,
    (~ product1(X_10, X_11, X_10) | equalish(X_10, X_11)),
    inference(resolve, [$cnf(product1(X_10, X_10, X_10))],
              [refute_0_4, refute_0_5])).

cnf(refute_0_7, plain, (~ product1(e_2, e_1, e_2) | equalish(e_2, e_1)),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_10, $fot(e_2)), bind(X_11, $fot(e_1))]])).

cnf(refute_0_8, plain, (~ product1(e_3, e_1, e_1) | equalish(e_1, e_3)),
    inference(subst, [],
              [refute_0_2 :
               [bind(X_27, $fot(e_1)), bind(X_29, $fot(e_3))]])).

cnf(refute_0_9, plain, (product2(X_37, X_37, X_37)),
    inference(subst, [], [product2_idempotence : [bind(X, $fot(X_37))]])).

cnf(refute_0_10, plain,
    (~ product2(X_37, X_37, X_34) | ~ product2(X_37, X_37, X_37) |
     equalish(X_34, X_37)),
    inference(subst, [],
              [product2_total_function2 :
               [bind(W, $fot(X_34)), bind(X, $fot(X_37)),
                bind(Y, $fot(X_37)), bind(Z, $fot(X_37))]])).

cnf(refute_0_11, plain,
    (~ product2(X_37, X_37, X_34) | equalish(X_34, X_37)),
    inference(resolve, [$cnf(product2(X_37, X_37, X_37))],
              [refute_0_9, refute_0_10])).

cnf(refute_0_12, plain, (~ product2(e_2, e_2, e_1) | equalish(e_1, e_2)),
    inference(subst, [],
              [refute_0_11 :
               [bind(X_34, $fot(e_1)), bind(X_37, $fot(e_2))]])).

cnf(refute_0_13, plain,
    (~ product1(X, e_1, e_3) | ~ product1(e_3, e_1, e_2) |
     product2(e_2, X, e_1)),
    inference(subst, [],
              [qg1a :
               [bind(Y, $fot(e_1)), bind(Z1, $fot(e_3)),
                bind(Z2, $fot(e_2))]])).

cnf(refute_0_14, plain, (~ product1(e_3, e_1, e_3) | equalish(e_3, e_1)),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_10, $fot(e_3)), bind(X_11, $fot(e_1))]])).

cnf(refute_0_15, plain,
    (~ group_element(X_67) | ~ group_element(e_3) |
     product1(e_3, X_67, e_1) | product1(e_3, X_67, e_2) |
     product1(e_3, X_67, e_3)),
    inference(subst, [],
              [product1_total_function1 :
               [bind(X, $fot(e_3)), bind(Y, $fot(X_67))]])).

cnf(refute_0_16, plain,
    (~ group_element(X_67) | product1(e_3, X_67, e_1) |
     product1(e_3, X_67, e_2) | product1(e_3, X_67, e_3)),
    inference(resolve, [$cnf(group_element(e_3))],
              [element_3, refute_0_15])).

cnf(refute_0_17, plain,
    (~ group_element(e_1) | product1(e_3, e_1, e_1) |
     product1(e_3, e_1, e_2) | product1(e_3, e_1, e_3)),
    inference(subst, [], [refute_0_16 : [bind(X_67, $fot(e_1))]])).

cnf(refute_0_18, plain,
    (product1(e_3, e_1, e_1) | product1(e_3, e_1, e_2) |
     product1(e_3, e_1, e_3)),
    inference(resolve, [$cnf(group_element(e_1))],
              [element_1, refute_0_17])).

cnf(refute_0_19, plain,
    (equalish(e_3, e_1) | product1(e_3, e_1, e_1) |
     product1(e_3, e_1, e_2)),
    inference(resolve, [$cnf(product1(e_3, e_1, e_3))],
              [refute_0_18, refute_0_14])).

cnf(refute_0_20, plain,
    (product1(e_3, e_1, e_1) | product1(e_3, e_1, e_2)),
    inference(resolve, [$cnf(equalish(e_3, e_1))],
              [refute_0_19, e_3_is_not_e_1])).

cnf(refute_0_21, plain,
    (~ product1(X, e_1, e_3) | product1(e_3, e_1, e_1) |
     product2(e_2, X, e_1)),
    inference(resolve, [$cnf(product1(e_3, e_1, e_2))],
              [refute_0_20, refute_0_13])).

cnf(refute_0_22, plain,
    (~ product1(e_2, e_1, e_3) | product1(e_3, e_1, e_1) |
     product2(e_2, e_2, e_1)),
    inference(subst, [], [refute_0_21 : [bind(X, $fot(e_2))]])).

cnf(refute_0_23, plain,
    (~ group_element(X_66) | ~ group_element(e_1) |
     product1(X_66, e_1, e_1) | product1(X_66, e_1, e_2) |
     product1(X_66, e_1, e_3)),
    inference(subst, [],
              [product1_total_function1 :
               [bind(X, $fot(X_66)), bind(Y, $fot(e_1))]])).

cnf(refute_0_24, plain,
    (~ group_element(X_66) | product1(X_66, e_1, e_1) |
     product1(X_66, e_1, e_2) | product1(X_66, e_1, e_3)),
    inference(resolve, [$cnf(group_element(e_1))],
              [element_1, refute_0_23])).

cnf(refute_0_25, plain,
    (~ group_element(e_2) | product1(e_2, e_1, e_1) |
     product1(e_2, e_1, e_2) | product1(e_2, e_1, e_3)),
    inference(subst, [], [refute_0_24 : [bind(X_66, $fot(e_2))]])).

cnf(refute_0_26, plain,
    (product1(e_2, e_1, e_1) | product1(e_2, e_1, e_2) |
     product1(e_2, e_1, e_3)),
    inference(resolve, [$cnf(group_element(e_2))],
              [element_2, refute_0_25])).

cnf(refute_0_27, plain,
    (product1(e_2, e_1, e_1) | product1(e_2, e_1, e_2) |
     product1(e_3, e_1, e_1) | product2(e_2, e_2, e_1)),
    inference(resolve, [$cnf(product1(e_2, e_1, e_3))],
              [refute_0_26, refute_0_22])).

cnf(refute_0_28, plain,
    (equalish(e_1, e_2) | product1(e_2, e_1, e_1) |
     product1(e_2, e_1, e_2) | product1(e_3, e_1, e_1)),
    inference(resolve, [$cnf(product2(e_2, e_2, e_1))],
              [refute_0_27, refute_0_12])).

cnf(refute_0_29, plain,
    (product1(e_2, e_1, e_1) | product1(e_2, e_1, e_2) |
     product1(e_3, e_1, e_1)),
    inference(resolve, [$cnf(equalish(e_1, e_2))],
              [refute_0_28, e_1_is_not_e_2])).

cnf(refute_0_30, plain,
    (equalish(e_1, e_3) | product1(e_2, e_1, e_1) |
     product1(e_2, e_1, e_2)),
    inference(resolve, [$cnf(product1(e_3, e_1, e_1))],
              [refute_0_29, refute_0_8])).

cnf(refute_0_31, plain,
    (product1(e_2, e_1, e_1) | product1(e_2, e_1, e_2)),
    inference(resolve, [$cnf(equalish(e_1, e_3))],
              [refute_0_30, e_1_is_not_e_3])).

cnf(refute_0_32, plain, (equalish(e_2, e_1) | product1(e_2, e_1, e_1)),
    inference(resolve, [$cnf(product1(e_2, e_1, e_2))],
              [refute_0_31, refute_0_7])).

cnf(refute_0_33, plain, (product1(e_2, e_1, e_1)),
    inference(resolve, [$cnf(equalish(e_2, e_1))],
              [refute_0_32, e_2_is_not_e_1])).

cnf(refute_0_34, plain, (equalish(e_1, e_2)),
    inference(resolve, [$cnf(product1(e_2, e_1, e_1))],
              [refute_0_33, refute_0_3])).

cnf(refute_0_35, plain, ($false),
    inference(resolve, [$cnf(equalish(e_1, e_2))],
              [refute_0_34, e_1_is_not_e_2])).
% SZS output end CNFRefutation for GRP123-6.003.tptp

