%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for GEO018-3.tptp

% SZS output start CNFRefutation for GEO018-3.tptp
cnf(d2, axiom, (~ equidistant(U, V, W, X) | equidistant(W, X, U, V))).

cnf(d3, axiom, (~ equidistant(U, V, W, X) | equidistant(V, U, W, X))).

cnf(u_to_v_equals_w_to_x, hypothesis, (equidistant(u, v, w, x))).

cnf(prove_symmetry, negated_conjecture, (~ equidistant(v, u, x, w))).

cnf(refute_0_0, plain,
    (~ equidistant(u, v, x, w) | equidistant(v, u, x, w)),
    inference(subst, [],
              [d3 :
               [bind(U, $fot(u)), bind(V, $fot(v)), bind(W, $fot(x)),
                bind(X, $fot(w))]])).

cnf(refute_0_1, plain,
    (~ equidistant(x, w, u, v) | equidistant(u, v, x, w)),
    inference(subst, [],
              [d2 :
               [bind(U, $fot(x)), bind(V, $fot(w)), bind(W, $fot(u)),
                bind(X, $fot(v))]])).

cnf(refute_0_2, plain,
    (~ equidistant(u, v, w, x) | equidistant(w, x, u, v)),
    inference(subst, [],
              [d2 :
               [bind(U, $fot(u)), bind(V, $fot(v)), bind(W, $fot(w)),
                bind(X, $fot(x))]])).

cnf(refute_0_3, plain, (equidistant(w, x, u, v)),
    inference(resolve, [$cnf(equidistant(u, v, w, x))],
              [u_to_v_equals_w_to_x, refute_0_2])).

cnf(refute_0_4, plain,
    (~ equidistant(w, x, u, v) | equidistant(x, w, u, v)),
    inference(subst, [],
              [d3 :
               [bind(U, $fot(w)), bind(V, $fot(x)), bind(W, $fot(u)),
                bind(X, $fot(v))]])).

cnf(refute_0_5, plain, (equidistant(x, w, u, v)),
    inference(resolve, [$cnf(equidistant(w, x, u, v))],
              [refute_0_3, refute_0_4])).

cnf(refute_0_6, plain, (equidistant(u, v, x, w)),
    inference(resolve, [$cnf(equidistant(x, w, u, v))],
              [refute_0_5, refute_0_1])).

cnf(refute_0_7, plain, (equidistant(v, u, x, w)),
    inference(resolve, [$cnf(equidistant(u, v, x, w))],
              [refute_0_6, refute_0_0])).

cnf(refute_0_8, plain, ($false),
    inference(resolve, [$cnf(equidistant(v, u, x, w))],
              [refute_0_7, prove_symmetry])).
% SZS output end CNFRefutation for GEO018-3.tptp

