%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for GEO065-2.tptp

% SZS output start CNFRefutation for GEO065-2.tptp
cnf(identity_for_equidistance, axiom, (~ equidistant(X, Y, Z, Z) | X = Y)).

cnf(segment_construction1, axiom, (between(X, Y, extension(X, Y, W, V)))).

cnf(segment_construction2, axiom,
    (equidistant(Y, extension(X, Y, W, V), W, V))).

cnf(identity_for_betweeness, axiom, (~ between(X, Y, X) | X = Y)).

cnf(inner_pasch1, axiom,
    (~ between(U, V, W) | ~ between(Y, X, W) |
     between(V, inner_pasch(U, V, W, X, Y), Y))).

cnf(inner_pasch2, axiom,
    (~ between(U, V, W) | ~ between(Y, X, W) |
     between(X, inner_pasch(U, V, W, X, Y), U))).

cnf(colinearity2, axiom, (~ between(Y, Z, X) | colinear(X, Y, Z))).

cnf(w_between_u_and_v, hypothesis, (between(u, w, v))).

cnf(prove_uvw_colinear, negated_conjecture, (~ colinear(u, v, w))).

cnf(refute_0_0, plain, (~ between(v, w, u) | colinear(u, v, w)),
    inference(subst, [],
              [colinearity2 :
               [bind(X, $fot(u)), bind(Y, $fot(v)), bind(Z, $fot(w))]])).

cnf(refute_0_1, plain, (between(X, Y, extension(X, Y, V, V))),
    inference(subst, [], [segment_construction1 : [bind(W, $fot(V))]])).

cnf(refute_0_2, plain,
    (equidistant(X_12, extension(X, X_12, X_14, X_14), X_14, X_14)),
    inference(subst, [],
              [segment_construction2 :
               [bind(V, $fot(X_14)), bind(W, $fot(X_14)),
                bind(Y, $fot(X_12))]])).

cnf(refute_0_3, plain,
    (~ equidistant(X_12, extension(X, X_12, X_14, X_14), X_14, X_14) |
     X_12 = extension(X, X_12, X_14, X_14)),
    inference(subst, [],
              [identity_for_equidistance :
               [bind(X, $fot(X_12)),
                bind(Y, $fot(extension(X, X_12, X_14, X_14))),
                bind(Z, $fot(X_14))]])).

cnf(refute_0_4, plain, (X_12 = extension(X, X_12, X_14, X_14)),
    inference(resolve,
              [$cnf(equidistant(X_12, extension(X, X_12, X_14, X_14), X_14,
                      X_14))], [refute_0_2, refute_0_3])).

cnf(refute_0_5, plain, (Y = extension(X, Y, V, V)),
    inference(subst, [],
              [refute_0_4 : [bind(X_12, $fot(Y)), bind(X_14, $fot(V))]])).

cnf(refute_0_6, plain, (X0 = X0),
    introduced(tautology, [refl, [$fot(X0)]])).

cnf(refute_0_7, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology,
               [equality, [$cnf($equal(X0, X0)), [0], $fot(Y0)]])).

cnf(refute_0_8, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [$cnf($equal(X0, X0))], [refute_0_6, refute_0_7])).

cnf(refute_0_9, plain,
    (Y != extension(X, Y, V, V) | extension(X, Y, V, V) = Y),
    inference(subst, [],
              [refute_0_8 :
               [bind(X0, $fot(Y)),
                bind(Y0, $fot(extension(X, Y, V, V)))]])).

cnf(refute_0_10, plain, (extension(X, Y, V, V) = Y),
    inference(resolve, [$cnf($equal(Y, extension(X, Y, V, V)))],
              [refute_0_5, refute_0_9])).

cnf(refute_0_11, plain,
    (extension(X, Y, V, V) != Y | ~ between(X, Y, extension(X, Y, V, V)) |
     between(X, Y, Y)),
    introduced(tautology,
               [equality,
                [$cnf(between(X, Y, extension(X, Y, V, V))), [2],
                 $fot(Y)]])).

cnf(refute_0_12, plain,
    (~ between(X, Y, extension(X, Y, V, V)) | between(X, Y, Y)),
    inference(resolve, [$cnf($equal(extension(X, Y, V, V), Y))],
              [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain, (between(X, Y, Y)),
    inference(resolve, [$cnf(between(X, Y, extension(X, Y, V, V)))],
              [refute_0_1, refute_0_12])).

cnf(refute_0_14, plain, (between(X_351, v, v)),
    inference(subst, [],
              [refute_0_13 : [bind(X, $fot(X_351)), bind(Y, $fot(v))]])).

cnf(refute_0_15, plain,
    (~ between(X_349, X_348, v) | ~ between(u, w, v) |
     between(X_348, inner_pasch(u, w, v, X_348, X_349), u)),
    inference(subst, [],
              [inner_pasch2 :
               [bind(U, $fot(u)), bind(V, $fot(w)), bind(W, $fot(v)),
                bind(X, $fot(X_348)), bind(Y, $fot(X_349))]])).

cnf(refute_0_16, plain,
    (~ between(X_349, X_348, v) |
     between(X_348, inner_pasch(u, w, v, X_348, X_349), u)),
    inference(resolve, [$cnf(between(u, w, v))],
              [w_between_u_and_v, refute_0_15])).

cnf(refute_0_17, plain,
    (~ between(X_351, v, v) |
     between(v, inner_pasch(u, w, v, v, X_351), u)),
    inference(subst, [],
              [refute_0_16 :
               [bind(X_348, $fot(v)), bind(X_349, $fot(X_351))]])).

cnf(refute_0_18, plain, (between(v, inner_pasch(u, w, v, v, X_351), u)),
    inference(resolve, [$cnf(between(X_351, v, v))],
              [refute_0_14, refute_0_17])).

cnf(refute_0_19, plain, (between(v, inner_pasch(u, w, v, v, w), u)),
    inference(subst, [], [refute_0_18 : [bind(X_351, $fot(w))]])).

cnf(refute_0_20, plain,
    (~ between(w, inner_pasch(u, w, v, v, w), w) |
     w = inner_pasch(u, w, v, v, w)),
    inference(subst, [],
              [identity_for_betweeness :
               [bind(X, $fot(w)),
                bind(Y, $fot(inner_pasch(u, w, v, v, w)))]])).

cnf(refute_0_21, plain, (between(X_228, v, v)),
    inference(subst, [],
              [refute_0_13 : [bind(X, $fot(X_228)), bind(Y, $fot(v))]])).

cnf(refute_0_22, plain,
    (~ between(X_212, X_211, v) | ~ between(u, w, v) |
     between(w, inner_pasch(u, w, v, X_211, X_212), X_212)),
    inference(subst, [],
              [inner_pasch1 :
               [bind(U, $fot(u)), bind(V, $fot(w)), bind(W, $fot(v)),
                bind(X, $fot(X_211)), bind(Y, $fot(X_212))]])).

cnf(refute_0_23, plain,
    (~ between(X_212, X_211, v) |
     between(w, inner_pasch(u, w, v, X_211, X_212), X_212)),
    inference(resolve, [$cnf(between(u, w, v))],
              [w_between_u_and_v, refute_0_22])).

cnf(refute_0_24, plain,
    (~ between(X_228, v, v) |
     between(w, inner_pasch(u, w, v, v, X_228), X_228)),
    inference(subst, [],
              [refute_0_23 :
               [bind(X_211, $fot(v)), bind(X_212, $fot(X_228))]])).

cnf(refute_0_25, plain,
    (between(w, inner_pasch(u, w, v, v, X_228), X_228)),
    inference(resolve, [$cnf(between(X_228, v, v))],
              [refute_0_21, refute_0_24])).

cnf(refute_0_26, plain, (between(w, inner_pasch(u, w, v, v, w), w)),
    inference(subst, [], [refute_0_25 : [bind(X_228, $fot(w))]])).

cnf(refute_0_27, plain, (w = inner_pasch(u, w, v, v, w)),
    inference(resolve, [$cnf(between(w, inner_pasch(u, w, v, v, w), w))],
              [refute_0_26, refute_0_20])).

cnf(refute_0_28, plain,
    (w != inner_pasch(u, w, v, v, w) | inner_pasch(u, w, v, v, w) = w),
    inference(subst, [],
              [refute_0_8 :
               [bind(X0, $fot(w)),
                bind(Y0, $fot(inner_pasch(u, w, v, v, w)))]])).

cnf(refute_0_29, plain, (inner_pasch(u, w, v, v, w) = w),
    inference(resolve, [$cnf($equal(w, inner_pasch(u, w, v, v, w)))],
              [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (inner_pasch(u, w, v, v, w) != w |
     ~ between(v, inner_pasch(u, w, v, v, w), u) | between(v, w, u)),
    introduced(tautology,
               [equality,
                [$cnf(between(v, inner_pasch(u, w, v, v, w), u)), [1],
                 $fot(w)]])).

cnf(refute_0_31, plain,
    (~ between(v, inner_pasch(u, w, v, v, w), u) | between(v, w, u)),
    inference(resolve, [$cnf($equal(inner_pasch(u, w, v, v, w), w))],
              [refute_0_29, refute_0_30])).

cnf(refute_0_32, plain, (between(v, w, u)),
    inference(resolve, [$cnf(between(v, inner_pasch(u, w, v, v, w), u))],
              [refute_0_19, refute_0_31])).

cnf(refute_0_33, plain, (colinear(u, v, w)),
    inference(resolve, [$cnf(between(v, w, u))],
              [refute_0_32, refute_0_0])).

cnf(refute_0_34, plain, ($false),
    inference(resolve, [$cnf(colinear(u, v, w))],
              [refute_0_33, prove_uvw_colinear])).
% SZS output end CNFRefutation for GEO065-2.tptp

