%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for GRP130-4.003.tptp

% SZS output start CNFRefutation for GRP130-4.003.tptp
cnf(row_surjectivity, axiom,
    (~ group_element(X) | ~ group_element(Y) | product(e_1, X, Y) |
     product(e_2, X, Y) | product(e_3, X, Y))).

cnf(column_surjectivity, axiom,
    (~ group_element(X) | ~ group_element(Y) | product(X, e_1, Y) |
     product(X, e_2, Y) | product(X, e_3, Y))).

cnf(qg3_1, negated_conjecture,
    (product(X, Y, Z1) | ~ product(Z2, Y, X) | ~ product(X, Z1, Z2))).

cnf(qg3_2, negated_conjecture,
    (product(X, Z1, Z2) | ~ product(Z2, Y, X) | ~ product(X, Y, Z1))).

cnf(element_1, axiom, (group_element(e_1))).

cnf(element_2, axiom, (group_element(e_2))).

cnf(element_3, axiom, (group_element(e_3))).

cnf(e_1_is_not_e_3, axiom, (~ equalish(e_1, e_3))).

cnf(e_2_is_not_e_1, axiom, (~ equalish(e_2, e_1))).

cnf(e_3_is_not_e_1, axiom, (~ equalish(e_3, e_1))).

cnf(e_3_is_not_e_2, axiom, (~ equalish(e_3, e_2))).

cnf(product_total_function2, axiom,
    (~ product(X, Y, W) | ~ product(X, Y, Z) | equalish(W, Z))).

cnf(product_right_cancellation, axiom,
    (~ product(X, W, Y) | ~ product(X, Z, Y) | equalish(W, Z))).

cnf(product_left_cancellation, axiom,
    (~ product(W, Y, X) | ~ product(Z, Y, X) | equalish(W, Z))).

cnf(qg3, negated_conjecture,
    (~ product(X, Y, Z1) | ~ product(X, Z1, Z2) | product(Z2, Y, X))).

cnf(refute_0_0, plain, (~ product(e_1, e_2, e_1) | product(e_1, e_1, e_1)),
    inference(subst, [],
              [qg3_2 :
               [bind(X, $fot(e_1)), bind(Y, $fot(e_2)),
                bind(Z1, $fot(e_1)), bind(Z2, $fot(e_1))]])).

cnf(refute_0_1, plain, (~ product(e_1, e_3, e_1) | product(e_1, e_1, e_1)),
    inference(subst, [],
              [qg3_2 :
               [bind(X, $fot(e_1)), bind(Y, $fot(e_3)),
                bind(Z1, $fot(e_1)), bind(Z2, $fot(e_1))]])).

cnf(refute_0_2, plain,
    (~ group_element(e_1) | product(e_1, e_1, e_1) |
     product(e_1, e_2, e_1) | product(e_1, e_3, e_1)),
    inference(subst, [],
              [column_surjectivity :
               [bind(X, $fot(e_1)), bind(Y, $fot(e_1))]])).

cnf(refute_0_3, plain,
    (product(e_1, e_1, e_1) | product(e_1, e_2, e_1) |
     product(e_1, e_3, e_1)),
    inference(resolve, [$cnf(group_element(e_1))],
              [element_1, refute_0_2])).

cnf(refute_0_4, plain, (product(e_1, e_1, e_1) | product(e_1, e_2, e_1)),
    inference(resolve, [$cnf(product(e_1, e_3, e_1))],
              [refute_0_3, refute_0_1])).

cnf(refute_0_5, plain, (product(e_1, e_1, e_1)),
    inference(resolve, [$cnf(product(e_1, e_2, e_1))],
              [refute_0_4, refute_0_0])).

cnf(refute_0_6, plain,
    (~ product(X_34, e_1, e_1) | ~ product(e_1, e_1, e_1) |
     equalish(X_34, e_1)),
    inference(subst, [],
              [product_left_cancellation :
               [bind(W, $fot(X_34)), bind(X, $fot(e_1)),
                bind(Y, $fot(e_1)), bind(Z, $fot(e_1))]])).

cnf(refute_0_7, plain, (~ product(X_34, e_1, e_1) | equalish(X_34, e_1)),
    inference(resolve, [$cnf(product(e_1, e_1, e_1))],
              [refute_0_5, refute_0_6])).

cnf(refute_0_8, plain, (~ product(e_3, e_1, e_1) | equalish(e_3, e_1)),
    inference(subst, [], [refute_0_7 : [bind(X_34, $fot(e_3))]])).

cnf(refute_0_9, plain, (~ product(e_3, e_1, e_3) | product(e_3, e_1, e_1)),
    inference(subst, [],
              [qg3_1 :
               [bind(X, $fot(e_3)), bind(Y, $fot(e_1)),
                bind(Z1, $fot(e_1)), bind(Z2, $fot(e_3))]])).

cnf(refute_0_10, plain,
    (~ product(Z, e_2, e_3) | ~ product(e_1, e_2, e_3) | equalish(e_1, Z)),
    inference(subst, [],
              [product_left_cancellation :
               [bind(W, $fot(e_1)), bind(X, $fot(e_3)),
                bind(Y, $fot(e_2))]])).

cnf(refute_0_11, plain, (~ product(e_2, e_1, e_1) | equalish(e_2, e_1)),
    inference(subst, [], [refute_0_7 : [bind(X_34, $fot(e_2))]])).

cnf(refute_0_12, plain,
    (~ product(e_2, e_1, e_2) | product(e_2, e_1, e_1)),
    inference(subst, [],
              [qg3_1 :
               [bind(X, $fot(e_2)), bind(Y, $fot(e_1)),
                bind(Z1, $fot(e_1)), bind(Z2, $fot(e_2))]])).

cnf(refute_0_13, plain,
    (~ product(e_2, e_3, e_2) | product(e_2, e_2, e_2)),
    inference(subst, [],
              [qg3_2 :
               [bind(X, $fot(e_2)), bind(Y, $fot(e_3)),
                bind(Z1, $fot(e_2)), bind(Z2, $fot(e_2))]])).

cnf(refute_0_14, plain,
    (~ group_element(e_2) | product(e_2, e_1, e_2) |
     product(e_2, e_2, e_2) | product(e_2, e_3, e_2)),
    inference(subst, [],
              [column_surjectivity :
               [bind(X, $fot(e_2)), bind(Y, $fot(e_2))]])).

cnf(refute_0_15, plain,
    (product(e_2, e_1, e_2) | product(e_2, e_2, e_2) |
     product(e_2, e_3, e_2)),
    inference(resolve, [$cnf(group_element(e_2))],
              [element_2, refute_0_14])).

cnf(refute_0_16, plain, (product(e_2, e_1, e_2) | product(e_2, e_2, e_2)),
    inference(resolve, [$cnf(product(e_2, e_3, e_2))],
              [refute_0_15, refute_0_13])).

cnf(refute_0_17, plain,
    (~ product(e_2, e_2, X_18) | ~ product(e_2, e_2, e_2) |
     equalish(X_18, e_2)),
    inference(subst, [],
              [product_total_function2 :
               [bind(W, $fot(X_18)), bind(X, $fot(e_2)),
                bind(Y, $fot(e_2)), bind(Z, $fot(e_2))]])).

cnf(refute_0_18, plain,
    (~ product(e_2, e_2, X_18) | equalish(X_18, e_2) |
     product(e_2, e_1, e_2)),
    inference(resolve, [$cnf(product(e_2, e_2, e_2))],
              [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain,
    (~ product(e_2, e_2, e_3) | equalish(e_3, e_2) |
     product(e_2, e_1, e_2)),
    inference(subst, [], [refute_0_18 : [bind(X_18, $fot(e_3))]])).

cnf(refute_0_20, plain,
    (~ product(e_3, e_2, e_2) | product(e_2, e_2, e_3)),
    inference(subst, [],
              [qg3 :
               [bind(X, $fot(e_3)), bind(Y, $fot(e_2)),
                bind(Z1, $fot(e_2)), bind(Z2, $fot(e_2))]])).

cnf(refute_0_21, plain,
    (~ product(e_3, e_2, e_3) | product(e_3, e_2, e_2)),
    inference(subst, [],
              [qg3_1 :
               [bind(X, $fot(e_3)), bind(Y, $fot(e_2)),
                bind(Z1, $fot(e_2)), bind(Z2, $fot(e_3))]])).

cnf(refute_0_22, plain,
    (~ group_element(X_106) | ~ group_element(e_2) |
     product(e_1, e_2, X_106) | product(e_2, e_2, X_106) |
     product(e_3, e_2, X_106)),
    inference(subst, [],
              [row_surjectivity :
               [bind(X, $fot(e_2)), bind(Y, $fot(X_106))]])).

cnf(refute_0_23, plain,
    (~ group_element(X_106) | product(e_1, e_2, X_106) |
     product(e_2, e_2, X_106) | product(e_3, e_2, X_106)),
    inference(resolve, [$cnf(group_element(e_2))],
              [element_2, refute_0_22])).

cnf(refute_0_24, plain,
    (~ group_element(e_3) | product(e_1, e_2, e_3) |
     product(e_2, e_2, e_3) | product(e_3, e_2, e_3)),
    inference(subst, [], [refute_0_23 : [bind(X_106, $fot(e_3))]])).

cnf(refute_0_25, plain,
    (product(e_1, e_2, e_3) | product(e_2, e_2, e_3) |
     product(e_3, e_2, e_3)),
    inference(resolve, [$cnf(group_element(e_3))],
              [element_3, refute_0_24])).

cnf(refute_0_26, plain,
    (product(e_1, e_2, e_3) | product(e_2, e_2, e_3) |
     product(e_3, e_2, e_2)),
    inference(resolve, [$cnf(product(e_3, e_2, e_3))],
              [refute_0_25, refute_0_21])).

cnf(refute_0_27, plain, (product(e_1, e_2, e_3) | product(e_2, e_2, e_3)),
    inference(resolve, [$cnf(product(e_3, e_2, e_2))],
              [refute_0_26, refute_0_20])).

cnf(refute_0_28, plain,
    (equalish(e_3, e_2) | product(e_1, e_2, e_3) | product(e_2, e_1, e_2)),
    inference(resolve, [$cnf(product(e_2, e_2, e_3))],
              [refute_0_27, refute_0_19])).

cnf(refute_0_29, plain, (product(e_1, e_2, e_3) | product(e_2, e_1, e_2)),
    inference(resolve, [$cnf(equalish(e_3, e_2))],
              [refute_0_28, e_3_is_not_e_2])).

cnf(refute_0_30, plain, (product(e_1, e_2, e_3) | product(e_2, e_1, e_1)),
    inference(resolve, [$cnf(product(e_2, e_1, e_2))],
              [refute_0_29, refute_0_12])).

cnf(refute_0_31, plain, (equalish(e_2, e_1) | product(e_1, e_2, e_3)),
    inference(resolve, [$cnf(product(e_2, e_1, e_1))],
              [refute_0_30, refute_0_11])).

cnf(refute_0_32, plain, (product(e_1, e_2, e_3)),
    inference(resolve, [$cnf(equalish(e_2, e_1))],
              [refute_0_31, e_2_is_not_e_1])).

cnf(refute_0_33, plain, (~ product(Z, e_2, e_3) | equalish(e_1, Z)),
    inference(resolve, [$cnf(product(e_1, e_2, e_3))],
              [refute_0_32, refute_0_10])).

cnf(refute_0_34, plain, (~ product(e_3, e_2, e_3) | equalish(e_1, e_3)),
    inference(subst, [], [refute_0_33 : [bind(Z, $fot(e_3))]])).

cnf(refute_0_35, plain,
    (~ group_element(e_3) | product(e_3, e_1, e_3) |
     product(e_3, e_2, e_3) | product(e_3, e_3, e_3)),
    inference(subst, [],
              [column_surjectivity :
               [bind(X, $fot(e_3)), bind(Y, $fot(e_3))]])).

cnf(refute_0_36, plain,
    (product(e_3, e_1, e_3) | product(e_3, e_2, e_3) |
     product(e_3, e_3, e_3)),
    inference(resolve, [$cnf(group_element(e_3))],
              [element_3, refute_0_35])).

cnf(refute_0_37, plain,
    (~ product(e_3, e_3, X_50) | ~ product(e_3, e_3, e_3) |
     product(e_3, X_50, e_3)),
    inference(subst, [],
              [qg3_2 :
               [bind(X, $fot(e_3)), bind(Y, $fot(e_3)),
                bind(Z1, $fot(X_50)), bind(Z2, $fot(e_3))]])).

cnf(refute_0_38, plain,
    (~ product(e_3, e_3, X_50) | product(e_3, X_50, e_3) |
     product(e_3, e_1, e_3) | product(e_3, e_2, e_3)),
    inference(resolve, [$cnf(product(e_3, e_3, e_3))],
              [refute_0_36, refute_0_37])).

cnf(refute_0_39, plain,
    (~ product(e_3, e_3, e_1) | product(e_3, e_1, e_3) |
     product(e_3, e_2, e_3)),
    inference(subst, [], [refute_0_38 : [bind(X_50, $fot(e_1))]])).

cnf(refute_0_40, plain,
    (~ product(e_1, Y, e_2) | ~ product(e_1, e_2, e_3) |
     product(e_3, Y, e_1)),
    inference(subst, [],
              [qg3 :
               [bind(X, $fot(e_1)), bind(Z1, $fot(e_2)),
                bind(Z2, $fot(e_3))]])).

cnf(refute_0_41, plain, (~ product(e_1, Y, e_2) | product(e_3, Y, e_1)),
    inference(resolve, [$cnf(product(e_1, e_2, e_3))],
              [refute_0_32, refute_0_40])).

cnf(refute_0_42, plain,
    (~ product(e_1, e_3, e_2) | product(e_3, e_3, e_1)),
    inference(subst, [], [refute_0_41 : [bind(Y, $fot(e_3))]])).

cnf(refute_0_43, plain,
    (~ product(e_2, X_26, e_2) | ~ product(e_2, e_2, e_2) |
     equalish(X_26, e_2)),
    inference(subst, [],
              [product_right_cancellation :
               [bind(W, $fot(X_26)), bind(X, $fot(e_2)),
                bind(Y, $fot(e_2)), bind(Z, $fot(e_2))]])).

cnf(refute_0_44, plain,
    (~ product(e_2, X_26, e_2) | equalish(X_26, e_2) |
     product(e_2, e_1, e_2)),
    inference(resolve, [$cnf(product(e_2, e_2, e_2))],
              [refute_0_16, refute_0_43])).

cnf(refute_0_45, plain,
    (~ product(e_2, e_3, e_2) | equalish(e_3, e_2) |
     product(e_2, e_1, e_2)),
    inference(subst, [], [refute_0_44 : [bind(X_26, $fot(e_3))]])).

cnf(refute_0_46, plain,
    (~ product(e_3, e_3, e_2) | product(e_3, e_1, e_3) |
     product(e_3, e_2, e_3)),
    inference(subst, [], [refute_0_38 : [bind(X_50, $fot(e_2))]])).

cnf(refute_0_47, plain,
    (~ group_element(X_106) | ~ group_element(e_3) |
     product(e_1, e_3, X_106) | product(e_2, e_3, X_106) |
     product(e_3, e_3, X_106)),
    inference(subst, [],
              [row_surjectivity :
               [bind(X, $fot(e_3)), bind(Y, $fot(X_106))]])).

cnf(refute_0_48, plain,
    (~ group_element(X_106) | product(e_1, e_3, X_106) |
     product(e_2, e_3, X_106) | product(e_3, e_3, X_106)),
    inference(resolve, [$cnf(group_element(e_3))],
              [element_3, refute_0_47])).

cnf(refute_0_49, plain,
    (~ group_element(e_2) | product(e_1, e_3, e_2) |
     product(e_2, e_3, e_2) | product(e_3, e_3, e_2)),
    inference(subst, [], [refute_0_48 : [bind(X_106, $fot(e_2))]])).

cnf(refute_0_50, plain,
    (product(e_1, e_3, e_2) | product(e_2, e_3, e_2) |
     product(e_3, e_3, e_2)),
    inference(resolve, [$cnf(group_element(e_2))],
              [element_2, refute_0_49])).

cnf(refute_0_51, plain,
    (product(e_1, e_3, e_2) | product(e_2, e_3, e_2) |
     product(e_3, e_1, e_3) | product(e_3, e_2, e_3)),
    inference(resolve, [$cnf(product(e_3, e_3, e_2))],
              [refute_0_50, refute_0_46])).

cnf(refute_0_52, plain,
    (equalish(e_1, e_3) | product(e_1, e_3, e_2) | product(e_2, e_3, e_2) |
     product(e_3, e_1, e_3)),
    inference(resolve, [$cnf(product(e_3, e_2, e_3))],
              [refute_0_51, refute_0_34])).

cnf(refute_0_53, plain,
    (product(e_1, e_3, e_2) | product(e_2, e_3, e_2) |
     product(e_3, e_1, e_3)),
    inference(resolve, [$cnf(equalish(e_1, e_3))],
              [refute_0_52, e_1_is_not_e_3])).

cnf(refute_0_54, plain,
    (product(e_1, e_3, e_2) | product(e_2, e_3, e_2) |
     product(e_3, e_1, e_1)),
    inference(resolve, [$cnf(product(e_3, e_1, e_3))],
              [refute_0_53, refute_0_9])).

cnf(refute_0_55, plain,
    (equalish(e_3, e_1) | product(e_1, e_3, e_2) | product(e_2, e_3, e_2)),
    inference(resolve, [$cnf(product(e_3, e_1, e_1))],
              [refute_0_54, refute_0_8])).

cnf(refute_0_56, plain, (product(e_1, e_3, e_2) | product(e_2, e_3, e_2)),
    inference(resolve, [$cnf(equalish(e_3, e_1))],
              [refute_0_55, e_3_is_not_e_1])).

cnf(refute_0_57, plain,
    (equalish(e_3, e_2) | product(e_1, e_3, e_2) | product(e_2, e_1, e_2)),
    inference(resolve, [$cnf(product(e_2, e_3, e_2))],
              [refute_0_56, refute_0_45])).

cnf(refute_0_58, plain, (product(e_1, e_3, e_2) | product(e_2, e_1, e_2)),
    inference(resolve, [$cnf(equalish(e_3, e_2))],
              [refute_0_57, e_3_is_not_e_2])).

cnf(refute_0_59, plain, (product(e_1, e_3, e_2) | product(e_2, e_1, e_1)),
    inference(resolve, [$cnf(product(e_2, e_1, e_2))],
              [refute_0_58, refute_0_12])).

cnf(refute_0_60, plain, (equalish(e_2, e_1) | product(e_1, e_3, e_2)),
    inference(resolve, [$cnf(product(e_2, e_1, e_1))],
              [refute_0_59, refute_0_11])).

cnf(refute_0_61, plain, (product(e_1, e_3, e_2)),
    inference(resolve, [$cnf(equalish(e_2, e_1))],
              [refute_0_60, e_2_is_not_e_1])).

cnf(refute_0_62, plain, (product(e_3, e_3, e_1)),
    inference(resolve, [$cnf(product(e_1, e_3, e_2))],
              [refute_0_61, refute_0_42])).

cnf(refute_0_63, plain, (product(e_3, e_1, e_3) | product(e_3, e_2, e_3)),
    inference(resolve, [$cnf(product(e_3, e_3, e_1))],
              [refute_0_62, refute_0_39])).

cnf(refute_0_64, plain, (equalish(e_1, e_3) | product(e_3, e_1, e_3)),
    inference(resolve, [$cnf(product(e_3, e_2, e_3))],
              [refute_0_63, refute_0_34])).

cnf(refute_0_65, plain, (product(e_3, e_1, e_3)),
    inference(resolve, [$cnf(equalish(e_1, e_3))],
              [refute_0_64, e_1_is_not_e_3])).

cnf(refute_0_66, plain, (product(e_3, e_1, e_1)),
    inference(resolve, [$cnf(product(e_3, e_1, e_3))],
              [refute_0_65, refute_0_9])).

cnf(refute_0_67, plain, (equalish(e_3, e_1)),
    inference(resolve, [$cnf(product(e_3, e_1, e_1))],
              [refute_0_66, refute_0_8])).

cnf(refute_0_68, plain, ($false),
    inference(resolve, [$cnf(equalish(e_3, e_1))],
              [refute_0_67, e_3_is_not_e_1])).
% SZS output end CNFRefutation for GRP130-4.003.tptp

