%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for GRP043-2.tptp

% SZS output start CNFRefutation for GRP043-2.tptp
cnf(left_identity, axiom, (product(identity, X, X))).

cnf(total_function2, axiom,
    (~ product(X, Y, Z) | ~ product(X, Y, W) | equalish(Z, W))).

cnf(product_substitution3, axiom,
    (~ equalish(X, Y) | ~ product(W, Z, X) | product(W, Z, Y))).

cnf(a_equals_b, hypothesis, (equalish(a, b))).

cnf(b_equals_c, hypothesis, (equalish(b, c))).

cnf(prove_transitivity, negated_conjecture, (~ equalish(a, c))).

cnf(refute_0_0, plain, (product(identity, X_52, X_52)),
    inference(subst, [], [left_identity : [bind(X, $fot(X_52))]])).

cnf(refute_0_1, plain,
    (~ product(identity, X_52, X_52) | ~ product(identity, X_52, X_55) |
     equalish(X_55, X_52)),
    inference(subst, [],
              [total_function2 :
               [bind(W, $fot(X_52)), bind(X, $fot(identity)),
                bind(Y, $fot(X_52)), bind(Z, $fot(X_55))]])).

cnf(refute_0_2, plain,
    (~ product(identity, X_52, X_55) | equalish(X_55, X_52)),
    inference(resolve, [$cnf(product(identity, X_52, X_52))],
              [refute_0_0, refute_0_1])).

cnf(refute_0_3, plain, (~ product(identity, c, a) | equalish(a, c)),
    inference(subst, [],
              [refute_0_2 : [bind(X_52, $fot(c)), bind(X_55, $fot(a))]])).

cnf(refute_0_4, plain, (product(identity, X_109, X_109)),
    inference(subst, [], [left_identity : [bind(X, $fot(X_109))]])).

cnf(refute_0_5, plain,
    (~ equalish(X_109, X_110) | ~ product(identity, X_109, X_109) |
     product(identity, X_109, X_110)),
    inference(subst, [],
              [product_substitution3 :
               [bind(W, $fot(identity)), bind(X, $fot(X_109)),
                bind(Y, $fot(X_110)), bind(Z, $fot(X_109))]])).

cnf(refute_0_6, plain,
    (~ equalish(X_109, X_110) | product(identity, X_109, X_110)),
    inference(resolve, [$cnf(product(identity, X_109, X_109))],
              [refute_0_4, refute_0_5])).

cnf(refute_0_7, plain, (~ equalish(c, a) | product(identity, c, a)),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_109, $fot(c)), bind(X_110, $fot(a))]])).

cnf(refute_0_8, plain, (~ equalish(b, a) | product(identity, b, a)),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_109, $fot(b)), bind(X_110, $fot(a))]])).

cnf(refute_0_9, plain, (~ product(identity, a, b) | equalish(b, a)),
    inference(subst, [],
              [refute_0_2 : [bind(X_52, $fot(a)), bind(X_55, $fot(b))]])).

cnf(refute_0_10, plain, (~ equalish(a, b) | product(identity, a, b)),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_109, $fot(a)), bind(X_110, $fot(b))]])).

cnf(refute_0_11, plain, (product(identity, a, b)),
    inference(resolve, [$cnf(equalish(a, b))], [a_equals_b, refute_0_10])).

cnf(refute_0_12, plain, (equalish(b, a)),
    inference(resolve, [$cnf(product(identity, a, b))],
              [refute_0_11, refute_0_9])).

cnf(refute_0_13, plain, (product(identity, b, a)),
    inference(resolve, [$cnf(equalish(b, a))], [refute_0_12, refute_0_8])).

cnf(refute_0_14, plain,
    (~ product(identity, b, W) | ~ product(identity, b, c) |
     equalish(c, W)),
    inference(subst, [],
              [total_function2 :
               [bind(X, $fot(identity)), bind(Y, $fot(b)),
                bind(Z, $fot(c))]])).

cnf(refute_0_15, plain, (~ equalish(b, c) | product(identity, b, c)),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_109, $fot(b)), bind(X_110, $fot(c))]])).

cnf(refute_0_16, plain, (product(identity, b, c)),
    inference(resolve, [$cnf(equalish(b, c))], [b_equals_c, refute_0_15])).

cnf(refute_0_17, plain, (~ product(identity, b, W) | equalish(c, W)),
    inference(resolve, [$cnf(product(identity, b, c))],
              [refute_0_16, refute_0_14])).

cnf(refute_0_18, plain, (~ product(identity, b, a) | equalish(c, a)),
    inference(subst, [], [refute_0_17 : [bind(W, $fot(a))]])).

cnf(refute_0_19, plain, (equalish(c, a)),
    inference(resolve, [$cnf(product(identity, b, a))],
              [refute_0_13, refute_0_18])).

cnf(refute_0_20, plain, (product(identity, c, a)),
    inference(resolve, [$cnf(equalish(c, a))], [refute_0_19, refute_0_7])).

cnf(refute_0_21, plain, (equalish(a, c)),
    inference(resolve, [$cnf(product(identity, c, a))],
              [refute_0_20, refute_0_3])).

cnf(refute_0_22, plain, ($false),
    inference(resolve, [$cnf(equalish(a, c))],
              [refute_0_21, prove_transitivity])).
% SZS output end CNFRefutation for GRP043-2.tptp

