%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for GRP017-1.tptp

% SZS output start CNFRefutation for GRP017-1.tptp
cnf(left_identity, axiom, (product(identity, X, X))).

cnf(right_identity, axiom, (product(X, identity, X))).

cnf(left_inverse, axiom, (product(inverse(X), X, identity))).

cnf(total_function1, axiom, (product(X, Y, multiply(X, Y)))).

cnf(total_function2, axiom,
    (~ product(X, Y, Z) | ~ product(X, Y, W) | Z = W)).

cnf(associativity1, axiom,
    (~ product(X, Y, U) | ~ product(Y, Z, V) | ~ product(U, Z, W) |
     product(X, V, W))).

cnf(a_times_b_is_identity, hypothesis, (product(a, b, identity))).

cnf(c_times_a_is_identity, hypothesis, (product(c, a, identity))).

cnf(prove_b_equals_c, negated_conjecture, (b != c)).

cnf(refute_0_0, plain, (product(X_53, X_54, multiply(X_53, X_54))),
    inference(subst, [],
              [total_function1 :
               [bind(X, $fot(X_53)), bind(Y, $fot(X_54))]])).

cnf(refute_0_1, plain,
    (~ product(X_53, X_54, X_55) |
     ~ product(X_53, X_54, multiply(X_53, X_54)) |
     X_55 = multiply(X_53, X_54)),
    inference(subst, [],
              [total_function2 :
               [bind(W, $fot(multiply(X_53, X_54))), bind(X, $fot(X_53)),
                bind(Y, $fot(X_54)), bind(Z, $fot(X_55))]])).

cnf(refute_0_2, plain,
    (~ product(X_53, X_54, X_55) | X_55 = multiply(X_53, X_54)),
    inference(resolve, [$cnf(product(X_53, X_54, multiply(X_53, X_54)))],
              [refute_0_0, refute_0_1])).

cnf(refute_0_3, plain,
    (~ product(c, identity, b) | b = multiply(c, identity)),
    inference(subst, [],
              [refute_0_2 :
               [bind(X_53, $fot(c)), bind(X_54, $fot(identity)),
                bind(X_55, $fot(b))]])).

cnf(refute_0_4, plain,
    (product(multiply(X_158, c), a, multiply(multiply(X_158, c), a))),
    inference(subst, [],
              [total_function1 :
               [bind(X, $fot(multiply(X_158, c))), bind(Y, $fot(a))]])).

cnf(refute_0_5, plain,
    (~ product(X_157, identity, multiply(multiply(X_157, c), a)) |
     multiply(multiply(X_157, c), a) = multiply(X_157, identity)),
    inference(subst, [],
              [refute_0_2 :
               [bind(X_53, $fot(X_157)), bind(X_54, $fot(identity)),
                bind(X_55, $fot(multiply(multiply(X_157, c), a)))]])).

cnf(refute_0_6, plain, (product(X_156, c, multiply(X_156, c))),
    inference(subst, [],
              [total_function1 :
               [bind(X, $fot(X_156)), bind(Y, $fot(c))]])).

cnf(refute_0_7, plain, (product(X_142, a, multiply(X_142, a))),
    inference(subst, [],
              [total_function1 :
               [bind(X, $fot(X_142)), bind(Y, $fot(a))]])).

cnf(refute_0_8, plain,
    (~ product(X_109, a, X_111) | ~ product(X_112, c, X_109) |
     ~ product(c, a, identity) | product(X_112, identity, X_111)),
    inference(subst, [],
              [associativity1 :
               [bind(U, $fot(X_109)), bind(V, $fot(identity)),
                bind(W, $fot(X_111)), bind(X, $fot(X_112)),
                bind(Y, $fot(c)), bind(Z, $fot(a))]])).

cnf(refute_0_9, plain,
    (~ product(X_109, a, X_111) | ~ product(X_112, c, X_109) |
     product(X_112, identity, X_111)),
    inference(resolve, [$cnf(product(c, a, identity))],
              [c_times_a_is_identity, refute_0_8])).

cnf(refute_0_10, plain,
    (~ product(X_142, a, multiply(X_142, a)) | ~ product(X_144, c, X_142) |
     product(X_144, identity, multiply(X_142, a))),
    inference(subst, [],
              [refute_0_9 :
               [bind(X_109, $fot(X_142)),
                bind(X_111, $fot(multiply(X_142, a))),
                bind(X_112, $fot(X_144))]])).

cnf(refute_0_11, plain,
    (~ product(X_144, c, X_142) |
     product(X_144, identity, multiply(X_142, a))),
    inference(resolve, [$cnf(product(X_142, a, multiply(X_142, a)))],
              [refute_0_7, refute_0_10])).

cnf(refute_0_12, plain,
    (~ product(X_156, c, multiply(X_156, c)) |
     product(X_156, identity, multiply(multiply(X_156, c), a))),
    inference(subst, [],
              [refute_0_11 :
               [bind(X_142, $fot(multiply(X_156, c))),
                bind(X_144, $fot(X_156))]])).

cnf(refute_0_13, plain,
    (product(X_156, identity, multiply(multiply(X_156, c), a))),
    inference(resolve, [$cnf(product(X_156, c, multiply(X_156, c)))],
              [refute_0_6, refute_0_12])).

cnf(refute_0_14, plain,
    (product(X_157, identity, multiply(multiply(X_157, c), a))),
    inference(subst, [], [refute_0_13 : [bind(X_156, $fot(X_157))]])).

cnf(refute_0_15, plain,
    (multiply(multiply(X_157, c), a) = multiply(X_157, identity)),
    inference(resolve,
              [$cnf(product(X_157, identity,
                      multiply(multiply(X_157, c), a)))],
              [refute_0_14, refute_0_5])).

cnf(refute_0_16, plain,
    (product(X_67, identity, multiply(X_67, identity))),
    inference(subst, [],
              [total_function1 :
               [bind(X, $fot(X_67)), bind(Y, $fot(identity))]])).

cnf(refute_0_17, plain,
    (product(identity, identity, multiply(identity, identity))),
    inference(subst, [],
              [total_function1 :
               [bind(X, $fot(identity)), bind(Y, $fot(identity))]])).

cnf(refute_0_18, plain, (product(X_7, identity, X_7)),
    inference(subst, [], [right_identity : [bind(X, $fot(X_7))]])).

cnf(refute_0_19, plain,
    (~ product(X_7, identity, X_7) | ~ product(identity, identity, X_6) |
     product(X_7, X_6, X_7)),
    inference(subst, [],
              [associativity1 :
               [bind(U, $fot(X_7)), bind(V, $fot(X_6)), bind(W, $fot(X_7)),
                bind(X, $fot(X_7)), bind(Y, $fot(identity)),
                bind(Z, $fot(identity))]])).

cnf(refute_0_20, plain,
    (~ product(identity, identity, X_6) | product(X_7, X_6, X_7)),
    inference(resolve, [$cnf(product(X_7, identity, X_7))],
              [refute_0_18, refute_0_19])).

cnf(refute_0_21, plain,
    (~ product(identity, identity, multiply(identity, identity)) |
     product(X_11, multiply(identity, identity), X_11)),
    inference(subst, [],
              [refute_0_20 :
               [bind(X_6, $fot(multiply(identity, identity))),
                bind(X_7, $fot(X_11))]])).

cnf(refute_0_22, plain,
    (product(X_11, multiply(identity, identity), X_11)),
    inference(resolve,
              [$cnf(product(identity, identity,
                      multiply(identity, identity)))],
              [refute_0_17, refute_0_21])).

cnf(refute_0_23, plain,
    (product(X_52, multiply(identity, identity), X_52)),
    inference(subst, [], [refute_0_22 : [bind(X_11, $fot(X_52))]])).

cnf(refute_0_24, plain,
    (~ product(X_52, multiply(identity, identity), X_52) |
     ~ product(X_52, multiply(identity, identity), X_55) | X_55 = X_52),
    inference(subst, [],
              [total_function2 :
               [bind(W, $fot(X_52)), bind(X, $fot(X_52)),
                bind(Y, $fot(multiply(identity, identity))),
                bind(Z, $fot(X_55))]])).

cnf(refute_0_25, plain,
    (~ product(X_52, multiply(identity, identity), X_55) | X_55 = X_52),
    inference(resolve,
              [$cnf(product(X_52, multiply(identity, identity), X_52))],
              [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (product(inverse(X_60), X_60, multiply(inverse(X_60), X_60))),
    inference(subst, [],
              [total_function1 :
               [bind(X, $fot(inverse(X_60))), bind(Y, $fot(X_60))]])).

cnf(refute_0_27, plain, (product(inverse(X_54), X_54, identity)),
    inference(subst, [], [left_inverse : [bind(X, $fot(X_54))]])).

cnf(refute_0_28, plain,
    (~ product(inverse(X_54), X_54, X_55) |
     ~ product(inverse(X_54), X_54, identity) | X_55 = identity),
    inference(subst, [],
              [total_function2 :
               [bind(W, $fot(identity)), bind(X, $fot(inverse(X_54))),
                bind(Y, $fot(X_54)), bind(Z, $fot(X_55))]])).

cnf(refute_0_29, plain,
    (~ product(inverse(X_54), X_54, X_55) | X_55 = identity),
    inference(resolve, [$cnf(product(inverse(X_54), X_54, identity))],
              [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (~ product(inverse(X_60), X_60, multiply(inverse(X_60), X_60)) |
     multiply(inverse(X_60), X_60) = identity),
    inference(subst, [],
              [refute_0_29 :
               [bind(X_54, $fot(X_60)),
                bind(X_55, $fot(multiply(inverse(X_60), X_60)))]])).

cnf(refute_0_31, plain, (multiply(inverse(X_60), X_60) = identity),
    inference(resolve,
              [$cnf(product(inverse(X_60), X_60,
                      multiply(inverse(X_60), X_60)))],
              [refute_0_26, refute_0_30])).

cnf(refute_0_32, plain, (multiply(inverse(identity), identity) = identity),
    inference(subst, [], [refute_0_31 : [bind(X_60, $fot(identity))]])).

cnf(refute_0_33, plain,
    (product(inverse(identity), identity, inverse(identity))),
    inference(subst, [],
              [right_identity : [bind(X, $fot(inverse(identity)))]])).

cnf(refute_0_34, plain,
    (~ product(inverse(identity), identity, inverse(identity)) |
     inverse(identity) = identity),
    inference(subst, [],
              [refute_0_29 :
               [bind(X_54, $fot(identity)),
                bind(X_55, $fot(inverse(identity)))]])).

cnf(refute_0_35, plain, (inverse(identity) = identity),
    inference(resolve,
              [$cnf(product(inverse(identity), identity,
                      inverse(identity)))], [refute_0_33, refute_0_34])).

cnf(refute_0_36, plain,
    (multiply(inverse(identity), identity) != identity |
     inverse(identity) != identity |
     multiply(identity, identity) = identity),
    introduced(tautology,
               [equality,
                [$cnf($equal(multiply(inverse(identity), identity),
                        identity)), [0, 0], $fot(identity)]])).

cnf(refute_0_37, plain,
    (multiply(inverse(identity), identity) != identity |
     multiply(identity, identity) = identity),
    inference(resolve, [$cnf($equal(inverse(identity), identity))],
              [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain, (multiply(identity, identity) = identity),
    inference(resolve,
              [$cnf($equal(multiply(inverse(identity), identity),
                      identity))], [refute_0_32, refute_0_37])).

cnf(refute_0_39, plain,
    (multiply(identity, identity) != identity |
     ~ product(X_52, identity, X_55) |
     product(X_52, multiply(identity, identity), X_55)),
    introduced(tautology,
               [equality,
                [$cnf(~ product(X_52, multiply(identity, identity), X_55)),
                 [1], $fot(identity)]])).

cnf(refute_0_40, plain,
    (~ product(X_52, identity, X_55) |
     product(X_52, multiply(identity, identity), X_55)),
    inference(resolve,
              [$cnf($equal(multiply(identity, identity), identity))],
              [refute_0_38, refute_0_39])).

cnf(refute_0_41, plain, (~ product(X_52, identity, X_55) | X_55 = X_52),
    inference(resolve,
              [$cnf(product(X_52, multiply(identity, identity), X_55))],
              [refute_0_40, refute_0_25])).

cnf(refute_0_42, plain,
    (~ product(X_67, identity, multiply(X_67, identity)) |
     multiply(X_67, identity) = X_67),
    inference(subst, [],
              [refute_0_41 :
               [bind(X_52, $fot(X_67)),
                bind(X_55, $fot(multiply(X_67, identity)))]])).

cnf(refute_0_43, plain, (multiply(X_67, identity) = X_67),
    inference(resolve,
              [$cnf(product(X_67, identity, multiply(X_67, identity)))],
              [refute_0_16, refute_0_42])).

cnf(refute_0_44, plain, (multiply(X_157, identity) = X_157),
    inference(subst, [], [refute_0_43 : [bind(X_67, $fot(X_157))]])).

cnf(refute_0_45, plain,
    (multiply(X_157, identity) != X_157 |
     multiply(multiply(X_157, c), a) != multiply(X_157, identity) |
     multiply(multiply(X_157, c), a) = X_157),
    introduced(tautology,
               [equality,
                [$cnf($equal(multiply(multiply(X_157, c), a),
                        multiply(X_157, identity))), [1], $fot(X_157)]])).

cnf(refute_0_46, plain,
    (multiply(multiply(X_157, c), a) != multiply(X_157, identity) |
     multiply(multiply(X_157, c), a) = X_157),
    inference(resolve, [$cnf($equal(multiply(X_157, identity), X_157))],
              [refute_0_44, refute_0_45])).

cnf(refute_0_47, plain, (multiply(multiply(X_157, c), a) = X_157),
    inference(resolve,
              [$cnf($equal(multiply(multiply(X_157, c), a),
                      multiply(X_157, identity)))],
              [refute_0_15, refute_0_46])).

cnf(refute_0_48, plain, (multiply(multiply(X_158, c), a) = X_158),
    inference(subst, [], [refute_0_47 : [bind(X_157, $fot(X_158))]])).

cnf(refute_0_49, plain,
    (multiply(multiply(X_158, c), a) != X_158 |
     ~ product(multiply(X_158, c), a, multiply(multiply(X_158, c), a)) |
     product(multiply(X_158, c), a, X_158)),
    introduced(tautology,
               [equality,
                [$cnf(product(multiply(X_158, c), a,
                        multiply(multiply(X_158, c), a))), [2],
                 $fot(X_158)]])).

cnf(refute_0_50, plain,
    (~ product(multiply(X_158, c), a, multiply(multiply(X_158, c), a)) |
     product(multiply(X_158, c), a, X_158)),
    inference(resolve,
              [$cnf($equal(multiply(multiply(X_158, c), a), X_158))],
              [refute_0_48, refute_0_49])).

cnf(refute_0_51, plain, (product(multiply(X_158, c), a, X_158)),
    inference(resolve,
              [$cnf(product(multiply(X_158, c), a,
                      multiply(multiply(X_158, c), a)))],
              [refute_0_4, refute_0_50])).

cnf(refute_0_52, plain, (product(multiply(identity, c), a, identity)),
    inference(subst, [], [refute_0_51 : [bind(X_158, $fot(identity))]])).

cnf(refute_0_53, plain, (product(identity, b, b)),
    inference(subst, [], [left_identity : [bind(X, $fot(b))]])).

cnf(refute_0_54, plain,
    (~ product(X_109, b, X_111) | ~ product(X_112, a, X_109) |
     ~ product(a, b, identity) | product(X_112, identity, X_111)),
    inference(subst, [],
              [associativity1 :
               [bind(U, $fot(X_109)), bind(V, $fot(identity)),
                bind(W, $fot(X_111)), bind(X, $fot(X_112)),
                bind(Y, $fot(a)), bind(Z, $fot(b))]])).

cnf(refute_0_55, plain,
    (~ product(X_109, b, X_111) | ~ product(X_112, a, X_109) |
     product(X_112, identity, X_111)),
    inference(resolve, [$cnf(product(a, b, identity))],
              [a_times_b_is_identity, refute_0_54])).

cnf(refute_0_56, plain,
    (~ product(X_173, a, identity) | ~ product(identity, b, b) |
     product(X_173, identity, b)),
    inference(subst, [],
              [refute_0_55 :
               [bind(X_109, $fot(identity)), bind(X_111, $fot(b)),
                bind(X_112, $fot(X_173))]])).

cnf(refute_0_57, plain,
    (~ product(X_173, a, identity) | product(X_173, identity, b)),
    inference(resolve, [$cnf(product(identity, b, b))],
              [refute_0_53, refute_0_56])).

cnf(refute_0_58, plain,
    (~ product(multiply(identity, c), a, identity) |
     product(multiply(identity, c), identity, b)),
    inference(subst, [],
              [refute_0_57 : [bind(X_173, $fot(multiply(identity, c)))]])).

cnf(refute_0_59, plain, (product(multiply(identity, c), identity, b)),
    inference(resolve, [$cnf(product(multiply(identity, c), a, identity))],
              [refute_0_52, refute_0_58])).

cnf(refute_0_60, plain,
    (product(identity, X_64, multiply(identity, X_64))),
    inference(subst, [],
              [total_function1 :
               [bind(X, $fot(identity)), bind(Y, $fot(X_64))]])).

cnf(refute_0_61, plain, (product(identity, X_52, X_52)),
    inference(subst, [], [left_identity : [bind(X, $fot(X_52))]])).

cnf(refute_0_62, plain,
    (~ product(identity, X_52, X_52) | ~ product(identity, X_52, X_55) |
     X_55 = X_52),
    inference(subst, [],
              [total_function2 :
               [bind(W, $fot(X_52)), bind(X, $fot(identity)),
                bind(Y, $fot(X_52)), bind(Z, $fot(X_55))]])).

cnf(refute_0_63, plain, (~ product(identity, X_52, X_55) | X_55 = X_52),
    inference(resolve, [$cnf(product(identity, X_52, X_52))],
              [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain,
    (~ product(identity, X_64, multiply(identity, X_64)) |
     multiply(identity, X_64) = X_64),
    inference(subst, [],
              [refute_0_63 :
               [bind(X_52, $fot(X_64)),
                bind(X_55, $fot(multiply(identity, X_64)))]])).

cnf(refute_0_65, plain, (multiply(identity, X_64) = X_64),
    inference(resolve,
              [$cnf(product(identity, X_64, multiply(identity, X_64)))],
              [refute_0_60, refute_0_64])).

cnf(refute_0_66, plain, (multiply(identity, c) = c),
    inference(subst, [], [refute_0_65 : [bind(X_64, $fot(c))]])).

cnf(refute_0_67, plain,
    (multiply(identity, c) != c |
     ~ product(multiply(identity, c), identity, b) |
     product(c, identity, b)),
    introduced(tautology,
               [equality,
                [$cnf(product(multiply(identity, c), identity, b)), [0],
                 $fot(c)]])).

cnf(refute_0_68, plain,
    (~ product(multiply(identity, c), identity, b) |
     product(c, identity, b)),
    inference(resolve, [$cnf($equal(multiply(identity, c), c))],
              [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain, (product(c, identity, b)),
    inference(resolve, [$cnf(product(multiply(identity, c), identity, b))],
              [refute_0_59, refute_0_68])).

cnf(refute_0_70, plain, (b = multiply(c, identity)),
    inference(resolve, [$cnf(product(c, identity, b))],
              [refute_0_69, refute_0_3])).

cnf(refute_0_71, plain, (multiply(c, identity) = c),
    inference(subst, [], [refute_0_43 : [bind(X_67, $fot(c))]])).

cnf(refute_0_72, plain,
    (multiply(c, identity) != c | b != multiply(c, identity) | b = c),
    introduced(tautology,
               [equality,
                [$cnf($equal(b, multiply(c, identity))), [1], $fot(c)]])).

cnf(refute_0_73, plain, (b != multiply(c, identity) | b = c),
    inference(resolve, [$cnf($equal(multiply(c, identity), c))],
              [refute_0_71, refute_0_72])).

cnf(refute_0_74, plain, (b = c),
    inference(resolve, [$cnf($equal(b, multiply(c, identity)))],
              [refute_0_70, refute_0_73])).

cnf(refute_0_75, plain, ($false),
    inference(resolve, [$cnf($equal(b, c))],
              [refute_0_74, prove_b_equals_c])).
% SZS output end CNFRefutation for GRP017-1.tptp

