%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for HEN008-4.tptp

% SZS output start CNFRefutation for HEN008-4.tptp
cnf(quotient_less_equal1, axiom,
    (~ less_equal(X, Y) | divide(X, Y) = zero)).

cnf(quotient_less_equal2, axiom,
    (divide(X, Y) != zero | less_equal(X, Y))).

cnf(quotient_property, axiom,
    (less_equal(divide(divide(X, Z), divide(Y, Z)),
       divide(divide(X, Y), Z)))).

cnf(zero_is_smallest, axiom, (less_equal(zero, X))).

cnf(less_equal_and_equal, axiom,
    (~ less_equal(X, Y) | ~ less_equal(Y, X) | X = Y)).

cnf(zero_divide_anything_is_zero, axiom, (divide(zero, X) = zero)).

cnf(a_LE_b, hypothesis, (less_equal(a, b))).

cnf(prove_a_divide_c_LE_b_divide_c, negated_conjecture,
    (~ less_equal(divide(a, c), divide(b, c)))).

cnf(refute_0_0, plain,
    (divide(divide(a, X_201), divide(b, X_201)) != zero |
     less_equal(divide(a, X_201), divide(b, X_201))),
    inference(subst, [],
              [quotient_less_equal2 :
               [bind(X, $fot(divide(a, X_201))),
                bind(Y, $fot(divide(b, X_201)))]])).

cnf(refute_0_1, plain,
    (less_equal(divide(divide(a, X_16), divide(b, X_16)),
       divide(divide(a, b), X_16))),
    inference(subst, [],
              [quotient_property :
               [bind(X, $fot(a)), bind(Y, $fot(b)),
                bind(Z, $fot(X_16))]])).

cnf(refute_0_2, plain, (~ less_equal(a, b) | divide(a, b) = zero),
    inference(subst, [],
              [quotient_less_equal1 :
               [bind(X, $fot(a)), bind(Y, $fot(b))]])).

cnf(refute_0_3, plain, (divide(a, b) = zero),
    inference(resolve, [$cnf(less_equal(a, b))], [a_LE_b, refute_0_2])).

cnf(refute_0_4, plain,
    (~
       less_equal(divide(divide(a, X_16), divide(b, X_16)),
         divide(divide(a, b), X_16)) | divide(a, b) != zero |
     less_equal(divide(divide(a, X_16), divide(b, X_16)),
       divide(zero, X_16))),
    introduced(tautology,
               [equality,
                [$cnf(less_equal(divide(divide(a, X_16), divide(b, X_16)),
                        divide(divide(a, b), X_16))), [1, 0],
                 $fot(zero)]])).

cnf(refute_0_5, plain,
    (~
       less_equal(divide(divide(a, X_16), divide(b, X_16)),
         divide(divide(a, b), X_16)) |
     less_equal(divide(divide(a, X_16), divide(b, X_16)),
       divide(zero, X_16))),
    inference(resolve, [$cnf($equal(divide(a, b), zero))],
              [refute_0_3, refute_0_4])).

cnf(refute_0_6, plain,
    (less_equal(divide(divide(a, X_16), divide(b, X_16)),
       divide(zero, X_16))),
    inference(resolve,
              [$cnf(less_equal(divide(divide(a, X_16), divide(b, X_16)),
                      divide(divide(a, b), X_16)))],
              [refute_0_1, refute_0_5])).

cnf(refute_0_7, plain, (divide(zero, X_16) = zero),
    inference(subst, [],
              [zero_divide_anything_is_zero : [bind(X, $fot(X_16))]])).

cnf(refute_0_8, plain,
    (~
       less_equal(divide(divide(a, X_16), divide(b, X_16)),
         divide(zero, X_16)) | divide(zero, X_16) != zero |
     less_equal(divide(divide(a, X_16), divide(b, X_16)), zero)),
    introduced(tautology,
               [equality,
                [$cnf(less_equal(divide(divide(a, X_16), divide(b, X_16)),
                        divide(zero, X_16))), [1], $fot(zero)]])).

cnf(refute_0_9, plain,
    (~
       less_equal(divide(divide(a, X_16), divide(b, X_16)),
         divide(zero, X_16)) |
     less_equal(divide(divide(a, X_16), divide(b, X_16)), zero)),
    inference(resolve, [$cnf($equal(divide(zero, X_16), zero))],
              [refute_0_7, refute_0_8])).

cnf(refute_0_10, plain,
    (less_equal(divide(divide(a, X_16), divide(b, X_16)), zero)),
    inference(resolve,
              [$cnf(less_equal(divide(divide(a, X_16), divide(b, X_16)),
                      divide(zero, X_16)))], [refute_0_6, refute_0_9])).

cnf(refute_0_11, plain,
    (~ less_equal(divide(divide(a, X_16), divide(b, X_16)), zero) |
     ~ less_equal(zero, divide(divide(a, X_16), divide(b, X_16))) |
     divide(divide(a, X_16), divide(b, X_16)) = zero),
    inference(subst, [],
              [less_equal_and_equal :
               [bind(X, $fot(divide(divide(a, X_16), divide(b, X_16)))),
                bind(Y, $fot(zero))]])).

cnf(refute_0_12, plain,
    (~ less_equal(zero, divide(divide(a, X_16), divide(b, X_16))) |
     divide(divide(a, X_16), divide(b, X_16)) = zero),
    inference(resolve,
              [$cnf(less_equal(divide(divide(a, X_16), divide(b, X_16)),
                      zero))], [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain,
    (less_equal(zero, divide(divide(a, X_16), divide(b, X_16)))),
    inference(subst, [],
              [zero_is_smallest :
               [bind(X,
                     $fot(divide(divide(a, X_16), divide(b, X_16))))]])).

cnf(refute_0_14, plain, (divide(divide(a, X_16), divide(b, X_16)) = zero),
    inference(resolve,
              [$cnf(less_equal(zero,
                      divide(divide(a, X_16), divide(b, X_16))))],
              [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (divide(divide(a, X_201), divide(b, X_201)) = zero),
    inference(subst, [], [refute_0_14 : [bind(X_16, $fot(X_201))]])).

cnf(refute_0_16, plain,
    (divide(divide(a, X_201), divide(b, X_201)) != zero | zero != zero |
     divide(divide(a, X_201), divide(b, X_201)) = zero),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(divide(divide(a, X_201), divide(b, X_201)),
                        zero)), [0], $fot(zero)]])).

cnf(refute_0_17, plain,
    (zero != zero | divide(divide(a, X_201), divide(b, X_201)) = zero),
    inference(resolve,
              [$cnf($equal(divide(divide(a, X_201), divide(b, X_201)),
                      zero))], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    (zero != zero | less_equal(divide(a, X_201), divide(b, X_201))),
    inference(resolve,
              [$cnf($equal(divide(divide(a, X_201), divide(b, X_201)),
                      zero))], [refute_0_17, refute_0_0])).

cnf(refute_0_19, plain, (zero = zero),
    introduced(tautology, [refl, [$fot(zero)]])).

cnf(refute_0_20, plain, (less_equal(divide(a, X_201), divide(b, X_201))),
    inference(resolve, [$cnf($equal(zero, zero))],
              [refute_0_19, refute_0_18])).

cnf(refute_0_21, plain, (less_equal(divide(a, c), divide(b, c))),
    inference(subst, [], [refute_0_20 : [bind(X_201, $fot(c))]])).

cnf(refute_0_22, plain, ($false),
    inference(resolve, [$cnf(less_equal(divide(a, c), divide(b, c)))],
              [refute_0_21, prove_a_divide_c_LE_b_divide_c])).
% SZS output end CNFRefutation for HEN008-4.tptp

