%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for LCL009-1.tptp

% SZS output start CNFRefutation for LCL009-1.tptp
cnf(condensed_detachment, axiom,
    (~ is_a_theorem(equivalent(X, Y)) | ~ is_a_theorem(X) |
     is_a_theorem(Y))).

cnf(yql, axiom,
    (is_a_theorem(equivalent(equivalent(X, Y),
         equivalent(equivalent(Z, Y), equivalent(X, Z)))))).

cnf(prove_ec_5, negated_conjecture,
    (~
       is_a_theorem(equivalent(equivalent(equivalent(a, b), c),
           equivalent(a, equivalent(b, c)))))).

cnf(refute_0_0, plain,
    (~
       is_a_theorem(equivalent(X_11,
           equivalent(equivalent(X_10, X_9), equivalent(X_8, X_10)))) |
     ~
       is_a_theorem(equivalent(equivalent(X_11,
             equivalent(equivalent(X_10, X_9), equivalent(X_8, X_10))),
           equivalent(equivalent(X_8, X_9), X_11))) |
     is_a_theorem(equivalent(equivalent(X_8, X_9), X_11))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X,
                     $fot(equivalent(X_11,
                            equivalent(equivalent(X_10, X_9),
                              equivalent(X_8, X_10))))),
                bind(Y, $fot(equivalent(equivalent(X_8, X_9), X_11)))]])).

cnf(refute_0_1, plain,
    (~ is_a_theorem(equivalent(X, Y)) |
     ~
       is_a_theorem(equivalent(equivalent(X, Y),
           equivalent(equivalent(Z, Y), equivalent(X, Z)))) |
     is_a_theorem(equivalent(equivalent(Z, Y), equivalent(X, Z)))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(equivalent(X, Y))),
                bind(Y,
                     $fot(equivalent(equivalent(Z, Y),
                            equivalent(X, Z))))]])).

cnf(refute_0_2, plain,
    (~ is_a_theorem(equivalent(X, Y)) |
     is_a_theorem(equivalent(equivalent(Z, Y), equivalent(X, Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X, Y),
                        equivalent(equivalent(Z, Y), equivalent(X, Z)))))],
              [yql, refute_0_1])).

cnf(refute_0_3, plain,
    (~
       is_a_theorem(equivalent(equivalent(X, Y),
           equivalent(equivalent(Z, Y), equivalent(X, Z)))) |
     is_a_theorem(equivalent(equivalent(X_7,
           equivalent(equivalent(Z, Y), equivalent(X, Z))),
         equivalent(equivalent(X, Y), X_7)))),
    inference(subst, [],
              [refute_0_2 :
               [bind(X, $fot(equivalent(X, Y))),
                bind(Y,
                     $fot(equivalent(equivalent(Z, Y), equivalent(X, Z)))),
                bind(Z, $fot(X_7))]])).

cnf(refute_0_4, plain,
    (is_a_theorem(equivalent(equivalent(X_7,
           equivalent(equivalent(Z, Y), equivalent(X, Z))),
         equivalent(equivalent(X, Y), X_7)))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X, Y),
                        equivalent(equivalent(Z, Y), equivalent(X, Z)))))],
              [yql, refute_0_3])).

cnf(refute_0_5, plain,
    (is_a_theorem(equivalent(equivalent(X_11,
           equivalent(equivalent(X_10, X_9), equivalent(X_8, X_10))),
         equivalent(equivalent(X_8, X_9), X_11)))),
    inference(subst, [],
              [refute_0_4 :
               [bind(X, $fot(X_8)), bind(Y, $fot(X_9)),
                bind(Z, $fot(X_10)), bind(X_7, $fot(X_11))]])).

cnf(refute_0_6, plain,
    (~
       is_a_theorem(equivalent(X_11,
           equivalent(equivalent(X_10, X_9), equivalent(X_8, X_10)))) |
     is_a_theorem(equivalent(equivalent(X_8, X_9), X_11))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_11,
                          equivalent(equivalent(X_10, X_9),
                            equivalent(X_8, X_10))),
                        equivalent(equivalent(X_8, X_9), X_11))))],
              [refute_0_5, refute_0_0])).

cnf(refute_0_7, plain,
    (~
       is_a_theorem(equivalent(equivalent(X_994, equivalent(X_992, X_9)),
           equivalent(equivalent(X_992, X_9),
             equivalent(equivalent(X_994, X_992), X_992)))) |
     is_a_theorem(equivalent(equivalent(equivalent(X_994, X_992), X_9),
         equivalent(X_994, equivalent(X_992, X_9))))),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_10, $fot(X_992)),
                bind(X_11,
                     $fot(equivalent(X_994, equivalent(X_992, X_9)))),
                bind(X_8, $fot(equivalent(X_994, X_992)))]])).

cnf(refute_0_8, plain,
    (~ is_a_theorem(equivalent(X_50, X_51)) |
     ~
       is_a_theorem(equivalent(equivalent(X_50, X_51),
           equivalent(X_51, X_50))) |
     is_a_theorem(equivalent(X_51, X_50))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(equivalent(X_50, X_51))),
                bind(Y, $fot(equivalent(X_51, X_50)))]])).

cnf(refute_0_9, plain,
    (~ is_a_theorem(equivalent(X_49, X_49)) |
     is_a_theorem(equivalent(equivalent(Z, X_49), equivalent(X_49, Z)))),
    inference(subst, [],
              [refute_0_2 : [bind(X, $fot(X_49)), bind(Y, $fot(X_49))]])).

cnf(refute_0_10, plain,
    (~
       is_a_theorem(equivalent(equivalent(equivalent(X_27, X_26), X_29),
           equivalent(equivalent(equivalent(X_27, X_26), X_27),
             equivalent(equivalent(X_26, X_29),
               equivalent(X_27, X_26))))) |
     is_a_theorem(equivalent(equivalent(equivalent(X_26, X_29), X_27),
         equivalent(equivalent(X_27, X_26), X_29)))),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_10, $fot(equivalent(X_27, X_26))),
                bind(X_11, $fot(equivalent(equivalent(X_27, X_26), X_29))),
                bind(X_8, $fot(equivalent(X_26, X_29))),
                bind(X_9, $fot(X_27))]])).

cnf(refute_0_11, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_14, X_12),
           equivalent(equivalent(Z, X_15), equivalent(X_12, Z))),
         equivalent(equivalent(X_12, X_15), equivalent(X_14, X_12))))),
    inference(subst, [],
              [refute_0_4 :
               [bind(X, $fot(X_12)), bind(Y, $fot(X_15)),
                bind(X_7, $fot(equivalent(X_14, X_12)))]])).

cnf(refute_0_12, plain,
    (~
       is_a_theorem(equivalent(equivalent(equivalent(X_14, X_12),
             equivalent(equivalent(Z, X_15), equivalent(X_12, Z))),
           equivalent(equivalent(X_12, X_15), equivalent(X_14, X_12)))) |
     is_a_theorem(equivalent(equivalent(X_14, X_15),
         equivalent(equivalent(X_14, X_12),
           equivalent(equivalent(Z, X_15), equivalent(X_12, Z)))))),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_10, $fot(X_12)),
                bind(X_11,
                     $fot(equivalent(equivalent(X_14, X_12),
                            equivalent(equivalent(Z, X_15),
                              equivalent(X_12, Z))))),
                bind(X_8, $fot(X_14)), bind(X_9, $fot(X_15))]])).

cnf(refute_0_13, plain,
    (is_a_theorem(equivalent(equivalent(X_14, X_15),
         equivalent(equivalent(X_14, X_12),
           equivalent(equivalent(Z, X_15), equivalent(X_12, Z)))))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(equivalent(X_14,
                            X_12),
                          equivalent(equivalent(Z, X_15),
                            equivalent(X_12, Z))),
                        equivalent(equivalent(X_12, X_15),
                          equivalent(X_14, X_12)))))],
              [refute_0_11, refute_0_12])).

cnf(refute_0_14, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_27, X_26), X_29),
         equivalent(equivalent(equivalent(X_27, X_26), X_27),
           equivalent(equivalent(X_26, X_29), equivalent(X_27, X_26)))))),
    inference(subst, [],
              [refute_0_13 :
               [bind(Z, $fot(X_26)), bind(X_12, $fot(X_27)),
                bind(X_14, $fot(equivalent(X_27, X_26))),
                bind(X_15, $fot(X_29))]])).

cnf(refute_0_15, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_26, X_29), X_27),
         equivalent(equivalent(X_27, X_26), X_29)))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(equivalent(X_27,
                            X_26), X_29),
                        equivalent(equivalent(equivalent(X_27, X_26),
                            X_27),
                          equivalent(equivalent(X_26, X_29),
                            equivalent(X_27, X_26))))))],
              [refute_0_14, refute_0_10])).

cnf(refute_0_16, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_29, X_29), X_29),
         equivalent(equivalent(X_29, X_29), X_29)))),
    inference(subst, [],
              [refute_0_15 :
               [bind(X_26, $fot(X_29)), bind(X_27, $fot(X_29))]])).

cnf(refute_0_17, plain,
    (~ is_a_theorem(equivalent(X_24, X_24)) |
     ~
       is_a_theorem(equivalent(equivalent(X_24, X_24),
           equivalent(X_25, X_25))) |
     is_a_theorem(equivalent(X_25, X_25))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(equivalent(X_24, X_24))),
                bind(Y, $fot(equivalent(X_25, X_25)))]])).

cnf(refute_0_18, plain,
    (~
       is_a_theorem(equivalent(equivalent(X_22, X_22),
           equivalent(equivalent(X_22, X_21), equivalent(X_21, X_22)))) |
     is_a_theorem(equivalent(equivalent(X_21, X_21),
         equivalent(X_22, X_22)))),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_10, $fot(X_22)),
                bind(X_11, $fot(equivalent(X_22, X_22))),
                bind(X_8, $fot(X_21)), bind(X_9, $fot(X_21))]])).

cnf(refute_0_19, plain,
    (~
       is_a_theorem(equivalent(equivalent(equivalent(X_8, X_19),
             equivalent(X_19, X_20)),
           equivalent(equivalent(X_19, X_20), equivalent(X_8, X_19)))) |
     is_a_theorem(equivalent(equivalent(X_8, X_20),
         equivalent(equivalent(X_8, X_19), equivalent(X_19, X_20))))),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_10, $fot(X_19)),
                bind(X_11,
                     $fot(equivalent(equivalent(X_8, X_19),
                            equivalent(X_19, X_20)))),
                bind(X_9, $fot(X_20))]])).

cnf(refute_0_20, plain,
    (~
       is_a_theorem(equivalent(equivalent(X_16, X_17),
           equivalent(X_16, X_17))) |
     is_a_theorem(equivalent(equivalent(Z, equivalent(X_16, X_17)),
         equivalent(equivalent(X_16, X_17), Z)))),
    inference(subst, [],
              [refute_0_2 :
               [bind(X, $fot(equivalent(X_16, X_17))),
                bind(Y, $fot(equivalent(X_16, X_17)))]])).

cnf(refute_0_21, plain,
    (is_a_theorem(equivalent(equivalent(X_14, X_15),
         equivalent(equivalent(X_12, X_15), equivalent(X_14, X_12))))),
    inference(subst, [],
              [yql :
               [bind(X, $fot(X_14)), bind(Y, $fot(X_15)),
                bind(Z, $fot(X_12))]])).

cnf(refute_0_22, plain,
    (~
       is_a_theorem(equivalent(equivalent(X_14, X_15),
           equivalent(equivalent(X_12, X_15), equivalent(X_14, X_12)))) |
     is_a_theorem(equivalent(equivalent(X_14, X_15),
         equivalent(X_14, X_15)))),
    inference(subst, [],
              [refute_0_6 :
               [bind(X_10, $fot(X_12)),
                bind(X_11, $fot(equivalent(X_14, X_15))),
                bind(X_8, $fot(X_14)), bind(X_9, $fot(X_15))]])).

cnf(refute_0_23, plain,
    (is_a_theorem(equivalent(equivalent(X_14, X_15),
         equivalent(X_14, X_15)))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_14, X_15),
                        equivalent(equivalent(X_12, X_15),
                          equivalent(X_14, X_12)))))],
              [refute_0_21, refute_0_22])).

cnf(refute_0_24, plain,
    (is_a_theorem(equivalent(equivalent(X_16, X_17),
         equivalent(X_16, X_17)))),
    inference(subst, [],
              [refute_0_23 :
               [bind(X_14, $fot(X_16)), bind(X_15, $fot(X_17))]])).

cnf(refute_0_25, plain,
    (is_a_theorem(equivalent(equivalent(Z, equivalent(X_16, X_17)),
         equivalent(equivalent(X_16, X_17), Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_16, X_17),
                        equivalent(X_16, X_17))))],
              [refute_0_24, refute_0_20])).

cnf(refute_0_26, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_8, X_19),
           equivalent(X_19, X_20)),
         equivalent(equivalent(X_19, X_20), equivalent(X_8, X_19))))),
    inference(subst, [],
              [refute_0_25 :
               [bind(Z, $fot(equivalent(X_8, X_19))),
                bind(X_16, $fot(X_19)), bind(X_17, $fot(X_20))]])).

cnf(refute_0_27, plain,
    (is_a_theorem(equivalent(equivalent(X_8, X_20),
         equivalent(equivalent(X_8, X_19), equivalent(X_19, X_20))))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(equivalent(X_8,
                            X_19), equivalent(X_19, X_20)),
                        equivalent(equivalent(X_19, X_20),
                          equivalent(X_8, X_19)))))],
              [refute_0_26, refute_0_19])).

cnf(refute_0_28, plain,
    (is_a_theorem(equivalent(equivalent(X_22, X_22),
         equivalent(equivalent(X_22, X_21), equivalent(X_21, X_22))))),
    inference(subst, [],
              [refute_0_27 :
               [bind(X_19, $fot(X_21)), bind(X_20, $fot(X_22)),
                bind(X_8, $fot(X_22))]])).

cnf(refute_0_29, plain,
    (is_a_theorem(equivalent(equivalent(X_21, X_21),
         equivalent(X_22, X_22)))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_22, X_22),
                        equivalent(equivalent(X_22, X_21),
                          equivalent(X_21, X_22)))))],
              [refute_0_28, refute_0_18])).

cnf(refute_0_30, plain,
    (is_a_theorem(equivalent(equivalent(X_24, X_24),
         equivalent(X_25, X_25)))),
    inference(subst, [],
              [refute_0_29 :
               [bind(X_21, $fot(X_24)), bind(X_22, $fot(X_25))]])).

cnf(refute_0_31, plain,
    (~ is_a_theorem(equivalent(X_24, X_24)) |
     is_a_theorem(equivalent(X_25, X_25))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_24, X_24),
                        equivalent(X_25, X_25))))],
              [refute_0_30, refute_0_17])).

cnf(refute_0_32, plain,
    (~
       is_a_theorem(equivalent(equivalent(equivalent(X_29, X_29), X_29),
           equivalent(equivalent(X_29, X_29), X_29))) |
     is_a_theorem(equivalent(X_48, X_48))),
    inference(subst, [],
              [refute_0_31 :
               [bind(X_24, $fot(equivalent(equivalent(X_29, X_29), X_29))),
                bind(X_25, $fot(X_48))]])).

cnf(refute_0_33, plain, (is_a_theorem(equivalent(X_48, X_48))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(equivalent(X_29,
                            X_29), X_29),
                        equivalent(equivalent(X_29, X_29), X_29))))],
              [refute_0_16, refute_0_32])).

cnf(refute_0_34, plain, (is_a_theorem(equivalent(X_49, X_49))),
    inference(subst, [], [refute_0_33 : [bind(X_48, $fot(X_49))]])).

cnf(refute_0_35, plain,
    (is_a_theorem(equivalent(equivalent(Z, X_49), equivalent(X_49, Z)))),
    inference(resolve, [$cnf(is_a_theorem(equivalent(X_49, X_49)))],
              [refute_0_34, refute_0_9])).

cnf(refute_0_36, plain,
    (is_a_theorem(equivalent(equivalent(X_50, X_51),
         equivalent(X_51, X_50)))),
    inference(subst, [],
              [refute_0_35 :
               [bind(Z, $fot(X_50)), bind(X_49, $fot(X_51))]])).

cnf(refute_0_37, plain,
    (~ is_a_theorem(equivalent(X_50, X_51)) |
     is_a_theorem(equivalent(X_51, X_50))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_50, X_51),
                        equivalent(X_51, X_50))))],
              [refute_0_36, refute_0_8])).

cnf(refute_0_38, plain,
    (~
       is_a_theorem(equivalent(equivalent(equivalent(X_108, X_109), X_109),
           X_108)) |
     is_a_theorem(equivalent(X_108,
         equivalent(equivalent(X_108, X_109), X_109)))),
    inference(subst, [],
              [refute_0_37 :
               [bind(X_50,
                     $fot(equivalent(equivalent(X_108, X_109), X_109))),
                bind(X_51, $fot(X_108))]])).

cnf(refute_0_39, plain,
    (is_a_theorem(equivalent(equivalent(X_99, X_101),
         equivalent(X_101, X_99)))),
    inference(subst, [],
              [refute_0_35 :
               [bind(Z, $fot(X_99)), bind(X_49, $fot(X_101))]])).

cnf(refute_0_40, plain,
    (~ is_a_theorem(equivalent(equivalent(X_30, X_32), X_31)) |
     ~
       is_a_theorem(equivalent(equivalent(equivalent(X_30, X_32), X_31),
           equivalent(equivalent(X_31, X_30), X_32))) |
     is_a_theorem(equivalent(equivalent(X_31, X_30), X_32))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(equivalent(equivalent(X_30, X_32), X_31))),
                bind(Y,
                     $fot(equivalent(equivalent(X_31, X_30), X_32)))]])).

cnf(refute_0_41, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_30, X_32), X_31),
         equivalent(equivalent(X_31, X_30), X_32)))),
    inference(subst, [],
              [refute_0_15 :
               [bind(X_26, $fot(X_30)), bind(X_27, $fot(X_31)),
                bind(X_29, $fot(X_32))]])).

cnf(refute_0_42, plain,
    (~ is_a_theorem(equivalent(equivalent(X_30, X_32), X_31)) |
     is_a_theorem(equivalent(equivalent(X_31, X_30), X_32))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(equivalent(X_30,
                            X_32), X_31),
                        equivalent(equivalent(X_31, X_30), X_32))))],
              [refute_0_41, refute_0_40])).

cnf(refute_0_43, plain,
    (~
       is_a_theorem(equivalent(equivalent(X_99, X_101),
           equivalent(X_101, X_99))) |
     is_a_theorem(equivalent(equivalent(equivalent(X_101, X_99), X_99),
         X_101))),
    inference(subst, [],
              [refute_0_42 :
               [bind(X_30, $fot(X_99)),
                bind(X_31, $fot(equivalent(X_101, X_99))),
                bind(X_32, $fot(X_101))]])).

cnf(refute_0_44, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_101, X_99), X_99),
         X_101))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_99, X_101),
                        equivalent(X_101, X_99))))],
              [refute_0_39, refute_0_43])).

cnf(refute_0_45, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_108, X_109), X_109),
         X_108))),
    inference(subst, [],
              [refute_0_44 :
               [bind(X_101, $fot(X_108)), bind(X_99, $fot(X_109))]])).

cnf(refute_0_46, plain,
    (is_a_theorem(equivalent(X_108,
         equivalent(equivalent(X_108, X_109), X_109)))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(equivalent(X_108,
                            X_109), X_109), X_108)))],
              [refute_0_45, refute_0_38])).

cnf(refute_0_47, plain,
    (is_a_theorem(equivalent(X_826,
         equivalent(equivalent(X_826, X_109), X_109)))),
    inference(subst, [], [refute_0_46 : [bind(X_108, $fot(X_826))]])).

cnf(refute_0_48, plain,
    (~ is_a_theorem(equivalent(X_23, X_22)) |
     ~
       is_a_theorem(equivalent(equivalent(X_23, X_22),
           equivalent(equivalent(X_23, X_21), equivalent(X_21, X_22)))) |
     is_a_theorem(equivalent(equivalent(X_23, X_21),
         equivalent(X_21, X_22)))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(equivalent(X_23, X_22))),
                bind(Y,
                     $fot(equivalent(equivalent(X_23, X_21),
                            equivalent(X_21, X_22))))]])).

cnf(refute_0_49, plain,
    (is_a_theorem(equivalent(equivalent(X_23, X_22),
         equivalent(equivalent(X_23, X_21), equivalent(X_21, X_22))))),
    inference(subst, [],
              [refute_0_27 :
               [bind(X_19, $fot(X_21)), bind(X_20, $fot(X_22)),
                bind(X_8, $fot(X_23))]])).

cnf(refute_0_50, plain,
    (~ is_a_theorem(equivalent(X_23, X_22)) |
     is_a_theorem(equivalent(equivalent(X_23, X_21),
         equivalent(X_21, X_22)))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_23, X_22),
                        equivalent(equivalent(X_23, X_21),
                          equivalent(X_21, X_22)))))],
              [refute_0_49, refute_0_48])).

cnf(refute_0_51, plain,
    (~
       is_a_theorem(equivalent(X_826,
           equivalent(equivalent(X_826, X_109), X_109))) |
     is_a_theorem(equivalent(equivalent(X_826, X_824),
         equivalent(X_824, equivalent(equivalent(X_826, X_109), X_109))))),
    inference(subst, [],
              [refute_0_50 :
               [bind(X_21, $fot(X_824)),
                bind(X_22,
                     $fot(equivalent(equivalent(X_826, X_109), X_109))),
                bind(X_23, $fot(X_826))]])).

cnf(refute_0_52, plain,
    (is_a_theorem(equivalent(equivalent(X_826, X_824),
         equivalent(X_824, equivalent(equivalent(X_826, X_109), X_109))))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(X_826,
                        equivalent(equivalent(X_826, X_109), X_109))))],
              [refute_0_47, refute_0_51])).

cnf(refute_0_53, plain,
    (is_a_theorem(equivalent(equivalent(X_994, equivalent(X_992, X_9)),
         equivalent(equivalent(X_992, X_9),
           equivalent(equivalent(X_994, X_992), X_992))))),
    inference(subst, [],
              [refute_0_52 :
               [bind(X_109, $fot(X_992)),
                bind(X_824, $fot(equivalent(X_992, X_9))),
                bind(X_826, $fot(X_994))]])).

cnf(refute_0_54, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(X_994, X_992), X_9),
         equivalent(X_994, equivalent(X_992, X_9))))),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(X_994,
                          equivalent(X_992, X_9)),
                        equivalent(equivalent(X_992, X_9),
                          equivalent(equivalent(X_994, X_992), X_992)))))],
              [refute_0_53, refute_0_7])).

cnf(refute_0_55, plain,
    (is_a_theorem(equivalent(equivalent(equivalent(a, b), c),
         equivalent(a, equivalent(b, c))))),
    inference(subst, [],
              [refute_0_54 :
               [bind(X_9, $fot(c)), bind(X_992, $fot(b)),
                bind(X_994, $fot(a))]])).

cnf(refute_0_56, plain, ($false),
    inference(resolve,
              [$cnf(is_a_theorem(equivalent(equivalent(equivalent(a, b),
                          c), equivalent(a, equivalent(b, c)))))],
              [refute_0_55, prove_ec_5])).
% SZS output end CNFRefutation for LCL009-1.tptp

