%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for LCL055-1.tptp

% SZS output start CNFRefutation for LCL055-1.tptp
cnf(condensed_detachment, axiom,
    (~ is_a_theorem(implies(X, Y)) | ~ is_a_theorem(X) | is_a_theorem(Y))).

cnf(cn_1, axiom,
    (is_a_theorem(implies(implies(X, Y),
         implies(implies(Y, Z), implies(X, Z)))))).

cnf(cn_2, axiom, (is_a_theorem(implies(implies(not(X), X), X)))).

cnf(cn_3, axiom, (is_a_theorem(implies(X, implies(not(X), Y))))).

cnf(prove_cn_37, negated_conjecture,
    (~
       is_a_theorem(implies(implies(implies(a, b), c),
           implies(not(a), c))))).

cnf(refute_0_0, plain,
    (~ is_a_theorem(implies(X, Y)) |
     ~
       is_a_theorem(implies(implies(X, Y),
           implies(implies(Y, Z), implies(X, Z)))) |
     is_a_theorem(implies(implies(Y, Z), implies(X, Z)))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(implies(X, Y))),
                bind(Y, $fot(implies(implies(Y, Z), implies(X, Z))))]])).

cnf(refute_0_1, plain,
    (~ is_a_theorem(implies(X, Y)) |
     is_a_theorem(implies(implies(Y, Z), implies(X, Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X, Y),
                        implies(implies(Y, Z), implies(X, Z)))))],
              [cn_1, refute_0_0])).

cnf(refute_0_2, plain,
    (~ is_a_theorem(implies(not(X_7369), implies(X_7369, X_7368))) |
     is_a_theorem(implies(implies(implies(X_7369, X_7368), Z),
         implies(not(X_7369), Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(not(X_7369))),
                bind(Y, $fot(implies(X_7369, X_7368)))]])).

cnf(refute_0_3, plain,
    (~
       is_a_theorem(implies(X_1940,
           implies(not(not(X_1942)), not(X_1942)))) |
     ~
       is_a_theorem(implies(implies(X_1940,
             implies(not(not(X_1942)), not(X_1942))),
           implies(X_1942, implies(X_1940, X_1941)))) |
     is_a_theorem(implies(X_1942, implies(X_1940, X_1941)))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X,
                     $fot(implies(X_1940,
                            implies(not(not(X_1942)), not(X_1942))))),
                bind(Y,
                     $fot(implies(X_1942, implies(X_1940, X_1941))))]])).

cnf(refute_0_4, plain,
    (~
       is_a_theorem(implies(X_231,
           implies(implies(not(not(X_231)), not(X_231)), X_230))) |
     is_a_theorem(implies(implies(implies(implies(not(not(X_231)),
               not(X_231)), X_230), Z), implies(X_231, Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(X_231)),
                bind(Y,
                     $fot(implies(implies(not(not(X_231)), not(X_231)),
                            X_230)))]])).

cnf(refute_0_5, plain,
    (is_a_theorem(implies(implies(not(X_28), X_28), X_28))),
    inference(subst, [], [cn_2 : [bind(X, $fot(X_28))]])).

cnf(refute_0_6, plain,
    (~ is_a_theorem(implies(implies(not(X_28), X_28), X_28)) |
     is_a_theorem(implies(implies(X_28, X_29),
         implies(implies(not(X_28), X_28), X_29)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(implies(not(X_28), X_28))),
                bind(Y, $fot(X_28)), bind(Z, $fot(X_29))]])).

cnf(refute_0_7, plain,
    (is_a_theorem(implies(implies(X_28, X_29),
         implies(implies(not(X_28), X_28), X_29)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(not(X_28), X_28),
                        X_28)))], [refute_0_5, refute_0_6])).

cnf(refute_0_8, plain,
    (is_a_theorem(implies(implies(not(X_219), X_218),
         implies(implies(not(not(X_219)), not(X_219)), X_218)))),
    inference(subst, [],
              [refute_0_7 :
               [bind(X_28, $fot(not(X_219))), bind(X_29, $fot(X_218))]])).

cnf(refute_0_9, plain,
    (~
       is_a_theorem(implies(implies(implies(not(X_33), X_32), X_34),
           implies(X_33, X_34))) |
     ~ is_a_theorem(implies(implies(not(X_33), X_32), X_34)) |
     is_a_theorem(implies(X_33, X_34))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(implies(implies(not(X_33), X_32), X_34))),
                bind(Y, $fot(implies(X_33, X_34)))]])).

cnf(refute_0_10, plain,
    (is_a_theorem(implies(X_27, implies(not(X_27), Y)))),
    inference(subst, [], [cn_3 : [bind(X, $fot(X_27))]])).

cnf(refute_0_11, plain,
    (~ is_a_theorem(implies(X_27, implies(not(X_27), Y))) |
     is_a_theorem(implies(implies(implies(not(X_27), Y), X_29),
         implies(X_27, X_29)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(X_27)), bind(Y, $fot(implies(not(X_27), Y))),
                bind(Z, $fot(X_29))]])).

cnf(refute_0_12, plain,
    (is_a_theorem(implies(implies(implies(not(X_27), Y), X_29),
         implies(X_27, X_29)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(X_27, implies(not(X_27), Y))))],
              [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain,
    (is_a_theorem(implies(implies(implies(not(X_33), X_32), X_34),
         implies(X_33, X_34)))),
    inference(subst, [],
              [refute_0_12 :
               [bind(Y, $fot(X_32)), bind(X_27, $fot(X_33)),
                bind(X_29, $fot(X_34))]])).

cnf(refute_0_14, plain,
    (~ is_a_theorem(implies(implies(not(X_33), X_32), X_34)) |
     is_a_theorem(implies(X_33, X_34))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(not(X_33), X_32),
                          X_34), implies(X_33, X_34))))],
              [refute_0_13, refute_0_9])).

cnf(refute_0_15, plain,
    (~
       is_a_theorem(implies(implies(not(X_219), X_218),
           implies(implies(not(not(X_219)), not(X_219)), X_218))) |
     is_a_theorem(implies(X_219,
         implies(implies(not(not(X_219)), not(X_219)), X_218)))),
    inference(subst, [],
              [refute_0_14 :
               [bind(X_32, $fot(X_218)), bind(X_33, $fot(X_219)),
                bind(X_34,
                     $fot(implies(implies(not(not(X_219)), not(X_219)),
                            X_218)))]])).

cnf(refute_0_16, plain,
    (is_a_theorem(implies(X_219,
         implies(implies(not(not(X_219)), not(X_219)), X_218)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(not(X_219), X_218),
                        implies(implies(not(not(X_219)), not(X_219)),
                          X_218))))], [refute_0_8, refute_0_15])).

cnf(refute_0_17, plain,
    (is_a_theorem(implies(X_231,
         implies(implies(not(not(X_231)), not(X_231)), X_230)))),
    inference(subst, [],
              [refute_0_16 :
               [bind(X_218, $fot(X_230)), bind(X_219, $fot(X_231))]])).

cnf(refute_0_18, plain,
    (is_a_theorem(implies(implies(implies(implies(not(not(X_231)),
               not(X_231)), X_230), Z), implies(X_231, Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(X_231,
                        implies(implies(not(not(X_231)), not(X_231)),
                          X_230))))], [refute_0_17, refute_0_4])).

cnf(refute_0_19, plain,
    (is_a_theorem(implies(implies(implies(implies(not(not(X_231)),
               not(X_231)), X_1772), implies(X_1770, X_1772)),
         implies(X_231, implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_18 :
               [bind(Z, $fot(implies(X_1770, X_1772))),
                bind(X_230, $fot(X_1772))]])).

cnf(refute_0_20, plain,
    (~
       is_a_theorem(implies(implies(implies(X_39, X_40),
             implies(X_38, X_40)), X_41)) |
     ~
       is_a_theorem(implies(implies(implies(implies(X_39, X_40),
               implies(X_38, X_40)), X_41),
           implies(implies(X_38, X_39), X_41))) |
     is_a_theorem(implies(implies(X_38, X_39), X_41))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X,
                     $fot(implies(implies(implies(X_39, X_40),
                              implies(X_38, X_40)), X_41))),
                bind(Y, $fot(implies(implies(X_38, X_39), X_41)))]])).

cnf(refute_0_21, plain,
    (~
       is_a_theorem(implies(implies(X, Y),
           implies(implies(Y, Z), implies(X, Z)))) |
     is_a_theorem(implies(implies(implies(implies(Y, Z), implies(X, Z)),
           X_29), implies(implies(X, Y), X_29)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(implies(X, Y))),
                bind(Y, $fot(implies(implies(Y, Z), implies(X, Z)))),
                bind(Z, $fot(X_29))]])).

cnf(refute_0_22, plain,
    (is_a_theorem(implies(implies(implies(implies(Y, Z), implies(X, Z)),
           X_29), implies(implies(X, Y), X_29)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X, Y),
                        implies(implies(Y, Z), implies(X, Z)))))],
              [cn_1, refute_0_21])).

cnf(refute_0_23, plain,
    (is_a_theorem(implies(implies(implies(implies(X_39, X_40),
             implies(X_38, X_40)), X_41),
         implies(implies(X_38, X_39), X_41)))),
    inference(subst, [],
              [refute_0_22 :
               [bind(X, $fot(X_38)), bind(Y, $fot(X_39)),
                bind(Z, $fot(X_40)), bind(X_29, $fot(X_41))]])).

cnf(refute_0_24, plain,
    (~
       is_a_theorem(implies(implies(implies(X_39, X_40),
             implies(X_38, X_40)), X_41)) |
     is_a_theorem(implies(implies(X_38, X_39), X_41))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(implies(X_39,
                              X_40), implies(X_38, X_40)), X_41),
                        implies(implies(X_38, X_39), X_41))))],
              [refute_0_23, refute_0_20])).

cnf(refute_0_25, plain,
    (~
       is_a_theorem(implies(implies(implies(implies(not(not(X_231)),
                 not(X_231)), X_1772), implies(X_1770, X_1772)),
           implies(X_231, implies(X_1770, X_1772)))) |
     is_a_theorem(implies(implies(X_1770,
           implies(not(not(X_231)), not(X_231))),
         implies(X_231, implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_38, $fot(X_1770)),
                bind(X_39, $fot(implies(not(not(X_231)), not(X_231)))),
                bind(X_40, $fot(X_1772)),
                bind(X_41,
                     $fot(implies(X_231, implies(X_1770, X_1772))))]])).

cnf(refute_0_26, plain,
    (is_a_theorem(implies(implies(X_1770,
           implies(not(not(X_231)), not(X_231))),
         implies(X_231, implies(X_1770, X_1772))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(implies(not(not(X_231)),
                              not(X_231)), X_1772),
                          implies(X_1770, X_1772)),
                        implies(X_231, implies(X_1770, X_1772)))))],
              [refute_0_19, refute_0_25])).

cnf(refute_0_27, plain,
    (is_a_theorem(implies(implies(X_1940,
           implies(not(not(X_1942)), not(X_1942))),
         implies(X_1942, implies(X_1940, X_1941))))),
    inference(subst, [],
              [refute_0_26 :
               [bind(X_1770, $fot(X_1940)), bind(X_1772, $fot(X_1941)),
                bind(X_231, $fot(X_1942))]])).

cnf(refute_0_28, plain,
    (~
       is_a_theorem(implies(X_1940,
           implies(not(not(X_1942)), not(X_1942)))) |
     is_a_theorem(implies(X_1942, implies(X_1940, X_1941)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_1940,
                          implies(not(not(X_1942)), not(X_1942))),
                        implies(X_1942, implies(X_1940, X_1941)))))],
              [refute_0_27, refute_0_3])).

cnf(refute_0_29, plain,
    (~
       is_a_theorem(implies(X_7366,
           implies(not(not(not(X_7366))), not(not(X_7366))))) |
     is_a_theorem(implies(not(X_7366), implies(X_7366, X_1941)))),
    inference(subst, [],
              [refute_0_28 :
               [bind(X_1940, $fot(X_7366)),
                bind(X_1942, $fot(not(X_7366)))]])).

cnf(refute_0_30, plain,
    (~ is_a_theorem(implies(X_1776, not(X_1778))) |
     ~
       is_a_theorem(implies(implies(X_1776, not(X_1778)),
           implies(X_1778, implies(X_1776, X_1777)))) |
     is_a_theorem(implies(X_1778, implies(X_1776, X_1777)))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(implies(X_1776, not(X_1778)))),
                bind(Y,
                     $fot(implies(X_1778, implies(X_1776, X_1777))))]])).

cnf(refute_0_31, plain,
    (is_a_theorem(implies(implies(implies(not(X_27), X_1772),
           implies(X_1770, X_1772)),
         implies(X_27, implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_12 :
               [bind(Y, $fot(X_1772)),
                bind(X_29, $fot(implies(X_1770, X_1772)))]])).

cnf(refute_0_32, plain,
    (~
       is_a_theorem(implies(implies(implies(not(X_27), X_1772),
             implies(X_1770, X_1772)),
           implies(X_27, implies(X_1770, X_1772)))) |
     is_a_theorem(implies(implies(X_1770, not(X_27)),
         implies(X_27, implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_38, $fot(X_1770)), bind(X_39, $fot(not(X_27))),
                bind(X_40, $fot(X_1772)),
                bind(X_41,
                     $fot(implies(X_27, implies(X_1770, X_1772))))]])).

cnf(refute_0_33, plain,
    (is_a_theorem(implies(implies(X_1770, not(X_27)),
         implies(X_27, implies(X_1770, X_1772))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(not(X_27),
                            X_1772), implies(X_1770, X_1772)),
                        implies(X_27, implies(X_1770, X_1772)))))],
              [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain,
    (is_a_theorem(implies(implies(X_1776, not(X_1778)),
         implies(X_1778, implies(X_1776, X_1777))))),
    inference(subst, [],
              [refute_0_33 :
               [bind(X_1770, $fot(X_1776)), bind(X_1772, $fot(X_1777)),
                bind(X_27, $fot(X_1778))]])).

cnf(refute_0_35, plain,
    (~ is_a_theorem(implies(X_1776, not(X_1778))) |
     is_a_theorem(implies(X_1778, implies(X_1776, X_1777)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_1776, not(X_1778)),
                        implies(X_1778, implies(X_1776, X_1777)))))],
              [refute_0_34, refute_0_30])).

cnf(refute_0_36, plain,
    (~ is_a_theorem(implies(not(not(not(X_1778))), not(X_1778))) |
     is_a_theorem(implies(X_1778,
         implies(not(not(not(X_1778))), X_1777)))),
    inference(subst, [],
              [refute_0_35 :
               [bind(X_1776, $fot(not(not(not(X_1778)))))]])).

cnf(refute_0_37, plain,
    (~
       is_a_theorem(implies(implies(implies(not(not(X_1167)), X_1168),
             X_1167), X_1167)) |
     ~ is_a_theorem(implies(implies(not(not(X_1167)), X_1168), X_1167)) |
     is_a_theorem(X_1167)),
    inference(subst, [],
              [condensed_detachment :
               [bind(X,
                     $fot(implies(implies(not(not(X_1167)), X_1168),
                            X_1167))), bind(Y, $fot(X_1167))]])).

cnf(refute_0_38, plain,
    (is_a_theorem(implies(implies(not(X_7), X_7), X_7))),
    inference(subst, [], [cn_2 : [bind(X, $fot(X_7))]])).

cnf(refute_0_39, plain,
    (~ is_a_theorem(implies(implies(not(X_7), X_7), X_7)) |
     ~ is_a_theorem(implies(not(X_7), X_7)) | is_a_theorem(X_7)),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(implies(not(X_7), X_7))),
                bind(Y, $fot(X_7))]])).

cnf(refute_0_40, plain,
    (~ is_a_theorem(implies(not(X_7), X_7)) | is_a_theorem(X_7)),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(not(X_7), X_7), X_7)))],
              [refute_0_38, refute_0_39])).

cnf(refute_0_41, plain,
    (~
       is_a_theorem(implies(not(implies(implies(implies(not(not(X_1164)),
                   X_1166), X_1164), X_1164)),
           implies(implies(implies(not(not(X_1164)), X_1166), X_1164),
             X_1164))) |
     is_a_theorem(implies(implies(implies(not(not(X_1164)), X_1166),
           X_1164), X_1164))),
    inference(subst, [],
              [refute_0_40 :
               [bind(X_7,
                     $fot(implies(implies(implies(not(not(X_1164)),
                                X_1166), X_1164), X_1164)))]])).

cnf(refute_0_42, plain,
    (~
       is_a_theorem(implies(implies(implies(not(X_33), X_32), X_34),
           implies(X_33, X_34))) |
     is_a_theorem(implies(implies(implies(X_33, X_34), Z),
         implies(implies(implies(not(X_33), X_32), X_34), Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(implies(implies(not(X_33), X_32), X_34))),
                bind(Y, $fot(implies(X_33, X_34)))]])).

cnf(refute_0_43, plain,
    (is_a_theorem(implies(implies(implies(X_33, X_34), Z),
         implies(implies(implies(not(X_33), X_32), X_34), Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(not(X_33), X_32),
                          X_34), implies(X_33, X_34))))],
              [refute_0_13, refute_0_42])).

cnf(refute_0_44, plain,
    (is_a_theorem(implies(implies(implies(not(X_697), X_697), X_698),
         implies(implies(implies(not(not(X_697)), X_32), X_697), X_698)))),
    inference(subst, [],
              [refute_0_43 :
               [bind(Z, $fot(X_698)), bind(X_33, $fot(not(X_697))),
                bind(X_34, $fot(X_697))]])).

cnf(refute_0_45, plain,
    (~
       is_a_theorem(implies(implies(implies(implies(not(X_62), X_62),
               X_63), X_61), implies(implies(X_62, X_63), X_61))) |
     ~
       is_a_theorem(implies(implies(implies(not(X_62), X_62), X_63),
           X_61)) | is_a_theorem(implies(implies(X_62, X_63), X_61))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X,
                     $fot(implies(implies(implies(not(X_62), X_62), X_63),
                            X_61))),
                bind(Y, $fot(implies(implies(X_62, X_63), X_61)))]])).

cnf(refute_0_46, plain,
    (~
       is_a_theorem(implies(implies(X_30, X_31),
           implies(implies(not(X_30), X_30), X_31))) |
     is_a_theorem(implies(implies(implies(implies(not(X_30), X_30), X_31),
           Z), implies(implies(X_30, X_31), Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(implies(X_30, X_31))),
                bind(Y, $fot(implies(implies(not(X_30), X_30), X_31)))]])).

cnf(refute_0_47, plain,
    (is_a_theorem(implies(implies(X_30, X_31),
         implies(implies(not(X_30), X_30), X_31)))),
    inference(subst, [],
              [refute_0_7 :
               [bind(X_28, $fot(X_30)), bind(X_29, $fot(X_31))]])).

cnf(refute_0_48, plain,
    (is_a_theorem(implies(implies(implies(implies(not(X_30), X_30), X_31),
           Z), implies(implies(X_30, X_31), Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_30, X_31),
                        implies(implies(not(X_30), X_30), X_31))))],
              [refute_0_47, refute_0_46])).

cnf(refute_0_49, plain,
    (is_a_theorem(implies(implies(implies(implies(not(X_62), X_62), X_63),
           X_61), implies(implies(X_62, X_63), X_61)))),
    inference(subst, [],
              [refute_0_48 :
               [bind(Z, $fot(X_61)), bind(X_30, $fot(X_62)),
                bind(X_31, $fot(X_63))]])).

cnf(refute_0_50, plain,
    (~
       is_a_theorem(implies(implies(implies(not(X_62), X_62), X_63),
           X_61)) | is_a_theorem(implies(implies(X_62, X_63), X_61))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(implies(not(X_62),
                              X_62), X_63), X_61),
                        implies(implies(X_62, X_63), X_61))))],
              [refute_0_49, refute_0_45])).

cnf(refute_0_51, plain,
    (~
       is_a_theorem(implies(implies(implies(not(X_697), X_697), X_698),
           implies(implies(implies(not(not(X_697)), X_32), X_697),
             X_698))) |
     is_a_theorem(implies(implies(X_697, X_698),
         implies(implies(implies(not(not(X_697)), X_32), X_697), X_698)))),
    inference(subst, [],
              [refute_0_50 :
               [bind(X_61,
                     $fot(implies(implies(implies(not(not(X_697)), X_32),
                              X_697), X_698))), bind(X_62, $fot(X_697)),
                bind(X_63, $fot(X_698))]])).

cnf(refute_0_52, plain,
    (is_a_theorem(implies(implies(X_697, X_698),
         implies(implies(implies(not(not(X_697)), X_32), X_697), X_698)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(not(X_697),
                            X_697), X_698),
                        implies(implies(implies(not(not(X_697)), X_32),
                            X_697), X_698))))],
              [refute_0_44, refute_0_51])).

cnf(refute_0_53, plain,
    (is_a_theorem(implies(implies(X_1044, X_1044),
         implies(implies(implies(not(not(X_1044)), X_32), X_1044),
           X_1044)))),
    inference(subst, [],
              [refute_0_52 :
               [bind(X_697, $fot(X_1044)), bind(X_698, $fot(X_1044))]])).

cnf(refute_0_54, plain,
    (~ is_a_theorem(implies(implies(X_702, X_702), X_704)) |
     ~
       is_a_theorem(implies(implies(implies(X_702, X_702), X_704),
           implies(X_703, X_704))) | is_a_theorem(implies(X_703, X_704))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(implies(implies(X_702, X_702), X_704))),
                bind(Y, $fot(implies(X_703, X_704)))]])).

cnf(refute_0_55, plain,
    (~
       is_a_theorem(implies(X_271,
           implies(not(implies(X_270, X_270)), X_269))) |
     is_a_theorem(implies(implies(implies(not(implies(X_270, X_270)),
             X_269), Z), implies(X_271, Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(X_271)),
                bind(Y,
                     $fot(implies(not(implies(X_270, X_270)), X_269)))]])).

cnf(refute_0_56, plain,
    (~
       is_a_theorem(implies(implies(not(X_33), X_266),
           implies(not(implies(X_268, X_268)), X_266))) |
     is_a_theorem(implies(X_33,
         implies(not(implies(X_268, X_268)), X_266)))),
    inference(subst, [],
              [refute_0_14 :
               [bind(X_32, $fot(X_266)),
                bind(X_34,
                     $fot(implies(not(implies(X_268, X_268)), X_266)))]])).

cnf(refute_0_57, plain,
    (~ is_a_theorem(implies(not(implies(X_223, X_223)), X_222)) |
     is_a_theorem(implies(implies(X_222, Z),
         implies(not(implies(X_223, X_223)), Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(not(implies(X_223, X_223)))),
                bind(Y, $fot(X_222))]])).

cnf(refute_0_58, plain, (is_a_theorem(implies(X_6, implies(not(X_6), Y)))),
    inference(subst, [], [cn_3 : [bind(X, $fot(X_6))]])).

cnf(refute_0_59, plain,
    (~ is_a_theorem(X_6) |
     ~ is_a_theorem(implies(X_6, implies(not(X_6), Y))) |
     is_a_theorem(implies(not(X_6), Y))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(X_6)),
                bind(Y, $fot(implies(not(X_6), Y)))]])).

cnf(refute_0_60, plain,
    (~ is_a_theorem(X_6) | is_a_theorem(implies(not(X_6), Y))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(X_6, implies(not(X_6), Y))))],
              [refute_0_58, refute_0_59])).

cnf(refute_0_61, plain,
    (~ is_a_theorem(implies(X_221, X_221)) |
     is_a_theorem(implies(not(implies(X_221, X_221)), Y))),
    inference(subst, [],
              [refute_0_60 : [bind(X_6, $fot(implies(X_221, X_221)))]])).

cnf(refute_0_62, plain,
    (is_a_theorem(implies(implies(not(X_220), X_220), X_220))),
    inference(subst, [], [cn_2 : [bind(X, $fot(X_220))]])).

cnf(refute_0_63, plain,
    (~ is_a_theorem(implies(implies(not(X_220), X_220), X_220)) |
     is_a_theorem(implies(X_220, X_220))),
    inference(subst, [],
              [refute_0_14 :
               [bind(X_32, $fot(X_220)), bind(X_33, $fot(X_220)),
                bind(X_34, $fot(X_220))]])).

cnf(refute_0_64, plain, (is_a_theorem(implies(X_220, X_220))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(not(X_220), X_220),
                        X_220)))], [refute_0_62, refute_0_63])).

cnf(refute_0_65, plain, (is_a_theorem(implies(X_221, X_221))),
    inference(subst, [], [refute_0_64 : [bind(X_220, $fot(X_221))]])).

cnf(refute_0_66, plain,
    (is_a_theorem(implies(not(implies(X_221, X_221)), Y))),
    inference(resolve, [$cnf(is_a_theorem(implies(X_221, X_221)))],
              [refute_0_65, refute_0_61])).

cnf(refute_0_67, plain,
    (is_a_theorem(implies(not(implies(X_223, X_223)), X_222))),
    inference(subst, [],
              [refute_0_66 :
               [bind(Y, $fot(X_222)), bind(X_221, $fot(X_223))]])).

cnf(refute_0_68, plain,
    (is_a_theorem(implies(implies(X_222, Z),
         implies(not(implies(X_223, X_223)), Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(not(implies(X_223, X_223)),
                        X_222)))], [refute_0_67, refute_0_57])).

cnf(refute_0_69, plain,
    (is_a_theorem(implies(implies(not(X_33), X_266),
         implies(not(implies(X_268, X_268)), X_266)))),
    inference(subst, [],
              [refute_0_68 :
               [bind(Z, $fot(X_266)), bind(X_222, $fot(not(X_33))),
                bind(X_223, $fot(X_268))]])).

cnf(refute_0_70, plain,
    (is_a_theorem(implies(X_33,
         implies(not(implies(X_268, X_268)), X_266)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(not(X_33), X_266),
                        implies(not(implies(X_268, X_268)), X_266))))],
              [refute_0_69, refute_0_56])).

cnf(refute_0_71, plain,
    (is_a_theorem(implies(X_271,
         implies(not(implies(X_270, X_270)), X_269)))),
    inference(subst, [],
              [refute_0_70 :
               [bind(X_266, $fot(X_269)), bind(X_268, $fot(X_270)),
                bind(X_33, $fot(X_271))]])).

cnf(refute_0_72, plain,
    (is_a_theorem(implies(implies(implies(not(implies(X_270, X_270)),
             X_269), Z), implies(X_271, Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(X_271,
                        implies(not(implies(X_270, X_270)), X_269))))],
              [refute_0_71, refute_0_55])).

cnf(refute_0_73, plain,
    (is_a_theorem(implies(implies(implies(not(implies(X_270, X_270)),
             implies(X_270, X_270)), X_698), implies(X_271, X_698)))),
    inference(subst, [],
              [refute_0_72 :
               [bind(Z, $fot(X_698)),
                bind(X_269, $fot(implies(X_270, X_270)))]])).

cnf(refute_0_74, plain,
    (~
       is_a_theorem(implies(implies(implies(not(implies(X_270, X_270)),
               implies(X_270, X_270)), X_698), implies(X_271, X_698))) |
     is_a_theorem(implies(implies(implies(X_270, X_270), X_698),
         implies(X_271, X_698)))),
    inference(subst, [],
              [refute_0_50 :
               [bind(X_61, $fot(implies(X_271, X_698))),
                bind(X_62, $fot(implies(X_270, X_270))),
                bind(X_63, $fot(X_698))]])).

cnf(refute_0_75, plain,
    (is_a_theorem(implies(implies(implies(X_270, X_270), X_698),
         implies(X_271, X_698)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(not(implies(X_270,
                                X_270)), implies(X_270, X_270)), X_698),
                        implies(X_271, X_698))))],
              [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain,
    (is_a_theorem(implies(implies(implies(X_702, X_702), X_704),
         implies(X_703, X_704)))),
    inference(subst, [],
              [refute_0_75 :
               [bind(X_270, $fot(X_702)), bind(X_271, $fot(X_703)),
                bind(X_698, $fot(X_704))]])).

cnf(refute_0_77, plain,
    (~ is_a_theorem(implies(implies(X_702, X_702), X_704)) |
     is_a_theorem(implies(X_703, X_704))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(X_702, X_702),
                          X_704), implies(X_703, X_704))))],
              [refute_0_76, refute_0_54])).

cnf(refute_0_78, plain,
    (~
       is_a_theorem(implies(implies(X_1044, X_1044),
           implies(implies(implies(not(not(X_1044)), X_32), X_1044),
             X_1044))) |
     is_a_theorem(implies(X_1045,
         implies(implies(implies(not(not(X_1044)), X_32), X_1044),
           X_1044)))),
    inference(subst, [],
              [refute_0_77 :
               [bind(X_702, $fot(X_1044)), bind(X_703, $fot(X_1045)),
                bind(X_704,
                     $fot(implies(implies(implies(not(not(X_1044)), X_32),
                              X_1044), X_1044)))]])).

cnf(refute_0_79, plain,
    (is_a_theorem(implies(X_1045,
         implies(implies(implies(not(not(X_1044)), X_32), X_1044),
           X_1044)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_1044, X_1044),
                        implies(implies(implies(not(not(X_1044)), X_32),
                            X_1044), X_1044))))],
              [refute_0_53, refute_0_78])).

cnf(refute_0_80, plain,
    (is_a_theorem(implies(not(implies(implies(implies(not(not(X_1164)),
                 X_1166), X_1164), X_1164)),
         implies(implies(implies(not(not(X_1164)), X_1166), X_1164),
           X_1164)))),
    inference(subst, [],
              [refute_0_79 :
               [bind(X_1044, $fot(X_1164)),
                bind(X_1045,
                     $fot(not(implies(implies(implies(not(not(X_1164)),
                                  X_1166), X_1164), X_1164)))),
                bind(X_32, $fot(X_1166))]])).

cnf(refute_0_81, plain,
    (is_a_theorem(implies(implies(implies(not(not(X_1164)), X_1166),
           X_1164), X_1164))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(not(implies(implies(implies(not(not(X_1164)),
                                X_1166), X_1164), X_1164)),
                        implies(implies(implies(not(not(X_1164)), X_1166),
                            X_1164), X_1164))))],
              [refute_0_80, refute_0_41])).

cnf(refute_0_82, plain,
    (is_a_theorem(implies(implies(implies(not(not(X_1167)), X_1168),
           X_1167), X_1167))),
    inference(subst, [],
              [refute_0_81 :
               [bind(X_1164, $fot(X_1167)), bind(X_1166, $fot(X_1168))]])).

cnf(refute_0_83, plain,
    (~ is_a_theorem(implies(implies(not(not(X_1167)), X_1168), X_1167)) |
     is_a_theorem(X_1167)),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(not(not(X_1167)),
                            X_1168), X_1167), X_1167)))],
              [refute_0_82, refute_0_37])).

cnf(refute_0_84, plain,
    (~
       is_a_theorem(implies(implies(not(not(implies(not(not(X_7354)),
                   X_7354))), X_1168),
           implies(not(not(X_7354)), X_7354))) |
     is_a_theorem(implies(not(not(X_7354)), X_7354))),
    inference(subst, [],
              [refute_0_83 :
               [bind(X_1167, $fot(implies(not(not(X_7354)), X_7354)))]])).

cnf(refute_0_85, plain,
    (~
       is_a_theorem(implies(implies(not(X_7033), X_7033),
           implies(implies(X_7033, X_7034), X_7034))) |
     is_a_theorem(implies(X_7033,
         implies(implies(X_7033, X_7034), X_7034)))),
    inference(subst, [],
              [refute_0_14 :
               [bind(X_32, $fot(X_7033)), bind(X_33, $fot(X_7033)),
                bind(X_34,
                     $fot(implies(implies(X_7033, X_7034), X_7034)))]])).

cnf(refute_0_86, plain,
    (~ is_a_theorem(implies(implies(X_1111, X_1111), X_1110)) |
     ~
       is_a_theorem(implies(implies(implies(X_1111, X_1111), X_1110),
           X_1110)) | is_a_theorem(X_1110)),
    inference(subst, [],
              [condensed_detachment :
               [bind(X, $fot(implies(implies(X_1111, X_1111), X_1110))),
                bind(Y, $fot(X_1110))]])).

cnf(refute_0_87, plain,
    (~
       is_a_theorem(implies(not(implies(implies(implies(X_1109, X_1109),
                 X_1107), X_1107)),
           implies(implies(implies(X_1109, X_1109), X_1107), X_1107))) |
     is_a_theorem(implies(implies(implies(X_1109, X_1109), X_1107),
         X_1107))),
    inference(subst, [],
              [refute_0_40 :
               [bind(X_7,
                     $fot(implies(implies(implies(X_1109, X_1109), X_1107),
                            X_1107)))]])).

cnf(refute_0_88, plain,
    (~
       is_a_theorem(implies(implies(implies(not(X_715), X_715), X_712),
           implies(implies(implies(X_713, X_713), X_715), X_712))) |
     is_a_theorem(implies(implies(X_715, X_712),
         implies(implies(implies(X_713, X_713), X_715), X_712)))),
    inference(subst, [],
              [refute_0_50 :
               [bind(X_61,
                     $fot(implies(implies(implies(X_713, X_713), X_715),
                            X_712))), bind(X_62, $fot(X_715)),
                bind(X_63, $fot(X_712))]])).

cnf(refute_0_89, plain,
    (~
       is_a_theorem(implies(implies(implies(X_702, X_702), X_704),
           implies(X_703, X_704))) |
     is_a_theorem(implies(implies(implies(X_703, X_704), Z),
         implies(implies(implies(X_702, X_702), X_704), Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(implies(implies(X_702, X_702), X_704))),
                bind(Y, $fot(implies(X_703, X_704)))]])).

cnf(refute_0_90, plain,
    (is_a_theorem(implies(implies(implies(X_703, X_704), Z),
         implies(implies(implies(X_702, X_702), X_704), Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(X_702, X_702),
                          X_704), implies(X_703, X_704))))],
              [refute_0_76, refute_0_89])).

cnf(refute_0_91, plain,
    (is_a_theorem(implies(implies(implies(not(X_715), X_715), X_712),
         implies(implies(implies(X_713, X_713), X_715), X_712)))),
    inference(subst, [],
              [refute_0_90 :
               [bind(Z, $fot(X_712)), bind(X_702, $fot(X_713)),
                bind(X_703, $fot(not(X_715))),
                bind(X_704, $fot(X_715))]])).

cnf(refute_0_92, plain,
    (is_a_theorem(implies(implies(X_715, X_712),
         implies(implies(implies(X_713, X_713), X_715), X_712)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(not(X_715),
                            X_715), X_712),
                        implies(implies(implies(X_713, X_713), X_715),
                          X_712))))], [refute_0_91, refute_0_88])).

cnf(refute_0_93, plain,
    (is_a_theorem(implies(implies(X_1044, X_1044),
         implies(implies(implies(X_713, X_713), X_1044), X_1044)))),
    inference(subst, [],
              [refute_0_92 :
               [bind(X_712, $fot(X_1044)), bind(X_715, $fot(X_1044))]])).

cnf(refute_0_94, plain,
    (~
       is_a_theorem(implies(implies(X_1044, X_1044),
           implies(implies(implies(X_713, X_713), X_1044), X_1044))) |
     is_a_theorem(implies(X_1045,
         implies(implies(implies(X_713, X_713), X_1044), X_1044)))),
    inference(subst, [],
              [refute_0_77 :
               [bind(X_702, $fot(X_1044)), bind(X_703, $fot(X_1045)),
                bind(X_704,
                     $fot(implies(implies(implies(X_713, X_713), X_1044),
                            X_1044)))]])).

cnf(refute_0_95, plain,
    (is_a_theorem(implies(X_1045,
         implies(implies(implies(X_713, X_713), X_1044), X_1044)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_1044, X_1044),
                        implies(implies(implies(X_713, X_713), X_1044),
                          X_1044))))], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    (is_a_theorem(implies(not(implies(implies(implies(X_1109, X_1109),
               X_1107), X_1107)),
         implies(implies(implies(X_1109, X_1109), X_1107), X_1107)))),
    inference(subst, [],
              [refute_0_95 :
               [bind(X_1044, $fot(X_1107)),
                bind(X_1045,
                     $fot(not(implies(implies(implies(X_1109, X_1109),
                                X_1107), X_1107)))),
                bind(X_713, $fot(X_1109))]])).

cnf(refute_0_97, plain,
    (is_a_theorem(implies(implies(implies(X_1109, X_1109), X_1107),
         X_1107))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(not(implies(implies(implies(X_1109,
                                X_1109), X_1107), X_1107)),
                        implies(implies(implies(X_1109, X_1109), X_1107),
                          X_1107))))], [refute_0_96, refute_0_87])).

cnf(refute_0_98, plain,
    (is_a_theorem(implies(implies(implies(X_1111, X_1111), X_1110),
         X_1110))),
    inference(subst, [],
              [refute_0_97 :
               [bind(X_1107, $fot(X_1110)), bind(X_1109, $fot(X_1111))]])).

cnf(refute_0_99, plain,
    (~ is_a_theorem(implies(implies(X_1111, X_1111), X_1110)) |
     is_a_theorem(X_1110)),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(X_1111, X_1111),
                          X_1110), X_1110)))],
              [refute_0_98, refute_0_86])).

cnf(refute_0_100, plain,
    (~
       is_a_theorem(implies(implies(X_7030, X_7030),
           implies(implies(not(X_7030), X_7030),
             implies(implies(X_7030, X_7032), X_7032)))) |
     is_a_theorem(implies(implies(not(X_7030), X_7030),
         implies(implies(X_7030, X_7032), X_7032)))),
    inference(subst, [],
              [refute_0_99 :
               [bind(X_1110,
                     $fot(implies(implies(not(X_7030), X_7030),
                            implies(implies(X_7030, X_7032), X_7032)))),
                bind(X_1111, $fot(X_7030))]])).

cnf(refute_0_101, plain,
    (~
       is_a_theorem(implies(implies(implies(not(X_6950), X_6950), X_6949),
           implies(implies(not(X_6950), X_6950),
             implies(implies(X_6949, X_6951), X_6951)))) |
     is_a_theorem(implies(implies(X_6950, X_6949),
         implies(implies(not(X_6950), X_6950),
           implies(implies(X_6949, X_6951), X_6951))))),
    inference(subst, [],
              [refute_0_50 :
               [bind(X_61,
                     $fot(implies(implies(not(X_6950), X_6950),
                            implies(implies(X_6949, X_6951), X_6951)))),
                bind(X_62, $fot(X_6950)), bind(X_63, $fot(X_6949))]])).

cnf(refute_0_102, plain,
    (~
       is_a_theorem(implies(implies(implies(X_39, X_6812),
             implies(implies(X_6813, X_6811), X_6812)),
           implies(implies(X_6813, X_6811),
             implies(implies(X_39, X_6812), X_6812)))) |
     is_a_theorem(implies(implies(implies(X_6813, X_6811), X_39),
         implies(implies(X_6813, X_6811),
           implies(implies(X_39, X_6812), X_6812))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_38, $fot(implies(X_6813, X_6811))),
                bind(X_40, $fot(X_6812)),
                bind(X_41,
                     $fot(implies(implies(X_6813, X_6811),
                            implies(implies(X_39, X_6812), X_6812))))]])).

cnf(refute_0_103, plain,
    (~
       is_a_theorem(implies(implies(implies(implies(implies(X_6758,
                   X_6756), X_6757), X_6757), implies(X_38, X_6757)),
           implies(implies(X_6758, X_6756), implies(X_38, X_6757)))) |
     is_a_theorem(implies(implies(X_38,
           implies(implies(X_6758, X_6756), X_6757)),
         implies(implies(X_6758, X_6756), implies(X_38, X_6757))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_39, $fot(implies(implies(X_6758, X_6756), X_6757))),
                bind(X_40, $fot(X_6757)),
                bind(X_41,
                     $fot(implies(implies(X_6758, X_6756),
                            implies(X_38, X_6757))))]])).

cnf(refute_0_104, plain,
    (~
       is_a_theorem(implies(implies(implies(X_2198, X_2198),
             implies(X_38, X_2198)),
           implies(X_2197,
             implies(implies(implies(X_38, X_2198), X_2199), X_2199)))) |
     is_a_theorem(implies(implies(X_38, X_2198),
         implies(X_2197,
           implies(implies(implies(X_38, X_2198), X_2199), X_2199))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_39, $fot(X_2198)), bind(X_40, $fot(X_2198)),
                bind(X_41,
                     $fot(implies(X_2197,
                            implies(implies(implies(X_38, X_2198), X_2199),
                              X_2199))))]])).

cnf(refute_0_105, plain,
    (~
       is_a_theorem(implies(implies(implies(X_39, X_2196),
             implies(implies(X_2194, X_2194), X_2196)),
           implies(X_2193, implies(implies(X_39, X_2196), X_2196)))) |
     is_a_theorem(implies(implies(implies(X_2194, X_2194), X_39),
         implies(X_2193, implies(implies(X_39, X_2196), X_2196))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_38, $fot(implies(X_2194, X_2194))),
                bind(X_40, $fot(X_2196)),
                bind(X_41,
                     $fot(implies(X_2193,
                            implies(implies(X_39, X_2196), X_2196))))]])).

cnf(refute_0_106, plain,
    (~
       is_a_theorem(implies(X_1108,
           implies(implies(implies(X_1109, X_1109), X_1107), X_1107))) |
     is_a_theorem(implies(implies(implies(implies(implies(X_1109, X_1109),
               X_1107), X_1107), Z), implies(X_1108, Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(X_1108)),
                bind(Y,
                     $fot(implies(implies(implies(X_1109, X_1109), X_1107),
                            X_1107)))]])).

cnf(refute_0_107, plain,
    (is_a_theorem(implies(X_1108,
         implies(implies(implies(X_1109, X_1109), X_1107), X_1107)))),
    inference(subst, [],
              [refute_0_95 :
               [bind(X_1044, $fot(X_1107)), bind(X_1045, $fot(X_1108)),
                bind(X_713, $fot(X_1109))]])).

cnf(refute_0_108, plain,
    (is_a_theorem(implies(implies(implies(implies(implies(X_1109, X_1109),
               X_1107), X_1107), Z), implies(X_1108, Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(X_1108,
                        implies(implies(implies(X_1109, X_1109), X_1107),
                          X_1107))))], [refute_0_107, refute_0_106])).

cnf(refute_0_109, plain,
    (is_a_theorem(implies(implies(implies(implies(implies(X_1109, X_1109),
               X_1772), X_1772), implies(X_1770, X_1772)),
         implies(X_1108, implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_108 :
               [bind(Z, $fot(implies(X_1770, X_1772))),
                bind(X_1107, $fot(X_1772))]])).

cnf(refute_0_110, plain,
    (~
       is_a_theorem(implies(implies(implies(implies(implies(X_1109,
                   X_1109), X_1772), X_1772), implies(X_1770, X_1772)),
           implies(X_1108, implies(X_1770, X_1772)))) |
     is_a_theorem(implies(implies(X_1770,
           implies(implies(X_1109, X_1109), X_1772)),
         implies(X_1108, implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_38, $fot(X_1770)),
                bind(X_39, $fot(implies(implies(X_1109, X_1109), X_1772))),
                bind(X_40, $fot(X_1772)),
                bind(X_41,
                     $fot(implies(X_1108, implies(X_1770, X_1772))))]])).

cnf(refute_0_111, plain,
    (is_a_theorem(implies(implies(X_1770,
           implies(implies(X_1109, X_1109), X_1772)),
         implies(X_1108, implies(X_1770, X_1772))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(implies(implies(X_1109,
                                X_1109), X_1772), X_1772),
                          implies(X_1770, X_1772)),
                        implies(X_1108, implies(X_1770, X_1772)))))],
              [refute_0_109, refute_0_110])).

cnf(refute_0_112, plain,
    (is_a_theorem(implies(implies(implies(X_39, X_2196),
           implies(implies(X_2194, X_2194), X_2196)),
         implies(X_2193, implies(implies(X_39, X_2196), X_2196))))),
    inference(subst, [],
              [refute_0_111 :
               [bind(X_1108, $fot(X_2193)), bind(X_1109, $fot(X_2194)),
                bind(X_1770, $fot(implies(X_39, X_2196))),
                bind(X_1772, $fot(X_2196))]])).

cnf(refute_0_113, plain,
    (is_a_theorem(implies(implies(implies(X_2194, X_2194), X_39),
         implies(X_2193, implies(implies(X_39, X_2196), X_2196))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(X_39, X_2196),
                          implies(implies(X_2194, X_2194), X_2196)),
                        implies(X_2193,
                          implies(implies(X_39, X_2196), X_2196)))))],
              [refute_0_112, refute_0_105])).

cnf(refute_0_114, plain,
    (is_a_theorem(implies(implies(implies(X_2198, X_2198),
           implies(X_38, X_2198)),
         implies(X_2197,
           implies(implies(implies(X_38, X_2198), X_2199), X_2199))))),
    inference(subst, [],
              [refute_0_113 :
               [bind(X_2193, $fot(X_2197)), bind(X_2194, $fot(X_2198)),
                bind(X_2196, $fot(X_2199)),
                bind(X_39, $fot(implies(X_38, X_2198)))]])).

cnf(refute_0_115, plain,
    (is_a_theorem(implies(implies(X_38, X_2198),
         implies(X_2197,
           implies(implies(implies(X_38, X_2198), X_2199), X_2199))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(X_2198, X_2198),
                          implies(X_38, X_2198)),
                        implies(X_2197,
                          implies(implies(implies(X_38, X_2198), X_2199),
                            X_2199)))))], [refute_0_114, refute_0_104])).

cnf(refute_0_116, plain,
    (is_a_theorem(implies(implies(X_38, X_2198),
         implies(implies(X_6726, X_6726),
           implies(implies(implies(X_38, X_2198), X_2199), X_2199))))),
    inference(subst, [],
              [refute_0_115 :
               [bind(X_2197, $fot(implies(X_6726, X_6726)))]])).

cnf(refute_0_117, plain,
    (~
       is_a_theorem(implies(X_1812,
           implies(implies(X_1814, X_1814), X_1815))) |
     ~
       is_a_theorem(implies(implies(X_1812,
             implies(implies(X_1814, X_1814), X_1815)),
           implies(implies(X_1815, X_1813), implies(X_1812, X_1813)))) |
     is_a_theorem(implies(implies(X_1815, X_1813),
         implies(X_1812, X_1813)))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X,
                     $fot(implies(X_1812,
                            implies(implies(X_1814, X_1814), X_1815)))),
                bind(Y,
                     $fot(implies(implies(X_1815, X_1813),
                            implies(X_1812, X_1813))))]])).

cnf(refute_0_118, plain,
    (~
       is_a_theorem(implies(implies(X_718, X_716),
           implies(implies(implies(X_717, X_717), X_718), X_716))) |
     is_a_theorem(implies(implies(implies(implies(implies(X_717, X_717),
               X_718), X_716), Z), implies(implies(X_718, X_716), Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(implies(X_718, X_716))),
                bind(Y,
                     $fot(implies(implies(implies(X_717, X_717), X_718),
                            X_716)))]])).

cnf(refute_0_119, plain,
    (is_a_theorem(implies(implies(X_718, X_716),
         implies(implies(implies(X_717, X_717), X_718), X_716)))),
    inference(subst, [],
              [refute_0_92 :
               [bind(X_712, $fot(X_716)), bind(X_713, $fot(X_717)),
                bind(X_715, $fot(X_718))]])).

cnf(refute_0_120, plain,
    (is_a_theorem(implies(implies(implies(implies(implies(X_717, X_717),
               X_718), X_716), Z), implies(implies(X_718, X_716), Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_718, X_716),
                        implies(implies(implies(X_717, X_717), X_718),
                          X_716))))], [refute_0_119, refute_0_118])).

cnf(refute_0_121, plain,
    (is_a_theorem(implies(implies(implies(implies(implies(X_717, X_717),
               X_718), X_1772), implies(X_1770, X_1772)),
         implies(implies(X_718, X_1772), implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_120 :
               [bind(Z, $fot(implies(X_1770, X_1772))),
                bind(X_716, $fot(X_1772))]])).

cnf(refute_0_122, plain,
    (~
       is_a_theorem(implies(implies(implies(implies(implies(X_717, X_717),
                 X_718), X_1772), implies(X_1770, X_1772)),
           implies(implies(X_718, X_1772), implies(X_1770, X_1772)))) |
     is_a_theorem(implies(implies(X_1770,
           implies(implies(X_717, X_717), X_718)),
         implies(implies(X_718, X_1772), implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_38, $fot(X_1770)),
                bind(X_39, $fot(implies(implies(X_717, X_717), X_718))),
                bind(X_40, $fot(X_1772)),
                bind(X_41,
                     $fot(implies(implies(X_718, X_1772),
                            implies(X_1770, X_1772))))]])).

cnf(refute_0_123, plain,
    (is_a_theorem(implies(implies(X_1770,
           implies(implies(X_717, X_717), X_718)),
         implies(implies(X_718, X_1772), implies(X_1770, X_1772))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(implies(implies(X_717,
                                X_717), X_718), X_1772),
                          implies(X_1770, X_1772)),
                        implies(implies(X_718, X_1772),
                          implies(X_1770, X_1772)))))],
              [refute_0_121, refute_0_122])).

cnf(refute_0_124, plain,
    (is_a_theorem(implies(implies(X_1812,
           implies(implies(X_1814, X_1814), X_1815)),
         implies(implies(X_1815, X_1813), implies(X_1812, X_1813))))),
    inference(subst, [],
              [refute_0_123 :
               [bind(X_1770, $fot(X_1812)), bind(X_1772, $fot(X_1813)),
                bind(X_717, $fot(X_1814)), bind(X_718, $fot(X_1815))]])).

cnf(refute_0_125, plain,
    (~
       is_a_theorem(implies(X_1812,
           implies(implies(X_1814, X_1814), X_1815))) |
     is_a_theorem(implies(implies(X_1815, X_1813),
         implies(X_1812, X_1813)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_1812,
                          implies(implies(X_1814, X_1814), X_1815)),
                        implies(implies(X_1815, X_1813),
                          implies(X_1812, X_1813)))))],
              [refute_0_124, refute_0_117])).

cnf(refute_0_126, plain,
    (~
       is_a_theorem(implies(implies(X_38, X_2198),
           implies(implies(X_6726, X_6726),
             implies(implies(implies(X_38, X_2198), X_2199), X_2199)))) |
     is_a_theorem(implies(implies(implies(implies(implies(X_38, X_2198),
               X_2199), X_2199), X_6725),
         implies(implies(X_38, X_2198), X_6725)))),
    inference(subst, [],
              [refute_0_125 :
               [bind(X_1812, $fot(implies(X_38, X_2198))),
                bind(X_1813, $fot(X_6725)), bind(X_1814, $fot(X_6726)),
                bind(X_1815,
                     $fot(implies(implies(implies(X_38, X_2198), X_2199),
                            X_2199)))]])).

cnf(refute_0_127, plain,
    (is_a_theorem(implies(implies(implies(implies(implies(X_38, X_2198),
               X_2199), X_2199), X_6725),
         implies(implies(X_38, X_2198), X_6725)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_38, X_2198),
                        implies(implies(X_6726, X_6726),
                          implies(implies(implies(X_38, X_2198), X_2199),
                            X_2199)))))], [refute_0_116, refute_0_126])).

cnf(refute_0_128, plain,
    (is_a_theorem(implies(implies(implies(implies(implies(X_6758, X_6756),
               X_6757), X_6757), implies(X_38, X_6757)),
         implies(implies(X_6758, X_6756), implies(X_38, X_6757))))),
    inference(subst, [],
              [refute_0_127 :
               [bind(X_2198, $fot(X_6756)), bind(X_2199, $fot(X_6757)),
                bind(X_38, $fot(X_6758)),
                bind(X_6725, $fot(implies(X_38, X_6757)))]])).

cnf(refute_0_129, plain,
    (is_a_theorem(implies(implies(X_38,
           implies(implies(X_6758, X_6756), X_6757)),
         implies(implies(X_6758, X_6756), implies(X_38, X_6757))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(implies(implies(X_6758,
                                X_6756), X_6757), X_6757),
                          implies(X_38, X_6757)),
                        implies(implies(X_6758, X_6756),
                          implies(X_38, X_6757)))))],
              [refute_0_128, refute_0_103])).

cnf(refute_0_130, plain,
    (is_a_theorem(implies(implies(implies(X_39, X_6812),
           implies(implies(X_6813, X_6811), X_6812)),
         implies(implies(X_6813, X_6811),
           implies(implies(X_39, X_6812), X_6812))))),
    inference(subst, [],
              [refute_0_129 :
               [bind(X_38, $fot(implies(X_39, X_6812))),
                bind(X_6756, $fot(X_6811)), bind(X_6757, $fot(X_6812)),
                bind(X_6758, $fot(X_6813))]])).

cnf(refute_0_131, plain,
    (is_a_theorem(implies(implies(implies(X_6813, X_6811), X_39),
         implies(implies(X_6813, X_6811),
           implies(implies(X_39, X_6812), X_6812))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(X_39, X_6812),
                          implies(implies(X_6813, X_6811), X_6812)),
                        implies(implies(X_6813, X_6811),
                          implies(implies(X_39, X_6812), X_6812)))))],
              [refute_0_130, refute_0_102])).

cnf(refute_0_132, plain,
    (is_a_theorem(implies(implies(implies(not(X_6950), X_6950), X_6949),
         implies(implies(not(X_6950), X_6950),
           implies(implies(X_6949, X_6951), X_6951))))),
    inference(subst, [],
              [refute_0_131 :
               [bind(X_39, $fot(X_6949)), bind(X_6811, $fot(X_6950)),
                bind(X_6812, $fot(X_6951)),
                bind(X_6813, $fot(not(X_6950)))]])).

cnf(refute_0_133, plain,
    (is_a_theorem(implies(implies(X_6950, X_6949),
         implies(implies(not(X_6950), X_6950),
           implies(implies(X_6949, X_6951), X_6951))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(not(X_6950),
                            X_6950), X_6949),
                        implies(implies(not(X_6950), X_6950),
                          implies(implies(X_6949, X_6951), X_6951)))))],
              [refute_0_132, refute_0_101])).

cnf(refute_0_134, plain,
    (is_a_theorem(implies(implies(X_7030, X_7030),
         implies(implies(not(X_7030), X_7030),
           implies(implies(X_7030, X_7032), X_7032))))),
    inference(subst, [],
              [refute_0_133 :
               [bind(X_6949, $fot(X_7030)), bind(X_6950, $fot(X_7030)),
                bind(X_6951, $fot(X_7032))]])).

cnf(refute_0_135, plain,
    (is_a_theorem(implies(implies(not(X_7030), X_7030),
         implies(implies(X_7030, X_7032), X_7032)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_7030, X_7030),
                        implies(implies(not(X_7030), X_7030),
                          implies(implies(X_7030, X_7032), X_7032)))))],
              [refute_0_134, refute_0_100])).

cnf(refute_0_136, plain,
    (is_a_theorem(implies(implies(not(X_7033), X_7033),
         implies(implies(X_7033, X_7034), X_7034)))),
    inference(subst, [],
              [refute_0_135 :
               [bind(X_7030, $fot(X_7033)), bind(X_7032, $fot(X_7034))]])).

cnf(refute_0_137, plain,
    (is_a_theorem(implies(X_7033,
         implies(implies(X_7033, X_7034), X_7034)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(not(X_7033), X_7033),
                        implies(implies(X_7033, X_7034), X_7034))))],
              [refute_0_136, refute_0_85])).

cnf(refute_0_138, plain,
    (is_a_theorem(implies(not(not(X_7349)),
         implies(implies(not(not(X_7349)), X_7349), X_7349)))),
    inference(subst, [],
              [refute_0_137 :
               [bind(X_7033, $fot(not(not(X_7349)))),
                bind(X_7034, $fot(X_7349))]])).

cnf(refute_0_139, plain,
    (~
       is_a_theorem(implies(X_2804,
           implies(implies(not(not(X_2805)), X_2803), X_2805))) |
     ~
       is_a_theorem(implies(implies(X_2804,
             implies(implies(not(not(X_2805)), X_2803), X_2805)),
           implies(X_2802, implies(X_2804, X_2805)))) |
     is_a_theorem(implies(X_2802, implies(X_2804, X_2805)))),
    inference(subst, [],
              [condensed_detachment :
               [bind(X,
                     $fot(implies(X_2804,
                            implies(implies(not(not(X_2805)), X_2803),
                              X_2805)))),
                bind(Y,
                     $fot(implies(X_2802, implies(X_2804, X_2805))))]])).

cnf(refute_0_140, plain,
    (~
       is_a_theorem(implies(X_1165,
           implies(implies(implies(not(not(X_1164)), X_1166), X_1164),
             X_1164))) |
     is_a_theorem(implies(implies(implies(implies(implies(not(not(X_1164)),
                 X_1166), X_1164), X_1164), Z), implies(X_1165, Z)))),
    inference(subst, [],
              [refute_0_1 :
               [bind(X, $fot(X_1165)),
                bind(Y,
                     $fot(implies(implies(implies(not(not(X_1164)),
                                X_1166), X_1164), X_1164)))]])).

cnf(refute_0_141, plain,
    (is_a_theorem(implies(X_1165,
         implies(implies(implies(not(not(X_1164)), X_1166), X_1164),
           X_1164)))),
    inference(subst, [],
              [refute_0_79 :
               [bind(X_1044, $fot(X_1164)), bind(X_1045, $fot(X_1165)),
                bind(X_32, $fot(X_1166))]])).

cnf(refute_0_142, plain,
    (is_a_theorem(implies(implies(implies(implies(implies(not(not(X_1164)),
                 X_1166), X_1164), X_1164), Z), implies(X_1165, Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(X_1165,
                        implies(implies(implies(not(not(X_1164)), X_1166),
                            X_1164), X_1164))))],
              [refute_0_141, refute_0_140])).

cnf(refute_0_143, plain,
    (is_a_theorem(implies(implies(implies(implies(implies(not(not(X_1772)),
                 X_1166), X_1772), X_1772), implies(X_1770, X_1772)),
         implies(X_1165, implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_142 :
               [bind(Z, $fot(implies(X_1770, X_1772))),
                bind(X_1164, $fot(X_1772))]])).

cnf(refute_0_144, plain,
    (~
       is_a_theorem(implies(implies(implies(implies(implies(not(not(X_1772)),
                   X_1166), X_1772), X_1772), implies(X_1770, X_1772)),
           implies(X_1165, implies(X_1770, X_1772)))) |
     is_a_theorem(implies(implies(X_1770,
           implies(implies(not(not(X_1772)), X_1166), X_1772)),
         implies(X_1165, implies(X_1770, X_1772))))),
    inference(subst, [],
              [refute_0_24 :
               [bind(X_38, $fot(X_1770)),
                bind(X_39,
                     $fot(implies(implies(not(not(X_1772)), X_1166),
                            X_1772))), bind(X_40, $fot(X_1772)),
                bind(X_41,
                     $fot(implies(X_1165, implies(X_1770, X_1772))))]])).

cnf(refute_0_145, plain,
    (is_a_theorem(implies(implies(X_1770,
           implies(implies(not(not(X_1772)), X_1166), X_1772)),
         implies(X_1165, implies(X_1770, X_1772))))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(implies(implies(not(not(X_1772)),
                                X_1166), X_1772), X_1772),
                          implies(X_1770, X_1772)),
                        implies(X_1165, implies(X_1770, X_1772)))))],
              [refute_0_143, refute_0_144])).

cnf(refute_0_146, plain,
    (is_a_theorem(implies(implies(X_2804,
           implies(implies(not(not(X_2805)), X_2803), X_2805)),
         implies(X_2802, implies(X_2804, X_2805))))),
    inference(subst, [],
              [refute_0_145 :
               [bind(X_1165, $fot(X_2802)), bind(X_1166, $fot(X_2803)),
                bind(X_1770, $fot(X_2804)), bind(X_1772, $fot(X_2805))]])).

cnf(refute_0_147, plain,
    (~
       is_a_theorem(implies(X_2804,
           implies(implies(not(not(X_2805)), X_2803), X_2805))) |
     is_a_theorem(implies(X_2802, implies(X_2804, X_2805)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(X_2804,
                          implies(implies(not(not(X_2805)), X_2803),
                            X_2805)),
                        implies(X_2802, implies(X_2804, X_2805)))))],
              [refute_0_146, refute_0_139])).

cnf(refute_0_148, plain,
    (~
       is_a_theorem(implies(not(not(X_7349)),
           implies(implies(not(not(X_7349)), X_7349), X_7349))) |
     is_a_theorem(implies(X_7346, implies(not(not(X_7349)), X_7349)))),
    inference(subst, [],
              [refute_0_147 :
               [bind(X_2802, $fot(X_7346)), bind(X_2803, $fot(X_7349)),
                bind(X_2804, $fot(not(not(X_7349)))),
                bind(X_2805, $fot(X_7349))]])).

cnf(refute_0_149, plain,
    (is_a_theorem(implies(X_7346, implies(not(not(X_7349)), X_7349)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(not(not(X_7349)),
                        implies(implies(not(not(X_7349)), X_7349),
                          X_7349))))], [refute_0_138, refute_0_148])).

cnf(refute_0_150, plain,
    (is_a_theorem(implies(implies(not(not(implies(not(not(X_7354)),
                 X_7354))), X_1168), implies(not(not(X_7354)), X_7354)))),
    inference(subst, [],
              [refute_0_149 :
               [bind(X_7346,
                     $fot(implies(not(not(implies(not(not(X_7354)),
                                  X_7354))), X_1168))),
                bind(X_7349, $fot(X_7354))]])).

cnf(refute_0_151, plain, (is_a_theorem(implies(not(not(X_7354)), X_7354))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(not(not(implies(not(not(X_7354)),
                                X_7354))), X_1168),
                        implies(not(not(X_7354)), X_7354))))],
              [refute_0_150, refute_0_84])).

cnf(refute_0_152, plain,
    (is_a_theorem(implies(not(not(not(X_1778))), not(X_1778)))),
    inference(subst, [],
              [refute_0_151 : [bind(X_7354, $fot(not(X_1778)))]])).

cnf(refute_0_153, plain,
    (is_a_theorem(implies(X_1778,
         implies(not(not(not(X_1778))), X_1777)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(not(not(not(X_1778))),
                        not(X_1778))))], [refute_0_152, refute_0_36])).

cnf(refute_0_154, plain,
    (is_a_theorem(implies(X_7366,
         implies(not(not(not(X_7366))), not(not(X_7366)))))),
    inference(subst, [],
              [refute_0_153 :
               [bind(X_1777, $fot(not(not(X_7366)))),
                bind(X_1778, $fot(X_7366))]])).

cnf(refute_0_155, plain,
    (is_a_theorem(implies(not(X_7366), implies(X_7366, X_1941)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(X_7366,
                        implies(not(not(not(X_7366))),
                          not(not(X_7366))))))],
              [refute_0_154, refute_0_29])).

cnf(refute_0_156, plain,
    (is_a_theorem(implies(not(X_7369), implies(X_7369, X_7368)))),
    inference(subst, [],
              [refute_0_155 :
               [bind(X_1941, $fot(X_7368)), bind(X_7366, $fot(X_7369))]])).

cnf(refute_0_157, plain,
    (is_a_theorem(implies(implies(implies(X_7369, X_7368), Z),
         implies(not(X_7369), Z)))),
    inference(resolve,
              [$cnf(is_a_theorem(implies(not(X_7369),
                        implies(X_7369, X_7368))))],
              [refute_0_156, refute_0_2])).

cnf(refute_0_158, plain,
    (is_a_theorem(implies(implies(implies(a, b), c), implies(not(a), c)))),
    inference(subst, [],
              [refute_0_157 :
               [bind(Z, $fot(c)), bind(X_7368, $fot(b)),
                bind(X_7369, $fot(a))]])).

cnf(refute_0_159, plain, ($false),
    inference(resolve,
              [$cnf(is_a_theorem(implies(implies(implies(a, b), c),
                        implies(not(a), c))))],
              [refute_0_158, prove_cn_37])).
% SZS output end CNFRefutation for LCL055-1.tptp

