%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for PUZ011-1.tptp

% SZS output start CNFRefutation for PUZ011-1.tptp
cnf(indian, hypothesis, (ocean(indian))).

cnf(indian_india, hypothesis, (borders(indian, india))).

cnf(indian_kenya, hypothesis, (borders(indian, kenya))).

cnf(kenya, hypothesis, (african(kenya))).

cnf(india, hypothesis, (asian(india))).

cnf(prove_there_is_a_common_ocean, negated_conjecture,
    (~ ocean(Ocean) | ~ borders(Ocean, African) | ~ african(African) |
     ~ borders(Ocean, Asian) | ~ asian(Asian))).

cnf(refute_0_0, plain,
    (~ african(kenya) | ~ asian(X_3) | ~ borders(indian, X_3) |
     ~ borders(indian, kenya) | ~ ocean(indian)),
    inference(subst, [],
              [prove_there_is_a_common_ocean :
               [bind(African, $fot(kenya)), bind(Asian, $fot(X_3)),
                bind(Ocean, $fot(indian))]])).

cnf(refute_0_1, plain,
    (~ african(kenya) | ~ asian(X_3) | ~ borders(indian, X_3) |
     ~ ocean(indian)),
    inference(resolve, [$cnf(borders(indian, kenya))],
              [indian_kenya, refute_0_0])).

cnf(refute_0_2, plain,
    (~ asian(X_3) | ~ borders(indian, X_3) | ~ ocean(indian)),
    inference(resolve, [$cnf(african(kenya))], [kenya, refute_0_1])).

cnf(refute_0_3, plain, (~ asian(X_3) | ~ borders(indian, X_3)),
    inference(resolve, [$cnf(ocean(indian))], [indian, refute_0_2])).

cnf(refute_0_4, plain, (~ asian(india) | ~ borders(indian, india)),
    inference(subst, [], [refute_0_3 : [bind(X_3, $fot(india))]])).

cnf(refute_0_5, plain, (~ asian(india)),
    inference(resolve, [$cnf(borders(indian, india))],
              [indian_india, refute_0_4])).

cnf(refute_0_6, plain, ($false),
    inference(resolve, [$cnf(asian(india))], [india, refute_0_5])).
% SZS output end CNFRefutation for PUZ011-1.tptp

