%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for SET090-7.tptp

% SZS output start CNFRefutation for SET090-7.tptp
cnf(set_in_its_singleton, axiom,
    (~ member(X, universal_class) | member(X, singleton(X)))).

cnf(only_member_in_singleton, axiom, (~ member(Y, singleton(X)) | Y = X)).

cnf(member_exists1, axiom,
    (~ member(Y, universal_class) |
     member(member_of(singleton(Y)), universal_class))).

cnf(member_exists2, axiom,
    (~ member(Y, universal_class) |
     singleton(member_of(singleton(Y))) = singleton(Y))).

cnf(prove_member_of_singleton_is_unique_1, negated_conjecture,
    (member(u, universal_class))).

cnf(prove_member_of_singleton_is_unique_2, negated_conjecture,
    (member_of(singleton(u)) != u)).

cnf(refute_0_0, plain,
    (~ member(member_of(singleton(u)), universal_class) |
     member(member_of(singleton(u)), singleton(member_of(singleton(u))))),
    inference(subst, [],
              [set_in_its_singleton :
               [bind(X, $fot(member_of(singleton(u))))]])).

cnf(refute_0_1, plain,
    (~ member(u, universal_class) |
     member(member_of(singleton(u)), universal_class)),
    inference(subst, [], [member_exists1 : [bind(Y, $fot(u))]])).

cnf(refute_0_2, plain, (member(member_of(singleton(u)), universal_class)),
    inference(resolve, [$cnf(member(u, universal_class))],
              [prove_member_of_singleton_is_unique_1, refute_0_1])).

cnf(refute_0_3, plain,
    (member(member_of(singleton(u)), singleton(member_of(singleton(u))))),
    inference(resolve,
              [$cnf(member(member_of(singleton(u)), universal_class))],
              [refute_0_2, refute_0_0])).

cnf(refute_0_4, plain,
    (~ member(u, universal_class) |
     singleton(member_of(singleton(u))) = singleton(u)),
    inference(subst, [], [member_exists2 : [bind(Y, $fot(u))]])).

cnf(refute_0_5, plain, (singleton(member_of(singleton(u))) = singleton(u)),
    inference(resolve, [$cnf(member(u, universal_class))],
              [prove_member_of_singleton_is_unique_1, refute_0_4])).

cnf(refute_0_6, plain,
    (singleton(member_of(singleton(u))) != singleton(u) |
     ~
       member(member_of(singleton(u)),
         singleton(member_of(singleton(u)))) |
     member(member_of(singleton(u)), singleton(u))),
    introduced(tautology,
               [equality,
                [$cnf(member(member_of(singleton(u)),
                        singleton(member_of(singleton(u))))), [1],
                 $fot(singleton(u))]])).

cnf(refute_0_7, plain,
    (~
       member(member_of(singleton(u)),
         singleton(member_of(singleton(u)))) |
     member(member_of(singleton(u)), singleton(u))),
    inference(resolve,
              [$cnf($equal(singleton(member_of(singleton(u))),
                      singleton(u)))], [refute_0_5, refute_0_6])).

cnf(refute_0_8, plain, (member(member_of(singleton(u)), singleton(u))),
    inference(resolve,
              [$cnf(member(member_of(singleton(u)),
                      singleton(member_of(singleton(u)))))],
              [refute_0_3, refute_0_7])).

cnf(refute_0_9, plain,
    (~ member(member_of(singleton(u)), singleton(u)) |
     member_of(singleton(u)) = u),
    inference(subst, [],
              [only_member_in_singleton :
               [bind(X, $fot(u)),
                bind(Y, $fot(member_of(singleton(u))))]])).

cnf(refute_0_10, plain, (member_of(singleton(u)) = u),
    inference(resolve,
              [$cnf(member(member_of(singleton(u)), singleton(u)))],
              [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain, ($false),
    inference(resolve, [$cnf($equal(member_of(singleton(u)), u))],
              [refute_0_10, prove_member_of_singleton_is_unique_2])).
% SZS output end CNFRefutation for SET090-7.tptp

