%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for SYN555-1.tptp

% SZS output start CNFRefutation for SYN555-1.tptp
cnf(p4_2, negated_conjecture, (p4(X12, X12))).

cnf(p4_4, negated_conjecture, (p4(f5(c7, c10, c9), c9))).

cnf(not_p4_5, negated_conjecture, (~ p4(f5(c7, c10, c8), c8))).

cnf(p4_6, negated_conjecture, (p4(f5(c7, X24, c8), f5(c7, X24, c9)))).

cnf(p4_8, negated_conjecture,
    (p4(X13, X14) | ~ p4(X12, X13) | ~ p4(X12, X14))).

cnf(p4_12, negated_conjecture,
    (p4(f5(c7, X21, X22), X22) | ~ p4(f5(c7, X21, X23), X23) |
     ~
       p4(f5(c7, f6(X21, X23, X22), X23),
         f5(c7, f6(X21, X23, X22), X22)))).

cnf(refute_0_0, plain, (p4(X_16, X_16)),
    inference(subst, [], [p4_2 : [bind(X12, $fot(X_16))]])).

cnf(refute_0_1, plain,
    (~ p4(X_16, X_15) | ~ p4(X_16, X_16) | p4(X_15, X_16)),
    inference(subst, [],
              [p4_8 :
               [bind(X12, $fot(X_16)), bind(X13, $fot(X_15)),
                bind(X14, $fot(X_16))]])).

cnf(refute_0_2, plain, (~ p4(X_16, X_15) | p4(X_15, X_16)),
    inference(resolve, [$cnf(p4(X_16, X_16))], [refute_0_0, refute_0_1])).

cnf(refute_0_3, plain,
    (~ p4(f5(c7, X24, c8), f5(c7, X24, c9)) |
     p4(f5(c7, X24, c9), f5(c7, X24, c8))),
    inference(subst, [],
              [refute_0_2 :
               [bind(X_15, $fot(f5(c7, X24, c9))),
                bind(X_16, $fot(f5(c7, X24, c8)))]])).

cnf(refute_0_4, plain, (p4(f5(c7, X24, c9), f5(c7, X24, c8))),
    inference(resolve, [$cnf(p4(f5(c7, X24, c8), f5(c7, X24, c9)))],
              [p4_6, refute_0_3])).

cnf(refute_0_5, plain,
    (p4(f5(c7, f6(X_605, c9, c8), c9), f5(c7, f6(X_605, c9, c8), c8))),
    inference(subst, [],
              [refute_0_4 : [bind(X24, $fot(f6(X_605, c9, c8)))]])).

cnf(refute_0_6, plain,
    (~ p4(f5(c7, X_605, c9), c9) |
     ~ p4(f5(c7, f6(X_605, c9, c8), c9), f5(c7, f6(X_605, c9, c8), c8)) |
     p4(f5(c7, X_605, c8), c8)),
    inference(subst, [],
              [p4_12 :
               [bind(X21, $fot(X_605)), bind(X22, $fot(c8)),
                bind(X23, $fot(c9))]])).

cnf(refute_0_7, plain,
    (~ p4(f5(c7, X_605, c9), c9) | p4(f5(c7, X_605, c8), c8)),
    inference(resolve,
              [$cnf(p4(f5(c7, f6(X_605, c9, c8), c9),
                      f5(c7, f6(X_605, c9, c8), c8)))],
              [refute_0_5, refute_0_6])).

cnf(refute_0_8, plain,
    (~ p4(f5(c7, c10, c9), c9) | p4(f5(c7, c10, c8), c8)),
    inference(subst, [], [refute_0_7 : [bind(X_605, $fot(c10))]])).

cnf(refute_0_9, plain, (p4(f5(c7, c10, c8), c8)),
    inference(resolve, [$cnf(p4(f5(c7, c10, c9), c9))],
              [p4_4, refute_0_8])).

cnf(refute_0_10, plain, ($false),
    inference(resolve, [$cnf(p4(f5(c7, c10, c8), c8))],
              [refute_0_9, not_p4_5])).
% SZS output end CNFRefutation for SYN555-1.tptp

