%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for SET864-2.tptp

% SZS output start CNFRefutation for SET864-2.tptp
cnf(cls_Zorn_Omaxchain__Zorn_0, axiom,
    (~ c_in(V_u, V_S, tc_set(T_a)) |
     ~ c_in(V_c, c_Zorn_Omaxchain(V_S, T_a), tc_set(tc_set(T_a))) |
     ~ c_lessequals(c_Union(V_c, T_a), V_u, tc_set(T_a)) |
     c_Union(V_c, T_a) = V_u)).

cnf(cls_conjecture_0, negated_conjecture,
    (c_in(v_c, c_Zorn_Omaxchain(v_S, t_a), tc_set(tc_set(t_a))))).

cnf(cls_conjecture_2, negated_conjecture,
    (c_in(c_Union(v_c, t_a), v_S, tc_set(t_a)))).

cnf(cls_conjecture_3, negated_conjecture,
    (c_in(v_x(V_U), v_S, tc_set(t_a)) | ~ c_in(V_U, v_S, tc_set(t_a)))).

cnf(cls_conjecture_4, negated_conjecture,
    (c_lessequals(V_U, v_x(V_U), tc_set(t_a)) |
     ~ c_in(V_U, v_S, tc_set(t_a)))).

cnf(cls_conjecture_5, negated_conjecture,
    (V_U != v_x(V_U) | ~ c_in(V_U, v_S, tc_set(t_a)))).

cnf(refute_0_0, plain,
    (~ c_in(c_Union(v_c, t_a), v_S, tc_set(t_a)) |
     c_in(v_x(c_Union(v_c, t_a)), v_S, tc_set(t_a))),
    inference(subst, [],
              [cls_conjecture_3 : [bind(V_U, $fot(c_Union(v_c, t_a)))]])).

cnf(refute_0_1, plain, (c_in(v_x(c_Union(v_c, t_a)), v_S, tc_set(t_a))),
    inference(resolve, [$cnf(c_in(c_Union(v_c, t_a), v_S, tc_set(t_a)))],
              [cls_conjecture_2, refute_0_0])).

cnf(refute_0_2, plain,
    (~ c_in(c_Union(v_c, t_a), v_S, tc_set(t_a)) |
     c_lessequals(c_Union(v_c, t_a), v_x(c_Union(v_c, t_a)), tc_set(t_a))),
    inference(subst, [],
              [cls_conjecture_4 : [bind(V_U, $fot(c_Union(v_c, t_a)))]])).

cnf(refute_0_3, plain,
    (c_lessequals(c_Union(v_c, t_a), v_x(c_Union(v_c, t_a)), tc_set(t_a))),
    inference(resolve, [$cnf(c_in(c_Union(v_c, t_a), v_S, tc_set(t_a)))],
              [cls_conjecture_2, refute_0_2])).

cnf(refute_0_4, plain,
    (~ c_in(v_c, c_Zorn_Omaxchain(X_4, t_a), tc_set(tc_set(t_a))) |
     ~ c_in(v_x(c_Union(v_c, t_a)), X_4, tc_set(t_a)) |
     ~
       c_lessequals(c_Union(v_c, t_a), v_x(c_Union(v_c, t_a)),
         tc_set(t_a)) | c_Union(v_c, t_a) = v_x(c_Union(v_c, t_a))),
    inference(subst, [],
              [cls_Zorn_Omaxchain__Zorn_0 :
               [bind(T_a, $fot(t_a)), bind(V_S, $fot(X_4)),
                bind(V_c, $fot(v_c)),
                bind(V_u, $fot(v_x(c_Union(v_c, t_a))))]])).

cnf(refute_0_5, plain,
    (~ c_in(v_c, c_Zorn_Omaxchain(X_4, t_a), tc_set(tc_set(t_a))) |
     ~ c_in(v_x(c_Union(v_c, t_a)), X_4, tc_set(t_a)) |
     c_Union(v_c, t_a) = v_x(c_Union(v_c, t_a))),
    inference(resolve,
              [$cnf(c_lessequals(c_Union(v_c, t_a), v_x(c_Union(v_c, t_a)),
                      tc_set(t_a)))], [refute_0_3, refute_0_4])).

cnf(refute_0_6, plain,
    (c_Union(v_c, t_a) != v_x(c_Union(v_c, t_a)) |
     ~ c_in(c_Union(v_c, t_a), v_S, tc_set(t_a))),
    inference(subst, [],
              [cls_conjecture_5 : [bind(V_U, $fot(c_Union(v_c, t_a)))]])).

cnf(refute_0_7, plain, (c_Union(v_c, t_a) != v_x(c_Union(v_c, t_a))),
    inference(resolve, [$cnf(c_in(c_Union(v_c, t_a), v_S, tc_set(t_a)))],
              [cls_conjecture_2, refute_0_6])).

cnf(refute_0_8, plain,
    (~ c_in(v_c, c_Zorn_Omaxchain(X_4, t_a), tc_set(tc_set(t_a))) |
     ~ c_in(v_x(c_Union(v_c, t_a)), X_4, tc_set(t_a))),
    inference(resolve,
              [$cnf($equal(c_Union(v_c, t_a), v_x(c_Union(v_c, t_a))))],
              [refute_0_5, refute_0_7])).

cnf(refute_0_9, plain,
    (~ c_in(v_c, c_Zorn_Omaxchain(v_S, t_a), tc_set(tc_set(t_a))) |
     ~ c_in(v_x(c_Union(v_c, t_a)), v_S, tc_set(t_a))),
    inference(subst, [], [refute_0_8 : [bind(X_4, $fot(v_S))]])).

cnf(refute_0_10, plain,
    (~ c_in(v_c, c_Zorn_Omaxchain(v_S, t_a), tc_set(tc_set(t_a)))),
    inference(resolve,
              [$cnf(c_in(v_x(c_Union(v_c, t_a)), v_S, tc_set(t_a)))],
              [refute_0_1, refute_0_9])).

cnf(refute_0_11, plain, ($false),
    inference(resolve,
              [$cnf(c_in(v_c, c_Zorn_Omaxchain(v_S, t_a),
                      tc_set(tc_set(t_a))))],
              [cls_conjecture_0, refute_0_10])).
% SZS output end CNFRefutation for SET864-2.tptp

