%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for SYN116-1.tptp

% SZS output start CNFRefutation for SYN116-1.tptp
cnf(axiom_1, axiom, (s0(d))).

cnf(axiom_14, axiom, (p0(b, X))).

cnf(axiom_34, axiom, (n0(c, d))).

cnf(rule_002, axiom, (l1(G, G) | ~ n0(H, G))).

cnf(rule_072, axiom, (p1(F, F, F) | ~ s0(G) | ~ s0(F))).

cnf(rule_085, axiom, (p1(B, B, B) | ~ p0(C, B))).

cnf(rule_125, axiom, (s1(I) | ~ p0(I, I))).

cnf(rule_137, axiom, (n2(A) | ~ p1(B, C, A))).

cnf(rule_181, axiom, (q2(I, I, I) | ~ p1(I, I, I))).

cnf(rule_189, axiom, (s2(H) | ~ q2(b, H, b) | ~ s1(b))).

cnf(rule_244, axiom, (p3(H, H, H) | ~ n2(H))).

cnf(rule_264, axiom, (r3(H, H, H) | ~ s2(H))).

cnf(rule_299, axiom, (s4(A) | ~ p3(B, C, D) | ~ l1(A, C))).

cnf(rule_300, axiom, (k5(E) | ~ s4(F) | ~ r3(G, E, E))).

cnf(prove_this, negated_conjecture, (~ k5(b))).

cnf(refute_0_0, plain, (~ n0(c, d) | l1(d, d)),
    inference(subst, [],
              [rule_002 : [bind(G, $fot(d)), bind(H, $fot(c))]])).

cnf(refute_0_1, plain, (l1(d, d)),
    inference(resolve, [$cnf(n0(c, d))], [axiom_34, refute_0_0])).

cnf(refute_0_2, plain, (~ n2(d) | p3(d, d, d)),
    inference(subst, [], [rule_244 : [bind(H, $fot(d))]])).

cnf(refute_0_3, plain, (~ s0(d) | p1(d, d, d)),
    inference(subst, [],
              [rule_072 : [bind(F, $fot(d)), bind(G, $fot(d))]])).

cnf(refute_0_4, plain, (p1(d, d, d)),
    inference(resolve, [$cnf(s0(d))], [axiom_1, refute_0_3])).

cnf(refute_0_5, plain, (~ p1(d, d, d) | n2(d)),
    inference(subst, [],
              [rule_137 :
               [bind(A, $fot(d)), bind(B, $fot(d)), bind(C, $fot(d))]])).

cnf(refute_0_6, plain, (n2(d)),
    inference(resolve, [$cnf(p1(d, d, d))], [refute_0_4, refute_0_5])).

cnf(refute_0_7, plain, (p3(d, d, d)),
    inference(resolve, [$cnf(n2(d))], [refute_0_6, refute_0_2])).

cnf(refute_0_8, plain, (~ s2(b) | r3(b, b, b)),
    inference(subst, [], [rule_264 : [bind(H, $fot(b))]])).

cnf(refute_0_9, plain, (~ p1(X_66, X_66, X_66) | q2(X_66, X_66, X_66)),
    inference(subst, [], [rule_181 : [bind(I, $fot(X_66))]])).

cnf(refute_0_10, plain, (p0(b, X_64)),
    inference(subst, [], [axiom_14 : [bind(X, $fot(X_64))]])).

cnf(refute_0_11, plain, (~ p0(b, X_64) | p1(X_64, X_64, X_64)),
    inference(subst, [],
              [rule_085 : [bind(B, $fot(X_64)), bind(C, $fot(b))]])).

cnf(refute_0_12, plain, (p1(X_64, X_64, X_64)),
    inference(resolve, [$cnf(p0(b, X_64))], [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain, (p1(X_66, X_66, X_66)),
    inference(subst, [], [refute_0_12 : [bind(X_64, $fot(X_66))]])).

cnf(refute_0_14, plain, (q2(X_66, X_66, X_66)),
    inference(resolve, [$cnf(p1(X_66, X_66, X_66))],
              [refute_0_13, refute_0_9])).

cnf(refute_0_15, plain, (q2(b, b, b)),
    inference(subst, [], [refute_0_14 : [bind(X_66, $fot(b))]])).

cnf(refute_0_16, plain, (p0(b, b)),
    inference(subst, [], [axiom_14 : [bind(X, $fot(b))]])).

cnf(refute_0_17, plain, (~ p0(b, b) | s1(b)),
    inference(subst, [], [rule_125 : [bind(I, $fot(b))]])).

cnf(refute_0_18, plain, (s1(b)),
    inference(resolve, [$cnf(p0(b, b))], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain, (~ q2(b, H, b) | s2(H)),
    inference(resolve, [$cnf(s1(b))], [refute_0_18, rule_189])).

cnf(refute_0_20, plain, (~ q2(b, b, b) | s2(b)),
    inference(subst, [], [refute_0_19 : [bind(H, $fot(b))]])).

cnf(refute_0_21, plain, (s2(b)),
    inference(resolve, [$cnf(q2(b, b, b))], [refute_0_15, refute_0_20])).

cnf(refute_0_22, plain, (r3(b, b, b)),
    inference(resolve, [$cnf(s2(b))], [refute_0_21, refute_0_8])).

cnf(refute_0_23, plain, (~ r3(b, b, b) | ~ s4(X_247) | k5(b)),
    inference(subst, [],
              [rule_300 :
               [bind(E, $fot(b)), bind(F, $fot(X_247)),
                bind(G, $fot(b))]])).

cnf(refute_0_24, plain, (~ s4(X_247) | k5(b)),
    inference(resolve, [$cnf(r3(b, b, b))], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain, (~ s4(X_247)),
    inference(resolve, [$cnf(k5(b))], [refute_0_24, prove_this])).

cnf(refute_0_26, plain, (~ s4(A)),
    inference(subst, [], [refute_0_25 : [bind(X_247, $fot(A))]])).

cnf(refute_0_27, plain, (~ l1(A, C) | ~ p3(B, C, D)),
    inference(resolve, [$cnf(s4(A))], [rule_299, refute_0_26])).

cnf(refute_0_28, plain, (~ l1(X_339, d) | ~ p3(d, d, d)),
    inference(subst, [],
              [refute_0_27 :
               [bind(A, $fot(X_339)), bind(B, $fot(d)), bind(C, $fot(d)),
                bind(D, $fot(d))]])).

cnf(refute_0_29, plain, (~ l1(X_339, d)),
    inference(resolve, [$cnf(p3(d, d, d))], [refute_0_7, refute_0_28])).

cnf(refute_0_30, plain, (~ l1(d, d)),
    inference(subst, [], [refute_0_29 : [bind(X_339, $fot(d))]])).

cnf(refute_0_31, plain, ($false),
    inference(resolve, [$cnf(l1(d, d))], [refute_0_1, refute_0_30])).
% SZS output end CNFRefutation for SYN116-1.tptp

