%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for SYN268-1.tptp

% SZS output start CNFRefutation for SYN268-1.tptp
cnf(axiom_5, axiom, (s0(b))).

cnf(axiom_37, axiom, (n0(b, a))).

cnf(rule_001, axiom, (k1(I) | ~ n0(J, I))).

cnf(rule_112, axiom, (q1(A, A, A) | ~ k1(A) | ~ s0(b))).

cnf(rule_129, axiom, (k2(J, J) | ~ q1(A, J, J))).

cnf(rule_194, axiom, (k3(F, F, G) | ~ k2(G, F))).

cnf(rule_287, axiom, (p4(B, C, B) | ~ k3(B, B, C))).

cnf(prove_this, negated_conjecture, (~ p4(a, X, Y))).

cnf(refute_0_0, plain, (~ k1(A) | q1(A, A, A)),
    inference(resolve, [$cnf(s0(b))], [axiom_5, rule_112])).

cnf(refute_0_1, plain, (~ k1(a) | q1(a, a, a)),
    inference(subst, [], [refute_0_0 : [bind(A, $fot(a))]])).

cnf(refute_0_2, plain, (~ n0(b, a) | k1(a)),
    inference(subst, [],
              [rule_001 : [bind(I, $fot(a)), bind(J, $fot(b))]])).

cnf(refute_0_3, plain, (k1(a)),
    inference(resolve, [$cnf(n0(b, a))], [axiom_37, refute_0_2])).

cnf(refute_0_4, plain, (q1(a, a, a)),
    inference(resolve, [$cnf(k1(a))], [refute_0_3, refute_0_1])).

cnf(refute_0_5, plain, (~ q1(a, a, a) | k2(a, a)),
    inference(subst, [],
              [rule_129 : [bind(A, $fot(a)), bind(J, $fot(a))]])).

cnf(refute_0_6, plain, (k2(a, a)),
    inference(resolve, [$cnf(q1(a, a, a))], [refute_0_4, refute_0_5])).

cnf(refute_0_7, plain, (~ k2(a, a) | k3(a, a, a)),
    inference(subst, [],
              [rule_194 : [bind(F, $fot(a)), bind(G, $fot(a))]])).

cnf(refute_0_8, plain, (k3(a, a, a)),
    inference(resolve, [$cnf(k2(a, a))], [refute_0_6, refute_0_7])).

cnf(refute_0_9, plain, (~ k3(a, a, a) | p4(a, a, a)),
    inference(subst, [],
              [rule_287 : [bind(B, $fot(a)), bind(C, $fot(a))]])).

cnf(refute_0_10, plain, (p4(a, a, a)),
    inference(resolve, [$cnf(k3(a, a, a))], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain, (~ p4(a, a, a)),
    inference(subst, [],
              [prove_this : [bind(X, $fot(a)), bind(Y, $fot(a))]])).

cnf(refute_0_12, plain, ($false),
    inference(resolve, [$cnf(p4(a, a, a))], [refute_0_10, refute_0_11])).
% SZS output end CNFRefutation for SYN268-1.tptp

