%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for SYN656-1.tptp

% SZS output start CNFRefutation for SYN656-1.tptp
cnf(p4_5, negated_conjecture, (p4(X75, X75))).

cnf(p3_6, negated_conjecture, (p3(X64, X64))).

cnf(p3_22, negated_conjecture,
    (p3(X65, X66) | ~ p3(X64, X65) | ~ p3(X64, X66))).

cnf(p25_31, negated_conjecture,
    (p25(f5(c29, f13(c31, f15(c32, X58))), f13(c31, X58)))).

cnf(p25_32, negated_conjecture,
    (p25(f5(f7(c28, X55), X56), X57) | ~ p25(f5(X55, X56), X57))).

cnf(p25_35, negated_conjecture,
    (p25(X51, X52) | ~ p3(X54, X52) | ~ p4(X53, X51) | ~ p25(X53, X54))).

cnf(not_p25_45, negated_conjecture,
    (~
       p25(f5(f7(c28, c29), f13(c31, f15(c32, f15(f17(c33, c34), c35)))),
         f13(c31, c34)))).

cnf(p25_46, negated_conjecture,
    (p25(f5(f7(c28, X55), X56), X57) | ~ p25(f5(X55, X56), X62) |
     ~ p25(f5(f7(c28, X55), X62), X57))).

cnf(p3_49, negated_conjecture,
    (p3(f13(c31, c34), f13(c31, f15(f17(c33, c34), X63))) |
     p25(f5(f7(c28, c29), f13(c31, f15(f17(c33, c34), X63))),
       f13(c31, c34)))).

cnf(refute_0_0, plain,
    (~ p25(f5(c29, f13(c31, f15(c32, X58))), f13(c31, X58)) |
     p25(f5(f7(c28, c29), f13(c31, f15(c32, X58))), f13(c31, X58))),
    inference(subst, [],
              [p25_32 :
               [bind(X55, $fot(c29)),
                bind(X56, $fot(f13(c31, f15(c32, X58)))),
                bind(X57, $fot(f13(c31, X58)))]])).

cnf(refute_0_1, plain,
    (p25(f5(f7(c28, c29), f13(c31, f15(c32, X58))), f13(c31, X58))),
    inference(resolve,
              [$cnf(p25(f5(c29, f13(c31, f15(c32, X58))), f13(c31, X58)))],
              [p25_31, refute_0_0])).

cnf(refute_0_2, plain, (p4(X_279, X_279)),
    inference(subst, [], [p4_5 : [bind(X75, $fot(X_279))]])).

cnf(refute_0_3, plain,
    (~ p25(X_279, X_282) | ~ p3(X_282, X_280) | ~ p4(X_279, X_279) |
     p25(X_279, X_280)),
    inference(subst, [],
              [p25_35 :
               [bind(X51, $fot(X_279)), bind(X52, $fot(X_280)),
                bind(X53, $fot(X_279)), bind(X54, $fot(X_282))]])).

cnf(refute_0_4, plain,
    (~ p25(X_279, X_282) | ~ p3(X_282, X_280) | p25(X_279, X_280)),
    inference(resolve, [$cnf(p4(X_279, X_279))],
              [refute_0_2, refute_0_3])).

cnf(refute_0_5, plain,
    (~ p25(f5(f7(c28, c29), f13(c31, f15(c32, X58))), f13(c31, X58)) |
     ~ p3(f13(c31, X58), X_284) |
     p25(f5(f7(c28, c29), f13(c31, f15(c32, X58))), X_284)),
    inference(subst, [],
              [refute_0_4 :
               [bind(X_279,
                     $fot(f5(f7(c28, c29), f13(c31, f15(c32, X58))))),
                bind(X_280, $fot(X_284)),
                bind(X_282, $fot(f13(c31, X58)))]])).

cnf(refute_0_6, plain,
    (~ p3(f13(c31, X58), X_284) |
     p25(f5(f7(c28, c29), f13(c31, f15(c32, X58))), X_284)),
    inference(resolve,
              [$cnf(p25(f5(f7(c28, c29), f13(c31, f15(c32, X58))),
                      f13(c31, X58)))], [refute_0_1, refute_0_5])).

cnf(refute_0_7, plain,
    (~ p3(f13(c31, f15(f17(c33, c34), c35)), f13(c31, c34)) |
     p25(f5(f7(c28, c29), f13(c31, f15(c32, f15(f17(c33, c34), c35)))),
       f13(c31, c34))),
    inference(subst, [],
              [refute_0_6 :
               [bind(X58, $fot(f15(f17(c33, c34), c35))),
                bind(X_284, $fot(f13(c31, c34)))]])).

cnf(refute_0_8, plain, (p3(X_79, X_79)),
    inference(subst, [], [p3_6 : [bind(X64, $fot(X_79))]])).

cnf(refute_0_9, plain,
    (~ p3(X_79, X_78) | ~ p3(X_79, X_79) | p3(X_78, X_79)),
    inference(subst, [],
              [p3_22 :
               [bind(X64, $fot(X_79)), bind(X65, $fot(X_78)),
                bind(X66, $fot(X_79))]])).

cnf(refute_0_10, plain, (~ p3(X_79, X_78) | p3(X_78, X_79)),
    inference(resolve, [$cnf(p3(X_79, X_79))], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain,
    (~ p3(f13(c31, c34), f13(c31, f15(f17(c33, c34), c35))) |
     p3(f13(c31, f15(f17(c33, c34), c35)), f13(c31, c34))),
    inference(subst, [],
              [refute_0_10 :
               [bind(X_78, $fot(f13(c31, f15(f17(c33, c34), c35)))),
                bind(X_79, $fot(f13(c31, c34)))]])).

cnf(refute_0_12, plain,
    (~ p25(f5(c29, f13(c31, f15(c32, X58))), f13(c31, X58)) |
     ~ p25(f5(f7(c28, c29), f13(c31, X58)), X_304) |
     p25(f5(f7(c28, c29), f13(c31, f15(c32, X58))), X_304)),
    inference(subst, [],
              [p25_46 :
               [bind(X55, $fot(c29)),
                bind(X56, $fot(f13(c31, f15(c32, X58)))),
                bind(X57, $fot(X_304)), bind(X62, $fot(f13(c31, X58)))]])).

cnf(refute_0_13, plain,
    (~ p25(f5(f7(c28, c29), f13(c31, X58)), X_304) |
     p25(f5(f7(c28, c29), f13(c31, f15(c32, X58))), X_304)),
    inference(resolve,
              [$cnf(p25(f5(c29, f13(c31, f15(c32, X58))), f13(c31, X58)))],
              [p25_31, refute_0_12])).

cnf(refute_0_14, plain,
    (~
       p25(f5(f7(c28, c29), f13(c31, f15(f17(c33, c34), X63))),
         f13(c31, c34)) |
     p25(f5(f7(c28, c29), f13(c31, f15(c32, f15(f17(c33, c34), X63)))),
       f13(c31, c34))),
    inference(subst, [],
              [refute_0_13 :
               [bind(X58, $fot(f15(f17(c33, c34), X63))),
                bind(X_304, $fot(f13(c31, c34)))]])).

cnf(refute_0_15, plain,
    (p25(f5(f7(c28, c29), f13(c31, f15(c32, f15(f17(c33, c34), X63)))),
       f13(c31, c34)) |
     p3(f13(c31, c34), f13(c31, f15(f17(c33, c34), X63)))),
    inference(resolve,
              [$cnf(p25(f5(f7(c28, c29),
                        f13(c31, f15(f17(c33, c34), X63))),
                      f13(c31, c34)))], [p3_49, refute_0_14])).

cnf(refute_0_16, plain,
    (p25(f5(f7(c28, c29), f13(c31, f15(c32, f15(f17(c33, c34), c35)))),
       f13(c31, c34)) |
     p3(f13(c31, c34), f13(c31, f15(f17(c33, c34), c35)))),
    inference(subst, [], [refute_0_15 : [bind(X63, $fot(c35))]])).

cnf(refute_0_17, plain,
    (p3(f13(c31, c34), f13(c31, f15(f17(c33, c34), c35)))),
    inference(resolve,
              [$cnf(p25(f5(f7(c28, c29),
                        f13(c31, f15(c32, f15(f17(c33, c34), c35)))),
                      f13(c31, c34)))], [refute_0_16, not_p25_45])).

cnf(refute_0_18, plain,
    (p3(f13(c31, f15(f17(c33, c34), c35)), f13(c31, c34))),
    inference(resolve,
              [$cnf(p3(f13(c31, c34), f13(c31, f15(f17(c33, c34), c35))))],
              [refute_0_17, refute_0_11])).

cnf(refute_0_19, plain,
    (p25(f5(f7(c28, c29), f13(c31, f15(c32, f15(f17(c33, c34), c35)))),
       f13(c31, c34))),
    inference(resolve,
              [$cnf(p3(f13(c31, f15(f17(c33, c34), c35)), f13(c31, c34)))],
              [refute_0_18, refute_0_7])).

cnf(refute_0_20, plain, ($false),
    inference(resolve,
              [$cnf(p25(f5(f7(c28, c29),
                        f13(c31, f15(c32, f15(f17(c33, c34), c35)))),
                      f13(c31, c34)))], [refute_0_19, not_p25_45])).
% SZS output end CNFRefutation for SYN656-1.tptp

