%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for SYN250-1.tptp

% SZS output start CNFRefutation for SYN250-1.tptp
cnf(axiom_9, axiom, (r0(b))).

cnf(axiom_12, axiom, (m0(a, X, a))).

cnf(axiom_13, axiom, (r0(e))).

cnf(axiom_17, axiom, (q0(X, d))).

cnf(axiom_21, axiom, (q0(b, e))).

cnf(axiom_24, axiom, (l0(c))).

cnf(axiom_29, axiom, (q0(d, b))).

cnf(axiom_37, axiom, (n0(b, a))).

cnf(rule_001, axiom, (k1(I) | ~ n0(J, I))).

cnf(rule_016, axiom, (m1(H, I, I) | ~ m1(J, I, H) | ~ m1(J, A, I))).

cnf(rule_024, axiom,
    (m1(F, a, G) | ~ m0(a, H, a) | ~ q0(F, G) | ~ m1(G, c, G))).

cnf(rule_034, axiom, (m1(A, B, B) | ~ k1(a) | ~ k1(B) | ~ q0(A, A))).

cnf(rule_035, axiom, (m1(I, J, I) | ~ r0(I) | ~ l0(J))).

cnf(prove_this, negated_conjecture, (~ m1(e, a, a))).

cnf(refute_0_0, plain, (~ n0(b, a) | k1(a)),
    inference(subst, [],
              [rule_001 : [bind(I, $fot(a)), bind(J, $fot(b))]])).

cnf(refute_0_1, plain, (k1(a)),
    inference(resolve, [$cnf(n0(b, a))], [axiom_37, refute_0_0])).

cnf(refute_0_2, plain, (q0(d, d)),
    inference(subst, [], [axiom_17 : [bind(X, $fot(d))]])).

cnf(refute_0_3, plain, (~ k1(B) | ~ q0(A, A) | m1(A, B, B)),
    inference(resolve, [$cnf(k1(a))], [refute_0_1, rule_034])).

cnf(refute_0_4, plain, (~ k1(X_331) | ~ q0(d, d) | m1(d, X_331, X_331)),
    inference(subst, [],
              [refute_0_3 : [bind(A, $fot(d)), bind(B, $fot(X_331))]])).

cnf(refute_0_5, plain, (~ k1(X_331) | m1(d, X_331, X_331)),
    inference(resolve, [$cnf(q0(d, d))], [refute_0_2, refute_0_4])).

cnf(refute_0_6, plain, (~ k1(a) | m1(d, a, a)),
    inference(subst, [], [refute_0_5 : [bind(X_331, $fot(a))]])).

cnf(refute_0_7, plain, (m1(d, a, a)),
    inference(resolve, [$cnf(k1(a))], [refute_0_1, refute_0_6])).

cnf(refute_0_8, plain, (m0(a, H, a)),
    inference(subst, [], [axiom_12 : [bind(X, $fot(H))]])).

cnf(refute_0_9, plain, (~ m1(G, c, G) | ~ q0(F, G) | m1(F, a, G)),
    inference(resolve, [$cnf(m0(a, H, a))], [refute_0_8, rule_024])).

cnf(refute_0_10, plain, (~ m1(b, c, b) | ~ q0(F, b) | m1(F, a, b)),
    inference(subst, [], [refute_0_9 : [bind(G, $fot(b))]])).

cnf(refute_0_11, plain, (~ l0(X_174) | ~ r0(b) | m1(b, X_174, b)),
    inference(subst, [],
              [rule_035 : [bind(I, $fot(b)), bind(J, $fot(X_174))]])).

cnf(refute_0_12, plain, (~ l0(X_174) | m1(b, X_174, b)),
    inference(resolve, [$cnf(r0(b))], [axiom_9, refute_0_11])).

cnf(refute_0_13, plain, (~ l0(c) | m1(b, c, b)),
    inference(subst, [], [refute_0_12 : [bind(X_174, $fot(c))]])).

cnf(refute_0_14, plain, (m1(b, c, b)),
    inference(resolve, [$cnf(l0(c))], [axiom_24, refute_0_13])).

cnf(refute_0_15, plain, (~ q0(F, b) | m1(F, a, b)),
    inference(resolve, [$cnf(m1(b, c, b))], [refute_0_14, refute_0_10])).

cnf(refute_0_16, plain, (~ q0(d, b) | m1(d, a, b)),
    inference(subst, [], [refute_0_15 : [bind(F, $fot(d))]])).

cnf(refute_0_17, plain, (m1(d, a, b)),
    inference(resolve, [$cnf(q0(d, b))], [axiom_29, refute_0_16])).

cnf(refute_0_18, plain, (~ m1(d, X_440, a) | ~ m1(d, a, b) | m1(b, a, a)),
    inference(subst, [],
              [rule_016 :
               [bind(A, $fot(X_440)), bind(H, $fot(b)), bind(I, $fot(a)),
                bind(J, $fot(d))]])).

cnf(refute_0_19, plain, (~ m1(d, X_440, a) | m1(b, a, a)),
    inference(resolve, [$cnf(m1(d, a, b))], [refute_0_17, refute_0_18])).

cnf(refute_0_20, plain, (~ m1(e, c, e) | ~ q0(F, e) | m1(F, a, e)),
    inference(subst, [], [refute_0_9 : [bind(G, $fot(e))]])).

cnf(refute_0_21, plain, (~ l0(X_174) | ~ r0(e) | m1(e, X_174, e)),
    inference(subst, [],
              [rule_035 : [bind(I, $fot(e)), bind(J, $fot(X_174))]])).

cnf(refute_0_22, plain, (~ l0(X_174) | m1(e, X_174, e)),
    inference(resolve, [$cnf(r0(e))], [axiom_13, refute_0_21])).

cnf(refute_0_23, plain, (~ l0(c) | m1(e, c, e)),
    inference(subst, [], [refute_0_22 : [bind(X_174, $fot(c))]])).

cnf(refute_0_24, plain, (m1(e, c, e)),
    inference(resolve, [$cnf(l0(c))], [axiom_24, refute_0_23])).

cnf(refute_0_25, plain, (~ q0(F, e) | m1(F, a, e)),
    inference(resolve, [$cnf(m1(e, c, e))], [refute_0_24, refute_0_20])).

cnf(refute_0_26, plain, (~ q0(b, e) | m1(b, a, e)),
    inference(subst, [], [refute_0_25 : [bind(F, $fot(b))]])).

cnf(refute_0_27, plain, (m1(b, a, e)),
    inference(resolve, [$cnf(q0(b, e))], [axiom_21, refute_0_26])).

cnf(refute_0_28, plain, (~ m1(b, X_440, a) | ~ m1(b, a, e) | m1(e, a, a)),
    inference(subst, [],
              [rule_016 :
               [bind(A, $fot(X_440)), bind(H, $fot(e)), bind(I, $fot(a)),
                bind(J, $fot(b))]])).

cnf(refute_0_29, plain, (~ m1(b, X_440, a) | m1(e, a, a)),
    inference(resolve, [$cnf(m1(b, a, e))], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain, (~ m1(b, X_440, a)),
    inference(resolve, [$cnf(m1(e, a, a))], [refute_0_29, prove_this])).

cnf(refute_0_31, plain, (~ m1(b, a, a)),
    inference(subst, [], [refute_0_30 : [bind(X_440, $fot(a))]])).

cnf(refute_0_32, plain, (~ m1(d, X_440, a)),
    inference(resolve, [$cnf(m1(b, a, a))], [refute_0_19, refute_0_31])).

cnf(refute_0_33, plain, (~ m1(d, a, a)),
    inference(subst, [], [refute_0_32 : [bind(X_440, $fot(a))]])).

cnf(refute_0_34, plain, ($false),
    inference(resolve, [$cnf(m1(d, a, a))], [refute_0_7, refute_0_33])).
% SZS output end CNFRefutation for SYN250-1.tptp

