%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
% SZS status Unsatisfiable for SYN206-1.tptp

% SZS output start CNFRefutation for SYN206-1.tptp
cnf(axiom_14, axiom, (p0(b, X))).

cnf(axiom_19, axiom, (m0(X, d, Y))).

cnf(axiom_32, axiom, (k0(b))).

cnf(rule_068, axiom, (p1(D, D, D) | ~ k0(D))).

cnf(rule_125, axiom, (s1(I) | ~ p0(I, I))).

cnf(rule_181, axiom, (q2(I, I, I) | ~ p1(I, I, I))).

cnf(rule_189, axiom, (s2(H) | ~ q2(b, H, b) | ~ s1(b))).

cnf(rule_273, axiom, (s3(I, J) | ~ q2(A, I, A) | ~ s2(I) | ~ m0(A, B, J))).

cnf(prove_this, negated_conjecture, (~ s3(b, b))).

cnf(refute_0_0, plain, (m0(b, d, X_884)),
    inference(subst, [],
              [axiom_19 : [bind(X, $fot(b)), bind(Y, $fot(X_884))]])).

cnf(refute_0_1, plain, (~ k0(b) | p1(b, b, b)),
    inference(subst, [], [rule_068 : [bind(D, $fot(b))]])).

cnf(refute_0_2, plain, (p1(b, b, b)),
    inference(resolve, [$cnf(k0(b))], [axiom_32, refute_0_1])).

cnf(refute_0_3, plain, (~ p1(b, b, b) | q2(b, b, b)),
    inference(subst, [], [rule_181 : [bind(I, $fot(b))]])).

cnf(refute_0_4, plain, (q2(b, b, b)),
    inference(resolve, [$cnf(p1(b, b, b))], [refute_0_2, refute_0_3])).

cnf(refute_0_5, plain,
    (~ m0(b, X_877, X_879) | ~ q2(b, b, b) | ~ s2(b) | s3(b, X_879)),
    inference(subst, [],
              [rule_273 :
               [bind(A, $fot(b)), bind(B, $fot(X_877)), bind(I, $fot(b)),
                bind(J, $fot(X_879))]])).

cnf(refute_0_6, plain, (~ m0(b, X_877, X_879) | ~ s2(b) | s3(b, X_879)),
    inference(resolve, [$cnf(q2(b, b, b))], [refute_0_4, refute_0_5])).

cnf(refute_0_7, plain, (p0(b, b)),
    inference(subst, [], [axiom_14 : [bind(X, $fot(b))]])).

cnf(refute_0_8, plain, (~ p0(b, b) | s1(b)),
    inference(subst, [], [rule_125 : [bind(I, $fot(b))]])).

cnf(refute_0_9, plain, (s1(b)),
    inference(resolve, [$cnf(p0(b, b))], [refute_0_7, refute_0_8])).

cnf(refute_0_10, plain, (~ q2(b, H, b) | s2(H)),
    inference(resolve, [$cnf(s1(b))], [refute_0_9, rule_189])).

cnf(refute_0_11, plain, (~ q2(b, b, b) | s2(b)),
    inference(subst, [], [refute_0_10 : [bind(H, $fot(b))]])).

cnf(refute_0_12, plain, (s2(b)),
    inference(resolve, [$cnf(q2(b, b, b))], [refute_0_4, refute_0_11])).

cnf(refute_0_13, plain, (~ m0(b, X_877, X_879) | s3(b, X_879)),
    inference(resolve, [$cnf(s2(b))], [refute_0_12, refute_0_6])).

cnf(refute_0_14, plain, (~ m0(b, d, X_884) | s3(b, X_884)),
    inference(subst, [],
              [refute_0_13 :
               [bind(X_877, $fot(d)), bind(X_879, $fot(X_884))]])).

cnf(refute_0_15, plain, (s3(b, X_884)),
    inference(resolve, [$cnf(m0(b, d, X_884))],
              [refute_0_0, refute_0_14])).

cnf(refute_0_16, plain, (s3(b, b)),
    inference(subst, [], [refute_0_15 : [bind(X_884, $fot(b))]])).

cnf(refute_0_17, plain, ($false),
    inference(resolve, [$cnf(s3(b, b))], [refute_0_16, prove_this])).
% SZS output end CNFRefutation for SYN206-1.tptp

