% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWV568-1.050
% SZS output start Proof for SWV568-1.050
fof(f86675,plain,(
  $false),
  inference(subsumption_resolution,[],[f86674,f498])).
fof(f498,axiom,(
  e16 != elem_352),
  file('SWV568-1.050.tptp',unknown)).
fof(f86674,plain,(
  e16 = elem_352),
  inference(backward_demodulation,[],[f86673,f86613])).
fof(f86613,plain,(
  elem_352 = select(earray_100,index_57)),
  inference(subsumption_resolution,[],[f86611,f3436])).
fof(f3436,plain,(
  index_105 != index_57),
  inference(forward_demodulation,[],[f3435,f903])).
fof(f903,plain,(
  index_105 = s(index_99)),
  inference(backward_demodulation,[],[f900,f197])).
fof(f197,axiom,(
  index_102 = s(index_99)),
  file('SWV568-1.050.tptp',unknown)).
fof(f900,plain,(
  index_105 = index_102),
  inference(forward_demodulation,[],[f850,f550])).
fof(f550,plain,(
  index_105 = rselect_tail(queue_103)),
  inference(definition_unfolding,[],[f198,f462])).
fof(f462,axiom,(
  q22 = queue_103),
  file('SWV568-1.050.tptp',unknown)).
fof(f198,axiom,(
  index_105 = rselect_tail(q22)),
  file('SWV568-1.050.tptp',unknown)).
fof(f850,plain,(
  index_102 = rselect_tail(queue_103)),
  inference(superposition,[],[f5,f332])).
fof(f332,axiom,(
  queue_103 = rstore_tail(queue_101,index_102)),
  file('SWV568-1.050.tptp',unknown)).
fof(f5,axiom,(
  ( ! [X2,X0] : (rselect_tail(rstore_tail(X0,X2)) = X2) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f3435,plain,(
  index_57 != s(index_99)),
  inference(forward_demodulation,[],[f3434,f2517])).
fof(f2517,plain,(
  index_99 = s(index_90)),
  inference(backward_demodulation,[],[f2516,f326])).
fof(f326,axiom,(
  index_93 = s(index_90)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2516,plain,(
  index_99 = index_93),
  inference(forward_demodulation,[],[f2515,f600])).
fof(f600,plain,(
  index_99 = rselect_tail(queue_97)),
  inference(definition_unfolding,[],[f329,f461])).
fof(f461,axiom,(
  q21 = queue_97),
  file('SWV568-1.050.tptp',unknown)).
fof(f329,axiom,(
  index_99 = rselect_tail(q21)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2515,plain,(
  index_93 = rselect_tail(queue_97)),
  inference(forward_demodulation,[],[f2071,f899])).
fof(f899,plain,(
  index_93 = rselect_tail(queue_94)),
  inference(superposition,[],[f5,f445])).
fof(f445,axiom,(
  queue_94 = rstore_tail(queue_92,index_93)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2071,plain,(
  rselect_tail(queue_97) = rselect_tail(queue_94)),
  inference(superposition,[],[f11,f1869])).
fof(f1869,plain,(
  queue_97 = rstore_head(queue_94,index_116)),
  inference(backward_demodulation,[],[f1868,f446])).
fof(f446,axiom,(
  queue_97 = rstore_head(queue_94,index_96)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1868,plain,(
  index_116 = index_96),
  inference(backward_demodulation,[],[f1867,f849])).
fof(f849,plain,(
  index_96 = rselect_head(queue_97)),
  inference(superposition,[],[f4,f446])).
fof(f4,axiom,(
  ( ! [X2,X0] : (rselect_head(rstore_head(X0,X2)) = X2) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f1867,plain,(
  index_116 = rselect_head(queue_97)),
  inference(forward_demodulation,[],[f1650,f1656])).
fof(f1656,plain,(
  index_116 = rselect_head(queue_101)),
  inference(backward_demodulation,[],[f1655,f1654])).
fof(f1654,plain,(
  rselect_head(queue_101) = rselect_head(queue_109)),
  inference(backward_demodulation,[],[f1653,f1535])).
fof(f1535,plain,(
  rselect_head(queue_107) = rselect_head(queue_109)),
  inference(superposition,[],[f8,f906])).
fof(f906,plain,(
  queue_109 = rstore_tail(queue_107,index_111)),
  inference(backward_demodulation,[],[f904,f334])).
fof(f334,axiom,(
  queue_109 = rstore_tail(queue_107,index_108)),
  file('SWV568-1.050.tptp',unknown)).
fof(f904,plain,(
  index_111 = index_108),
  inference(forward_demodulation,[],[f851,f551])).
fof(f551,plain,(
  index_111 = rselect_tail(queue_109)),
  inference(definition_unfolding,[],[f200,f463])).
fof(f463,axiom,(
  q23 = queue_109),
  file('SWV568-1.050.tptp',unknown)).
fof(f200,axiom,(
  index_111 = rselect_tail(q23)),
  file('SWV568-1.050.tptp',unknown)).
fof(f851,plain,(
  index_108 = rselect_tail(queue_109)),
  inference(superposition,[],[f5,f334])).
fof(f8,axiom,(
  ( ! [X2,X0] : (rselect_head(rstore_tail(X0,X2)) = rselect_head(X0)) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f1653,plain,(
  rselect_head(queue_101) = rselect_head(queue_107)),
  inference(forward_demodulation,[],[f1602,f1534])).
fof(f1534,plain,(
  rselect_head(queue_101) = rselect_head(queue_103)),
  inference(superposition,[],[f8,f902])).
fof(f902,plain,(
  queue_103 = rstore_tail(queue_101,index_105)),
  inference(backward_demodulation,[],[f900,f332])).
fof(f1602,plain,(
  rselect_head(queue_103) = rselect_head(queue_107)),
  inference(superposition,[],[f9,f602])).
fof(f602,plain,(
  queue_107 = rstore_seq(queue_103,earray_106)),
  inference(definition_unfolding,[],[f333,f462])).
fof(f333,axiom,(
  queue_107 = rstore_seq(q22,earray_106)),
  file('SWV568-1.050.tptp',unknown)).
fof(f9,axiom,(
  ( ! [X2,X0] : (rselect_head(X0) = rselect_head(rstore_seq(X0,X2))) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f1655,plain,(
  index_116 = rselect_head(queue_109)),
  inference(forward_demodulation,[],[f1603,f1584])).
fof(f1584,plain,(
  index_116 = rselect_head(queue_113)),
  inference(forward_demodulation,[],[f1537,f202])).
fof(f202,axiom,(
  index_116 = rselect_head(queue_115)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1537,plain,(
  rselect_head(queue_115) = rselect_head(queue_113)),
  inference(superposition,[],[f8,f337])).
fof(f337,axiom,(
  queue_115 = rstore_tail(queue_113,index_114)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1603,plain,(
  rselect_head(queue_109) = rselect_head(queue_113)),
  inference(superposition,[],[f9,f604])).
fof(f604,plain,(
  queue_113 = rstore_seq(queue_109,earray_112)),
  inference(definition_unfolding,[],[f336,f463])).
fof(f336,axiom,(
  queue_113 = rstore_seq(q23,earray_112)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1650,plain,(
  rselect_head(queue_97) = rselect_head(queue_101)),
  inference(superposition,[],[f9,f601])).
fof(f601,plain,(
  queue_101 = rstore_seq(queue_97,earray_100)),
  inference(definition_unfolding,[],[f331,f461])).
fof(f331,axiom,(
  queue_101 = rstore_seq(q21,earray_100)),
  file('SWV568-1.050.tptp',unknown)).
fof(f11,axiom,(
  ( ! [X2,X0] : (rselect_tail(X0) = rselect_tail(rstore_head(X0,X2))) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f3434,plain,(
  index_57 != s(s(index_90))),
  inference(forward_demodulation,[],[f3433,f1033])).
fof(f1033,plain,(
  index_90 = s(index_84)),
  inference(backward_demodulation,[],[f1030,f323])).
fof(f323,axiom,(
  index_87 = s(index_84)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1030,plain,(
  index_90 = index_87),
  inference(forward_demodulation,[],[f898,f599])).
fof(f599,plain,(
  index_90 = rselect_tail(queue_88)),
  inference(definition_unfolding,[],[f325,f460])).
fof(f460,axiom,(
  q20 = queue_88),
  file('SWV568-1.050.tptp',unknown)).
fof(f325,axiom,(
  index_90 = rselect_tail(q20)),
  file('SWV568-1.050.tptp',unknown)).
fof(f898,plain,(
  index_87 = rselect_tail(queue_88)),
  inference(superposition,[],[f5,f443])).
fof(f443,axiom,(
  queue_88 = rstore_tail(queue_86,index_87)),
  file('SWV568-1.050.tptp',unknown)).
fof(f3433,plain,(
  index_57 != s(s(s(index_84)))),
  inference(forward_demodulation,[],[f3432,f1023])).
fof(f1023,plain,(
  index_84 = s(index_72)),
  inference(backward_demodulation,[],[f1020,f319])).
fof(f319,axiom,(
  index_75 = s(index_72)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1020,plain,(
  index_84 = index_75),
  inference(forward_demodulation,[],[f896,f597])).
fof(f597,plain,(
  index_84 = rselect_tail(queue_76)),
  inference(definition_unfolding,[],[f322,f458])).
fof(f458,axiom,(
  q19 = queue_76),
  file('SWV568-1.050.tptp',unknown)).
fof(f322,axiom,(
  index_84 = rselect_tail(q19)),
  file('SWV568-1.050.tptp',unknown)).
fof(f896,plain,(
  index_75 = rselect_tail(queue_76)),
  inference(superposition,[],[f5,f439])).
fof(f439,axiom,(
  queue_76 = rstore_tail(queue_74,index_75)),
  file('SWV568-1.050.tptp',unknown)).
fof(f3432,plain,(
  index_57 != s(s(s(s(index_72))))),
  inference(forward_demodulation,[],[f3254,f2508])).
fof(f2508,plain,(
  index_72 = s(index_63)),
  inference(backward_demodulation,[],[f2507,f315])).
fof(f315,axiom,(
  index_66 = s(index_63)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2507,plain,(
  index_72 = index_66),
  inference(forward_demodulation,[],[f2506,f595])).
fof(f595,plain,(
  index_72 = rselect_tail(queue_70)),
  inference(definition_unfolding,[],[f318,f457])).
fof(f457,axiom,(
  q18 = queue_70),
  file('SWV568-1.050.tptp',unknown)).
fof(f318,axiom,(
  index_72 = rselect_tail(q18)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2506,plain,(
  index_66 = rselect_tail(queue_70)),
  inference(forward_demodulation,[],[f2070,f895])).
fof(f895,plain,(
  index_66 = rselect_tail(queue_67)),
  inference(superposition,[],[f5,f435])).
fof(f435,axiom,(
  queue_67 = rstore_tail(queue_65,index_66)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2070,plain,(
  rselect_tail(queue_70) = rselect_tail(queue_67)),
  inference(superposition,[],[f11,f437])).
fof(f437,axiom,(
  queue_70 = rstore_head(queue_67,index_69)),
  file('SWV568-1.050.tptp',unknown)).
fof(f3254,plain,(
  index_57 != s(s(s(s(s(index_63)))))),
  inference(superposition,[],[f88,f1019])).
fof(f1019,plain,(
  index_63 = s(index_57)),
  inference(backward_demodulation,[],[f1016,f313])).
fof(f313,axiom,(
  index_60 = s(index_57)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1016,plain,(
  index_63 = index_60),
  inference(forward_demodulation,[],[f894,f594])).
fof(f594,plain,(
  index_63 = rselect_tail(queue_61)),
  inference(definition_unfolding,[],[f314,f456])).
fof(f456,axiom,(
  q17 = queue_61),
  file('SWV568-1.050.tptp',unknown)).
fof(f314,axiom,(
  index_63 = rselect_tail(q17)),
  file('SWV568-1.050.tptp',unknown)).
fof(f894,plain,(
  index_60 = rselect_tail(queue_61)),
  inference(superposition,[],[f5,f433])).
fof(f433,axiom,(
  queue_61 = rstore_tail(queue_59,index_60)),
  file('SWV568-1.050.tptp',unknown)).
fof(f88,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(X4)))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86611,plain,(
  elem_352 = select(earray_100,index_57) | index_105 = index_57),
  inference(superposition,[],[f86610,f4486])).
fof(f4486,plain,(
  ( ! [X2] : (select(earray_100,X2) = select(earray_106,X2) | index_105 = X2) )),
  inference(superposition,[],[f2,f1351])).
fof(f1351,plain,(
  earray_106 = store(earray_100,index_105,e22)),
  inference(backward_demodulation,[],[f1350,f97])).
fof(f97,axiom,(
  earray_106 = store(earray_104,index_105,e22)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1350,plain,(
  earray_100 = earray_104),
  inference(backward_demodulation,[],[f1349,f499])).
fof(f499,plain,(
  earray_104 = rselect_seq(queue_103)),
  inference(definition_unfolding,[],[f96,f462])).
fof(f96,axiom,(
  earray_104 = rselect_seq(q22)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1349,plain,(
  earray_100 = rselect_seq(queue_103)),
  inference(forward_demodulation,[],[f1299,f832])).
fof(f832,plain,(
  earray_100 = rselect_seq(queue_101)),
  inference(superposition,[],[f3,f601])).
fof(f3,axiom,(
  ( ! [X2,X0] : (rselect_seq(rstore_seq(X0,X2)) = X2) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f1299,plain,(
  rselect_seq(queue_103) = rselect_seq(queue_101)),
  inference(superposition,[],[f7,f902])).
fof(f7,axiom,(
  ( ! [X2,X0] : (rselect_seq(X0) = rselect_seq(rstore_tail(X0,X2))) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f2,axiom,(
  ( ! [X2,X0,X3,X1] : (select(store(X0,X1,X2),X3) = select(X0,X3) | X1 = X3) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86610,plain,(
  elem_352 = select(earray_106,index_57)),
  inference(subsumption_resolution,[],[f86608,f3729])).
fof(f3729,plain,(
  index_111 != index_57),
  inference(forward_demodulation,[],[f3728,f907])).
fof(f907,plain,(
  index_111 = s(index_105)),
  inference(backward_demodulation,[],[f904,f199])).
fof(f199,axiom,(
  index_108 = s(index_105)),
  file('SWV568-1.050.tptp',unknown)).
fof(f3728,plain,(
  index_57 != s(index_105)),
  inference(forward_demodulation,[],[f3727,f903])).
fof(f3727,plain,(
  index_57 != s(s(index_99))),
  inference(forward_demodulation,[],[f3726,f2517])).
fof(f3726,plain,(
  index_57 != s(s(s(index_90)))),
  inference(forward_demodulation,[],[f3725,f1033])).
fof(f3725,plain,(
  index_57 != s(s(s(s(index_84))))),
  inference(forward_demodulation,[],[f3724,f1023])).
fof(f3724,plain,(
  index_57 != s(s(s(s(s(index_72)))))),
  inference(forward_demodulation,[],[f3506,f2508])).
fof(f3506,plain,(
  index_57 != s(s(s(s(s(s(index_63))))))),
  inference(superposition,[],[f87,f1019])).
fof(f87,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(X4))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86608,plain,(
  elem_352 = select(earray_106,index_57) | index_111 = index_57),
  inference(superposition,[],[f86607,f4488])).
fof(f4488,plain,(
  ( ! [X4] : (select(earray_106,X4) = select(earray_112,X4) | index_111 = X4) )),
  inference(superposition,[],[f2,f1354])).
fof(f1354,plain,(
  earray_112 = store(earray_106,index_111,e23)),
  inference(backward_demodulation,[],[f1353,f99])).
fof(f99,axiom,(
  earray_112 = store(earray_110,index_111,e23)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1353,plain,(
  earray_106 = earray_110),
  inference(backward_demodulation,[],[f1352,f500])).
fof(f500,plain,(
  earray_110 = rselect_seq(queue_109)),
  inference(definition_unfolding,[],[f98,f463])).
fof(f98,axiom,(
  earray_110 = rselect_seq(q23)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1352,plain,(
  earray_106 = rselect_seq(queue_109)),
  inference(forward_demodulation,[],[f1300,f784])).
fof(f784,plain,(
  earray_106 = rselect_seq(queue_107)),
  inference(superposition,[],[f3,f602])).
fof(f1300,plain,(
  rselect_seq(queue_109) = rselect_seq(queue_107)),
  inference(superposition,[],[f7,f906])).
fof(f86607,plain,(
  elem_352 = select(earray_112,index_57)),
  inference(subsumption_resolution,[],[f86605,f4065])).
fof(f4065,plain,(
  index_120 != index_57),
  inference(forward_demodulation,[],[f4064,f2139])).
fof(f2139,plain,(
  index_120 = s(index_111)),
  inference(backward_demodulation,[],[f2138,f201])).
fof(f201,axiom,(
  index_114 = s(index_111)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2138,plain,(
  index_120 = index_114),
  inference(forward_demodulation,[],[f2137,f552])).
fof(f552,plain,(
  index_120 = rselect_tail(queue_118)),
  inference(definition_unfolding,[],[f205,f464])).
fof(f464,axiom,(
  q24 = queue_118),
  file('SWV568-1.050.tptp',unknown)).
fof(f205,axiom,(
  index_120 = rselect_tail(q24)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2137,plain,(
  index_114 = rselect_tail(queue_118)),
  inference(forward_demodulation,[],[f2056,f853])).
fof(f853,plain,(
  index_114 = rselect_tail(queue_115)),
  inference(superposition,[],[f5,f337])).
fof(f2056,plain,(
  rselect_tail(queue_118) = rselect_tail(queue_115)),
  inference(superposition,[],[f11,f338])).
fof(f338,axiom,(
  queue_118 = rstore_head(queue_115,index_117)),
  file('SWV568-1.050.tptp',unknown)).
fof(f4064,plain,(
  index_57 != s(index_111)),
  inference(forward_demodulation,[],[f4063,f907])).
fof(f4063,plain,(
  index_57 != s(s(index_105))),
  inference(forward_demodulation,[],[f4062,f903])).
fof(f4062,plain,(
  index_57 != s(s(s(index_99)))),
  inference(forward_demodulation,[],[f4061,f2517])).
fof(f4061,plain,(
  index_57 != s(s(s(s(index_90))))),
  inference(forward_demodulation,[],[f4060,f1033])).
fof(f4060,plain,(
  index_57 != s(s(s(s(s(index_84)))))),
  inference(forward_demodulation,[],[f4059,f1023])).
fof(f4059,plain,(
  index_57 != s(s(s(s(s(s(index_72))))))),
  inference(forward_demodulation,[],[f3804,f2508])).
fof(f3804,plain,(
  index_57 != s(s(s(s(s(s(s(index_63)))))))),
  inference(superposition,[],[f86,f1019])).
fof(f86,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(X4)))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86605,plain,(
  elem_352 = select(earray_112,index_57) | index_120 = index_57),
  inference(superposition,[],[f86604,f4489])).
fof(f4489,plain,(
  ( ! [X5] : (select(earray_112,X5) = select(earray_121,X5) | index_120 = X5) )),
  inference(superposition,[],[f2,f1361])).
fof(f1361,plain,(
  earray_121 = store(earray_112,index_120,e24)),
  inference(backward_demodulation,[],[f1359,f101])).
fof(f101,axiom,(
  earray_121 = store(earray_119,index_120,e24)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1359,plain,(
  earray_112 = earray_119),
  inference(forward_demodulation,[],[f1358,f785])).
fof(f785,plain,(
  earray_112 = rselect_seq(queue_113)),
  inference(superposition,[],[f3,f604])).
fof(f1358,plain,(
  earray_119 = rselect_seq(queue_113)),
  inference(forward_demodulation,[],[f1302,f1283])).
fof(f1283,plain,(
  earray_119 = rselect_seq(queue_115)),
  inference(forward_demodulation,[],[f1266,f501])).
fof(f501,plain,(
  earray_119 = rselect_seq(queue_118)),
  inference(definition_unfolding,[],[f100,f464])).
fof(f100,axiom,(
  earray_119 = rselect_seq(q24)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1266,plain,(
  rselect_seq(queue_118) = rselect_seq(queue_115)),
  inference(superposition,[],[f6,f338])).
fof(f6,axiom,(
  ( ! [X2,X0] : (rselect_seq(rstore_head(X0,X2)) = rselect_seq(X0)) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f1302,plain,(
  rselect_seq(queue_113) = rselect_seq(queue_115)),
  inference(superposition,[],[f7,f337])).
fof(f86604,plain,(
  elem_352 = select(earray_121,index_57)),
  inference(subsumption_resolution,[],[f86602,f4443])).
fof(f4443,plain,(
  index_126 != index_57),
  inference(forward_demodulation,[],[f4442,f915])).
fof(f915,plain,(
  index_126 = s(index_120)),
  inference(backward_demodulation,[],[f912,f206])).
fof(f206,axiom,(
  index_123 = s(index_120)),
  file('SWV568-1.050.tptp',unknown)).
fof(f912,plain,(
  index_126 = index_123),
  inference(forward_demodulation,[],[f854,f553])).
fof(f553,plain,(
  index_126 = rselect_tail(queue_124)),
  inference(definition_unfolding,[],[f207,f465])).
fof(f465,axiom,(
  q25 = queue_124),
  file('SWV568-1.050.tptp',unknown)).
fof(f207,axiom,(
  index_126 = rselect_tail(q25)),
  file('SWV568-1.050.tptp',unknown)).
fof(f854,plain,(
  index_123 = rselect_tail(queue_124)),
  inference(superposition,[],[f5,f340])).
fof(f340,axiom,(
  queue_124 = rstore_tail(queue_122,index_123)),
  file('SWV568-1.050.tptp',unknown)).
fof(f4442,plain,(
  index_57 != s(index_120)),
  inference(forward_demodulation,[],[f4441,f2139])).
fof(f4441,plain,(
  index_57 != s(s(index_111))),
  inference(forward_demodulation,[],[f4440,f907])).
fof(f4440,plain,(
  index_57 != s(s(s(index_105)))),
  inference(forward_demodulation,[],[f4439,f903])).
fof(f4439,plain,(
  index_57 != s(s(s(s(index_99))))),
  inference(forward_demodulation,[],[f4438,f2517])).
fof(f4438,plain,(
  index_57 != s(s(s(s(s(index_90)))))),
  inference(forward_demodulation,[],[f4437,f1033])).
fof(f4437,plain,(
  index_57 != s(s(s(s(s(s(index_84))))))),
  inference(forward_demodulation,[],[f4436,f1023])).
fof(f4436,plain,(
  index_57 != s(s(s(s(s(s(s(index_72)))))))),
  inference(forward_demodulation,[],[f4145,f2508])).
fof(f4145,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_63))))))))),
  inference(superposition,[],[f85,f1019])).
fof(f85,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(X4))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86602,plain,(
  elem_352 = select(earray_121,index_57) | index_126 = index_57),
  inference(superposition,[],[f86601,f4490])).
fof(f4490,plain,(
  ( ! [X6] : (select(earray_121,X6) = select(earray_127,X6) | index_126 = X6) )),
  inference(superposition,[],[f2,f1365])).
fof(f1365,plain,(
  earray_127 = store(earray_121,index_126,e25)),
  inference(backward_demodulation,[],[f1364,f103])).
fof(f103,axiom,(
  earray_127 = store(earray_125,index_126,e25)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1364,plain,(
  earray_121 = earray_125),
  inference(backward_demodulation,[],[f1363,f502])).
fof(f502,plain,(
  earray_125 = rselect_seq(queue_124)),
  inference(definition_unfolding,[],[f102,f465])).
fof(f102,axiom,(
  earray_125 = rselect_seq(q25)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1363,plain,(
  earray_121 = rselect_seq(queue_124)),
  inference(forward_demodulation,[],[f1303,f786])).
fof(f786,plain,(
  earray_121 = rselect_seq(queue_122)),
  inference(superposition,[],[f3,f605])).
fof(f605,plain,(
  queue_122 = rstore_seq(queue_118,earray_121)),
  inference(definition_unfolding,[],[f339,f464])).
fof(f339,axiom,(
  queue_122 = rstore_seq(q24,earray_121)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1303,plain,(
  rselect_seq(queue_124) = rselect_seq(queue_122)),
  inference(superposition,[],[f7,f914])).
fof(f914,plain,(
  queue_124 = rstore_tail(queue_122,index_126)),
  inference(backward_demodulation,[],[f912,f340])).
fof(f86601,plain,(
  elem_352 = select(earray_127,index_57)),
  inference(subsumption_resolution,[],[f86599,f4916])).
fof(f4916,plain,(
  index_132 != index_57),
  inference(forward_demodulation,[],[f4915,f919])).
fof(f919,plain,(
  index_132 = s(index_126)),
  inference(backward_demodulation,[],[f916,f208])).
fof(f208,axiom,(
  index_129 = s(index_126)),
  file('SWV568-1.050.tptp',unknown)).
fof(f916,plain,(
  index_132 = index_129),
  inference(forward_demodulation,[],[f855,f554])).
fof(f554,plain,(
  index_132 = rselect_tail(queue_130)),
  inference(definition_unfolding,[],[f209,f466])).
fof(f466,axiom,(
  q26 = queue_130),
  file('SWV568-1.050.tptp',unknown)).
fof(f209,axiom,(
  index_132 = rselect_tail(q26)),
  file('SWV568-1.050.tptp',unknown)).
fof(f855,plain,(
  index_129 = rselect_tail(queue_130)),
  inference(superposition,[],[f5,f343])).
fof(f343,axiom,(
  queue_130 = rstore_tail(queue_128,index_129)),
  file('SWV568-1.050.tptp',unknown)).
fof(f4915,plain,(
  index_57 != s(index_126)),
  inference(forward_demodulation,[],[f4914,f915])).
fof(f4914,plain,(
  index_57 != s(s(index_120))),
  inference(forward_demodulation,[],[f4913,f2139])).
fof(f4913,plain,(
  index_57 != s(s(s(index_111)))),
  inference(forward_demodulation,[],[f4912,f907])).
fof(f4912,plain,(
  index_57 != s(s(s(s(index_105))))),
  inference(forward_demodulation,[],[f4911,f903])).
fof(f4911,plain,(
  index_57 != s(s(s(s(s(index_99)))))),
  inference(forward_demodulation,[],[f4910,f2517])).
fof(f4910,plain,(
  index_57 != s(s(s(s(s(s(index_90))))))),
  inference(forward_demodulation,[],[f4909,f1033])).
fof(f4909,plain,(
  index_57 != s(s(s(s(s(s(s(index_84)))))))),
  inference(forward_demodulation,[],[f4908,f1023])).
fof(f4908,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_72))))))))),
  inference(forward_demodulation,[],[f4580,f2508])).
fof(f4580,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_63)))))))))),
  inference(superposition,[],[f84,f1019])).
fof(f84,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(X4)))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86599,plain,(
  elem_352 = select(earray_127,index_57) | index_132 = index_57),
  inference(superposition,[],[f86598,f4491])).
fof(f4491,plain,(
  ( ! [X7] : (select(earray_127,X7) = select(earray_133,X7) | index_132 = X7) )),
  inference(superposition,[],[f2,f1368])).
fof(f1368,plain,(
  earray_133 = store(earray_127,index_132,e26)),
  inference(backward_demodulation,[],[f1367,f105])).
fof(f105,axiom,(
  earray_133 = store(earray_131,index_132,e26)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1367,plain,(
  earray_127 = earray_131),
  inference(backward_demodulation,[],[f1366,f503])).
fof(f503,plain,(
  earray_131 = rselect_seq(queue_130)),
  inference(definition_unfolding,[],[f104,f466])).
fof(f104,axiom,(
  earray_131 = rselect_seq(q26)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1366,plain,(
  earray_127 = rselect_seq(queue_130)),
  inference(forward_demodulation,[],[f1304,f787])).
fof(f787,plain,(
  earray_127 = rselect_seq(queue_128)),
  inference(superposition,[],[f3,f606])).
fof(f606,plain,(
  queue_128 = rstore_seq(queue_124,earray_127)),
  inference(definition_unfolding,[],[f341,f465])).
fof(f341,axiom,(
  queue_128 = rstore_seq(q25,earray_127)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1304,plain,(
  rselect_seq(queue_130) = rselect_seq(queue_128)),
  inference(superposition,[],[f7,f918])).
fof(f918,plain,(
  queue_130 = rstore_tail(queue_128,index_132)),
  inference(backward_demodulation,[],[f916,f343])).
fof(f86598,plain,(
  elem_352 = select(earray_133,index_57)),
  inference(subsumption_resolution,[],[f86596,f5379])).
fof(f5379,plain,(
  index_141 != index_57),
  inference(forward_demodulation,[],[f5378,f2148])).
fof(f2148,plain,(
  index_141 = s(index_132)),
  inference(backward_demodulation,[],[f2147,f210])).
fof(f210,axiom,(
  index_135 = s(index_132)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2147,plain,(
  index_141 = index_135),
  inference(forward_demodulation,[],[f2146,f555])).
fof(f555,plain,(
  index_141 = rselect_tail(queue_139)),
  inference(definition_unfolding,[],[f213,f467])).
fof(f467,axiom,(
  q27 = queue_139),
  file('SWV568-1.050.tptp',unknown)).
fof(f213,axiom,(
  index_141 = rselect_tail(q27)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2146,plain,(
  index_135 = rselect_tail(queue_139)),
  inference(forward_demodulation,[],[f2057,f856])).
fof(f856,plain,(
  index_135 = rselect_tail(queue_136)),
  inference(superposition,[],[f5,f345])).
fof(f345,axiom,(
  queue_136 = rstore_tail(queue_134,index_135)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2057,plain,(
  rselect_tail(queue_139) = rselect_tail(queue_136)),
  inference(superposition,[],[f11,f346])).
fof(f346,axiom,(
  queue_139 = rstore_head(queue_136,index_138)),
  file('SWV568-1.050.tptp',unknown)).
fof(f5378,plain,(
  index_57 != s(index_132)),
  inference(forward_demodulation,[],[f5377,f919])).
fof(f5377,plain,(
  index_57 != s(s(index_126))),
  inference(forward_demodulation,[],[f5376,f915])).
fof(f5376,plain,(
  index_57 != s(s(s(index_120)))),
  inference(forward_demodulation,[],[f5375,f2139])).
fof(f5375,plain,(
  index_57 != s(s(s(s(index_111))))),
  inference(forward_demodulation,[],[f5374,f907])).
fof(f5374,plain,(
  index_57 != s(s(s(s(s(index_105)))))),
  inference(forward_demodulation,[],[f5373,f903])).
fof(f5373,plain,(
  index_57 != s(s(s(s(s(s(index_99))))))),
  inference(forward_demodulation,[],[f5372,f2517])).
fof(f5372,plain,(
  index_57 != s(s(s(s(s(s(s(index_90)))))))),
  inference(forward_demodulation,[],[f5371,f1033])).
fof(f5371,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_84))))))))),
  inference(forward_demodulation,[],[f5370,f1023])).
fof(f5370,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_72)))))))))),
  inference(forward_demodulation,[],[f5006,f2508])).
fof(f5006,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_63))))))))))),
  inference(superposition,[],[f83,f1019])).
fof(f83,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(X4))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86596,plain,(
  elem_352 = select(earray_133,index_57) | index_141 = index_57),
  inference(superposition,[],[f86595,f4492])).
fof(f4492,plain,(
  ( ! [X8] : (select(earray_133,X8) = select(earray_142,X8) | index_141 = X8) )),
  inference(superposition,[],[f2,f1372])).
fof(f1372,plain,(
  earray_142 = store(earray_133,index_141,e27)),
  inference(backward_demodulation,[],[f1370,f108])).
fof(f108,axiom,(
  earray_142 = store(earray_140,index_141,e27)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1370,plain,(
  earray_133 = earray_140),
  inference(forward_demodulation,[],[f1369,f789])).
fof(f789,plain,(
  earray_133 = rselect_seq(queue_134)),
  inference(superposition,[],[f3,f607])).
fof(f607,plain,(
  queue_134 = rstore_seq(queue_130,earray_133)),
  inference(definition_unfolding,[],[f344,f466])).
fof(f344,axiom,(
  queue_134 = rstore_seq(q26,earray_133)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1369,plain,(
  earray_140 = rselect_seq(queue_134)),
  inference(forward_demodulation,[],[f1305,f1284])).
fof(f1284,plain,(
  earray_140 = rselect_seq(queue_136)),
  inference(forward_demodulation,[],[f1267,f505])).
fof(f505,plain,(
  earray_140 = rselect_seq(queue_139)),
  inference(definition_unfolding,[],[f107,f467])).
fof(f107,axiom,(
  earray_140 = rselect_seq(q27)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1267,plain,(
  rselect_seq(queue_139) = rselect_seq(queue_136)),
  inference(superposition,[],[f6,f346])).
fof(f1305,plain,(
  rselect_seq(queue_134) = rselect_seq(queue_136)),
  inference(superposition,[],[f7,f345])).
fof(f86595,plain,(
  elem_352 = select(earray_142,index_57)),
  inference(subsumption_resolution,[],[f86593,f5882])).
fof(f5882,plain,(
  index_147 != index_57),
  inference(forward_demodulation,[],[f5881,f923])).
fof(f923,plain,(
  index_147 = s(index_141)),
  inference(backward_demodulation,[],[f920,f214])).
fof(f214,axiom,(
  index_144 = s(index_141)),
  file('SWV568-1.050.tptp',unknown)).
fof(f920,plain,(
  index_147 = index_144),
  inference(forward_demodulation,[],[f857,f556])).
fof(f556,plain,(
  index_147 = rselect_tail(queue_145)),
  inference(definition_unfolding,[],[f215,f468])).
fof(f468,axiom,(
  q28 = queue_145),
  file('SWV568-1.050.tptp',unknown)).
fof(f215,axiom,(
  index_147 = rselect_tail(q28)),
  file('SWV568-1.050.tptp',unknown)).
fof(f857,plain,(
  index_144 = rselect_tail(queue_145)),
  inference(superposition,[],[f5,f348])).
fof(f348,axiom,(
  queue_145 = rstore_tail(queue_143,index_144)),
  file('SWV568-1.050.tptp',unknown)).
fof(f5881,plain,(
  index_57 != s(index_141)),
  inference(forward_demodulation,[],[f5880,f2148])).
fof(f5880,plain,(
  index_57 != s(s(index_132))),
  inference(forward_demodulation,[],[f5879,f919])).
fof(f5879,plain,(
  index_57 != s(s(s(index_126)))),
  inference(forward_demodulation,[],[f5878,f915])).
fof(f5878,plain,(
  index_57 != s(s(s(s(index_120))))),
  inference(forward_demodulation,[],[f5877,f2139])).
fof(f5877,plain,(
  index_57 != s(s(s(s(s(index_111)))))),
  inference(forward_demodulation,[],[f5876,f907])).
fof(f5876,plain,(
  index_57 != s(s(s(s(s(s(index_105))))))),
  inference(forward_demodulation,[],[f5875,f903])).
fof(f5875,plain,(
  index_57 != s(s(s(s(s(s(s(index_99)))))))),
  inference(forward_demodulation,[],[f5874,f2517])).
fof(f5874,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_90))))))))),
  inference(forward_demodulation,[],[f5873,f1033])).
fof(f5873,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_84)))))))))),
  inference(forward_demodulation,[],[f5872,f1023])).
fof(f5872,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_72))))))))))),
  inference(forward_demodulation,[],[f5474,f2508])).
fof(f5474,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))),
  inference(superposition,[],[f82,f1019])).
fof(f82,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86593,plain,(
  elem_352 = select(earray_142,index_57) | index_147 = index_57),
  inference(superposition,[],[f86592,f4493])).
fof(f4493,plain,(
  ( ! [X9] : (select(earray_142,X9) = select(earray_148,X9) | index_147 = X9) )),
  inference(superposition,[],[f2,f1376])).
fof(f1376,plain,(
  earray_148 = store(earray_142,index_147,e28)),
  inference(backward_demodulation,[],[f1375,f110])).
fof(f110,axiom,(
  earray_148 = store(earray_146,index_147,e28)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1375,plain,(
  earray_142 = earray_146),
  inference(backward_demodulation,[],[f1374,f506])).
fof(f506,plain,(
  earray_146 = rselect_seq(queue_145)),
  inference(definition_unfolding,[],[f109,f468])).
fof(f109,axiom,(
  earray_146 = rselect_seq(q28)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1374,plain,(
  earray_142 = rselect_seq(queue_145)),
  inference(forward_demodulation,[],[f1306,f790])).
fof(f790,plain,(
  earray_142 = rselect_seq(queue_143)),
  inference(superposition,[],[f3,f608])).
fof(f608,plain,(
  queue_143 = rstore_seq(queue_139,earray_142)),
  inference(definition_unfolding,[],[f347,f467])).
fof(f347,axiom,(
  queue_143 = rstore_seq(q27,earray_142)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1306,plain,(
  rselect_seq(queue_145) = rselect_seq(queue_143)),
  inference(superposition,[],[f7,f922])).
fof(f922,plain,(
  queue_145 = rstore_tail(queue_143,index_147)),
  inference(backward_demodulation,[],[f920,f348])).
fof(f86592,plain,(
  elem_352 = select(earray_148,index_57)),
  inference(subsumption_resolution,[],[f86590,f6425])).
fof(f6425,plain,(
  index_162 != index_57),
  inference(forward_demodulation,[],[f6424,f927])).
fof(f927,plain,(
  index_162 = s(index_147)),
  inference(backward_demodulation,[],[f924,f217])).
fof(f217,axiom,(
  index_150 = s(index_147)),
  file('SWV568-1.050.tptp',unknown)).
fof(f924,plain,(
  index_162 = index_150),
  inference(forward_demodulation,[],[f858,f559])).
fof(f559,plain,(
  index_162 = rselect_tail(queue_151)),
  inference(definition_unfolding,[],[f222,f469])).
fof(f469,axiom,(
  q29 = queue_151),
  file('SWV568-1.050.tptp',unknown)).
fof(f222,axiom,(
  index_162 = rselect_tail(q29)),
  file('SWV568-1.050.tptp',unknown)).
fof(f858,plain,(
  index_150 = rselect_tail(queue_151)),
  inference(superposition,[],[f5,f350])).
fof(f350,axiom,(
  queue_151 = rstore_tail(queue_149,index_150)),
  file('SWV568-1.050.tptp',unknown)).
fof(f6424,plain,(
  index_57 != s(index_147)),
  inference(forward_demodulation,[],[f6423,f923])).
fof(f6423,plain,(
  index_57 != s(s(index_141))),
  inference(forward_demodulation,[],[f6422,f2148])).
fof(f6422,plain,(
  index_57 != s(s(s(index_132)))),
  inference(forward_demodulation,[],[f6421,f919])).
fof(f6421,plain,(
  index_57 != s(s(s(s(index_126))))),
  inference(forward_demodulation,[],[f6420,f915])).
fof(f6420,plain,(
  index_57 != s(s(s(s(s(index_120)))))),
  inference(forward_demodulation,[],[f6419,f2139])).
fof(f6419,plain,(
  index_57 != s(s(s(s(s(s(index_111))))))),
  inference(forward_demodulation,[],[f6418,f907])).
fof(f6418,plain,(
  index_57 != s(s(s(s(s(s(s(index_105)))))))),
  inference(forward_demodulation,[],[f6417,f903])).
fof(f6417,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_99))))))))),
  inference(forward_demodulation,[],[f6416,f2517])).
fof(f6416,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_90)))))))))),
  inference(forward_demodulation,[],[f6415,f1033])).
fof(f6415,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_84))))))))))),
  inference(forward_demodulation,[],[f6414,f1023])).
fof(f6414,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))),
  inference(forward_demodulation,[],[f5982,f2508])).
fof(f5982,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))),
  inference(superposition,[],[f81,f1019])).
fof(f81,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86590,plain,(
  elem_352 = select(earray_148,index_57) | index_162 = index_57),
  inference(superposition,[],[f86589,f4494])).
fof(f4494,plain,(
  ( ! [X10] : (select(earray_148,X10) = select(earray_163,X10) | index_162 = X10) )),
  inference(superposition,[],[f2,f1379])).
fof(f1379,plain,(
  earray_163 = store(earray_148,index_162,e29)),
  inference(backward_demodulation,[],[f1378,f115])).
fof(f115,axiom,(
  earray_163 = store(earray_161,index_162,e29)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1378,plain,(
  earray_148 = earray_161),
  inference(backward_demodulation,[],[f1377,f508])).
fof(f508,plain,(
  earray_161 = rselect_seq(queue_151)),
  inference(definition_unfolding,[],[f114,f469])).
fof(f114,axiom,(
  earray_161 = rselect_seq(q29)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1377,plain,(
  earray_148 = rselect_seq(queue_151)),
  inference(forward_demodulation,[],[f1307,f791])).
fof(f791,plain,(
  earray_148 = rselect_seq(queue_149)),
  inference(superposition,[],[f3,f609])).
fof(f609,plain,(
  queue_149 = rstore_seq(queue_145,earray_148)),
  inference(definition_unfolding,[],[f349,f468])).
fof(f349,axiom,(
  queue_149 = rstore_seq(q28,earray_148)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1307,plain,(
  rselect_seq(queue_151) = rselect_seq(queue_149)),
  inference(superposition,[],[f7,f926])).
fof(f926,plain,(
  queue_151 = rstore_tail(queue_149,index_162)),
  inference(backward_demodulation,[],[f924,f350])).
fof(f86589,plain,(
  elem_352 = select(earray_163,index_57)),
  inference(subsumption_resolution,[],[f86587,f7006])).
fof(f7006,plain,(
  index_171 != index_57),
  inference(forward_demodulation,[],[f7005,f2226])).
fof(f2226,plain,(
  index_171 = s(index_162)),
  inference(backward_demodulation,[],[f2225,f223])).
fof(f223,axiom,(
  index_165 = s(index_162)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2225,plain,(
  index_171 = index_165),
  inference(forward_demodulation,[],[f2224,f560])).
fof(f560,plain,(
  index_171 = rselect_tail(queue_169)),
  inference(definition_unfolding,[],[f226,f471])).
fof(f471,axiom,(
  q30 = queue_169),
  file('SWV568-1.050.tptp',unknown)).
fof(f226,axiom,(
  index_171 = rselect_tail(q30)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2224,plain,(
  index_165 = rselect_tail(queue_169)),
  inference(forward_demodulation,[],[f2059,f860])).
fof(f860,plain,(
  index_165 = rselect_tail(queue_166)),
  inference(superposition,[],[f5,f355])).
fof(f355,axiom,(
  queue_166 = rstore_tail(queue_164,index_165)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2059,plain,(
  rselect_tail(queue_169) = rselect_tail(queue_166)),
  inference(superposition,[],[f11,f356])).
fof(f356,axiom,(
  queue_169 = rstore_head(queue_166,index_168)),
  file('SWV568-1.050.tptp',unknown)).
fof(f7005,plain,(
  index_57 != s(index_162)),
  inference(forward_demodulation,[],[f7004,f927])).
fof(f7004,plain,(
  index_57 != s(s(index_147))),
  inference(forward_demodulation,[],[f7003,f923])).
fof(f7003,plain,(
  index_57 != s(s(s(index_141)))),
  inference(forward_demodulation,[],[f7002,f2148])).
fof(f7002,plain,(
  index_57 != s(s(s(s(index_132))))),
  inference(forward_demodulation,[],[f7001,f919])).
fof(f7001,plain,(
  index_57 != s(s(s(s(s(index_126)))))),
  inference(forward_demodulation,[],[f7000,f915])).
fof(f7000,plain,(
  index_57 != s(s(s(s(s(s(index_120))))))),
  inference(forward_demodulation,[],[f6999,f2139])).
fof(f6999,plain,(
  index_57 != s(s(s(s(s(s(s(index_111)))))))),
  inference(forward_demodulation,[],[f6998,f907])).
fof(f6998,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_105))))))))),
  inference(forward_demodulation,[],[f6997,f903])).
fof(f6997,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_99)))))))))),
  inference(forward_demodulation,[],[f6996,f2517])).
fof(f6996,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_90))))))))))),
  inference(forward_demodulation,[],[f6995,f1033])).
fof(f6995,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))),
  inference(forward_demodulation,[],[f6994,f1023])).
fof(f6994,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))),
  inference(forward_demodulation,[],[f6530,f2508])).
fof(f6530,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))),
  inference(superposition,[],[f80,f1019])).
fof(f80,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86587,plain,(
  elem_352 = select(earray_163,index_57) | index_171 = index_57),
  inference(superposition,[],[f86586,f4498])).
fof(f4498,plain,(
  ( ! [X14] : (select(earray_163,X14) = select(earray_172,X14) | index_171 = X14) )),
  inference(superposition,[],[f2,f1388])).
fof(f1388,plain,(
  earray_172 = store(earray_163,index_171,e30)),
  inference(backward_demodulation,[],[f1386,f117])).
fof(f117,axiom,(
  earray_172 = store(earray_170,index_171,e30)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1386,plain,(
  earray_163 = earray_170),
  inference(forward_demodulation,[],[f1385,f792])).
fof(f792,plain,(
  earray_163 = rselect_seq(queue_164)),
  inference(superposition,[],[f3,f611])).
fof(f611,plain,(
  queue_164 = rstore_seq(queue_151,earray_163)),
  inference(definition_unfolding,[],[f354,f469])).
fof(f354,axiom,(
  queue_164 = rstore_seq(q29,earray_163)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1385,plain,(
  earray_170 = rselect_seq(queue_164)),
  inference(forward_demodulation,[],[f1309,f1286])).
fof(f1286,plain,(
  earray_170 = rselect_seq(queue_166)),
  inference(forward_demodulation,[],[f1269,f509])).
fof(f509,plain,(
  earray_170 = rselect_seq(queue_169)),
  inference(definition_unfolding,[],[f116,f471])).
fof(f116,axiom,(
  earray_170 = rselect_seq(q30)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1269,plain,(
  rselect_seq(queue_169) = rselect_seq(queue_166)),
  inference(superposition,[],[f6,f356])).
fof(f1309,plain,(
  rselect_seq(queue_164) = rselect_seq(queue_166)),
  inference(superposition,[],[f7,f355])).
fof(f86586,plain,(
  elem_352 = select(earray_172,index_57)),
  inference(subsumption_resolution,[],[f86584,f7625])).
fof(f7625,plain,(
  index_177 != index_57),
  inference(forward_demodulation,[],[f7624,f935])).
fof(f935,plain,(
  index_177 = s(index_171)),
  inference(backward_demodulation,[],[f932,f227])).
fof(f227,axiom,(
  index_174 = s(index_171)),
  file('SWV568-1.050.tptp',unknown)).
fof(f932,plain,(
  index_177 = index_174),
  inference(forward_demodulation,[],[f862,f561])).
fof(f561,plain,(
  index_177 = rselect_tail(queue_175)),
  inference(definition_unfolding,[],[f228,f472])).
fof(f472,axiom,(
  q31 = queue_175),
  file('SWV568-1.050.tptp',unknown)).
fof(f228,axiom,(
  index_177 = rselect_tail(q31)),
  file('SWV568-1.050.tptp',unknown)).
fof(f862,plain,(
  index_174 = rselect_tail(queue_175)),
  inference(superposition,[],[f5,f359])).
fof(f359,axiom,(
  queue_175 = rstore_tail(queue_173,index_174)),
  file('SWV568-1.050.tptp',unknown)).
fof(f7624,plain,(
  index_57 != s(index_171)),
  inference(forward_demodulation,[],[f7623,f2226])).
fof(f7623,plain,(
  index_57 != s(s(index_162))),
  inference(forward_demodulation,[],[f7622,f927])).
fof(f7622,plain,(
  index_57 != s(s(s(index_147)))),
  inference(forward_demodulation,[],[f7621,f923])).
fof(f7621,plain,(
  index_57 != s(s(s(s(index_141))))),
  inference(forward_demodulation,[],[f7620,f2148])).
fof(f7620,plain,(
  index_57 != s(s(s(s(s(index_132)))))),
  inference(forward_demodulation,[],[f7619,f919])).
fof(f7619,plain,(
  index_57 != s(s(s(s(s(s(index_126))))))),
  inference(forward_demodulation,[],[f7618,f915])).
fof(f7618,plain,(
  index_57 != s(s(s(s(s(s(s(index_120)))))))),
  inference(forward_demodulation,[],[f7617,f2139])).
fof(f7617,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_111))))))))),
  inference(forward_demodulation,[],[f7616,f907])).
fof(f7616,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_105)))))))))),
  inference(forward_demodulation,[],[f7615,f903])).
fof(f7615,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_99))))))))))),
  inference(forward_demodulation,[],[f7614,f2517])).
fof(f7614,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))),
  inference(forward_demodulation,[],[f7613,f1033])).
fof(f7613,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))),
  inference(forward_demodulation,[],[f7612,f1023])).
fof(f7612,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))),
  inference(forward_demodulation,[],[f7116,f2508])).
fof(f7116,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))),
  inference(superposition,[],[f79,f1019])).
fof(f79,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86584,plain,(
  elem_352 = select(earray_172,index_57) | index_177 = index_57),
  inference(superposition,[],[f86583,f4499])).
fof(f4499,plain,(
  ( ! [X15] : (select(earray_172,X15) = select(earray_178,X15) | index_177 = X15) )),
  inference(superposition,[],[f2,f1395])).
fof(f1395,plain,(
  earray_178 = store(earray_172,index_177,e31)),
  inference(backward_demodulation,[],[f1394,f119])).
fof(f119,axiom,(
  earray_178 = store(earray_176,index_177,e31)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1394,plain,(
  earray_172 = earray_176),
  inference(backward_demodulation,[],[f1393,f510])).
fof(f510,plain,(
  earray_176 = rselect_seq(queue_175)),
  inference(definition_unfolding,[],[f118,f472])).
fof(f118,axiom,(
  earray_176 = rselect_seq(q31)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1393,plain,(
  earray_172 = rselect_seq(queue_175)),
  inference(forward_demodulation,[],[f1311,f794])).
fof(f794,plain,(
  earray_172 = rselect_seq(queue_173)),
  inference(superposition,[],[f3,f613])).
fof(f613,plain,(
  queue_173 = rstore_seq(queue_169,earray_172)),
  inference(definition_unfolding,[],[f358,f471])).
fof(f358,axiom,(
  queue_173 = rstore_seq(q30,earray_172)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1311,plain,(
  rselect_seq(queue_175) = rselect_seq(queue_173)),
  inference(superposition,[],[f7,f934])).
fof(f934,plain,(
  queue_175 = rstore_tail(queue_173,index_177)),
  inference(backward_demodulation,[],[f932,f359])).
fof(f86583,plain,(
  elem_352 = select(earray_178,index_57)),
  inference(subsumption_resolution,[],[f86581,f8281])).
fof(f8281,plain,(
  index_183 != index_57),
  inference(forward_demodulation,[],[f8280,f939])).
fof(f939,plain,(
  index_183 = s(index_177)),
  inference(backward_demodulation,[],[f936,f230])).
fof(f230,axiom,(
  index_180 = s(index_177)),
  file('SWV568-1.050.tptp',unknown)).
fof(f936,plain,(
  index_183 = index_180),
  inference(forward_demodulation,[],[f863,f562])).
fof(f562,plain,(
  index_183 = rselect_tail(queue_181)),
  inference(definition_unfolding,[],[f231,f473])).
fof(f473,axiom,(
  q32 = queue_181),
  file('SWV568-1.050.tptp',unknown)).
fof(f231,axiom,(
  index_183 = rselect_tail(q32)),
  file('SWV568-1.050.tptp',unknown)).
fof(f863,plain,(
  index_180 = rselect_tail(queue_181)),
  inference(superposition,[],[f5,f361])).
fof(f361,axiom,(
  queue_181 = rstore_tail(queue_179,index_180)),
  file('SWV568-1.050.tptp',unknown)).
fof(f8280,plain,(
  index_57 != s(index_177)),
  inference(forward_demodulation,[],[f8279,f935])).
fof(f8279,plain,(
  index_57 != s(s(index_171))),
  inference(forward_demodulation,[],[f8278,f2226])).
fof(f8278,plain,(
  index_57 != s(s(s(index_162)))),
  inference(forward_demodulation,[],[f8277,f927])).
fof(f8277,plain,(
  index_57 != s(s(s(s(index_147))))),
  inference(forward_demodulation,[],[f8276,f923])).
fof(f8276,plain,(
  index_57 != s(s(s(s(s(index_141)))))),
  inference(forward_demodulation,[],[f8275,f2148])).
fof(f8275,plain,(
  index_57 != s(s(s(s(s(s(index_132))))))),
  inference(forward_demodulation,[],[f8274,f919])).
fof(f8274,plain,(
  index_57 != s(s(s(s(s(s(s(index_126)))))))),
  inference(forward_demodulation,[],[f8273,f915])).
fof(f8273,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_120))))))))),
  inference(forward_demodulation,[],[f8272,f2139])).
fof(f8272,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_111)))))))))),
  inference(forward_demodulation,[],[f8271,f907])).
fof(f8271,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_105))))))))))),
  inference(forward_demodulation,[],[f8270,f903])).
fof(f8270,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))),
  inference(forward_demodulation,[],[f8269,f2517])).
fof(f8269,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))),
  inference(forward_demodulation,[],[f8268,f1033])).
fof(f8268,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))),
  inference(forward_demodulation,[],[f8267,f1023])).
fof(f8267,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))),
  inference(forward_demodulation,[],[f7740,f2508])).
fof(f7740,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))),
  inference(superposition,[],[f78,f1019])).
fof(f78,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86581,plain,(
  elem_352 = select(earray_178,index_57) | index_183 = index_57),
  inference(superposition,[],[f86580,f4500])).
fof(f4500,plain,(
  ( ! [X16] : (select(earray_178,X16) = select(earray_184,X16) | index_183 = X16) )),
  inference(superposition,[],[f2,f1398])).
fof(f1398,plain,(
  earray_184 = store(earray_178,index_183,e32)),
  inference(backward_demodulation,[],[f1397,f121])).
fof(f121,axiom,(
  earray_184 = store(earray_182,index_183,e32)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1397,plain,(
  earray_178 = earray_182),
  inference(backward_demodulation,[],[f1396,f511])).
fof(f511,plain,(
  earray_182 = rselect_seq(queue_181)),
  inference(definition_unfolding,[],[f120,f473])).
fof(f120,axiom,(
  earray_182 = rselect_seq(q32)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1396,plain,(
  earray_178 = rselect_seq(queue_181)),
  inference(forward_demodulation,[],[f1312,f795])).
fof(f795,plain,(
  earray_178 = rselect_seq(queue_179)),
  inference(superposition,[],[f3,f614])).
fof(f614,plain,(
  queue_179 = rstore_seq(queue_175,earray_178)),
  inference(definition_unfolding,[],[f360,f472])).
fof(f360,axiom,(
  queue_179 = rstore_seq(q31,earray_178)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1312,plain,(
  rselect_seq(queue_181) = rselect_seq(queue_179)),
  inference(superposition,[],[f7,f938])).
fof(f938,plain,(
  queue_181 = rstore_tail(queue_179,index_183)),
  inference(backward_demodulation,[],[f936,f361])).
fof(f86580,plain,(
  elem_352 = select(earray_184,index_57)),
  inference(subsumption_resolution,[],[f86578,f8973])).
fof(f8973,plain,(
  index_192 != index_57),
  inference(forward_demodulation,[],[f8972,f2235])).
fof(f2235,plain,(
  index_192 = s(index_183)),
  inference(backward_demodulation,[],[f2234,f232])).
fof(f232,axiom,(
  index_186 = s(index_183)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2234,plain,(
  index_192 = index_186),
  inference(forward_demodulation,[],[f2233,f563])).
fof(f563,plain,(
  index_192 = rselect_tail(queue_190)),
  inference(definition_unfolding,[],[f235,f474])).
fof(f474,axiom,(
  q33 = queue_190),
  file('SWV568-1.050.tptp',unknown)).
fof(f235,axiom,(
  index_192 = rselect_tail(q33)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2233,plain,(
  index_186 = rselect_tail(queue_190)),
  inference(forward_demodulation,[],[f2060,f864])).
fof(f864,plain,(
  index_186 = rselect_tail(queue_187)),
  inference(superposition,[],[f5,f363])).
fof(f363,axiom,(
  queue_187 = rstore_tail(queue_185,index_186)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2060,plain,(
  rselect_tail(queue_190) = rselect_tail(queue_187)),
  inference(superposition,[],[f11,f365])).
fof(f365,axiom,(
  queue_190 = rstore_head(queue_187,index_189)),
  file('SWV568-1.050.tptp',unknown)).
fof(f8972,plain,(
  index_57 != s(index_183)),
  inference(forward_demodulation,[],[f8971,f939])).
fof(f8971,plain,(
  index_57 != s(s(index_177))),
  inference(forward_demodulation,[],[f8970,f935])).
fof(f8970,plain,(
  index_57 != s(s(s(index_171)))),
  inference(forward_demodulation,[],[f8969,f2226])).
fof(f8969,plain,(
  index_57 != s(s(s(s(index_162))))),
  inference(forward_demodulation,[],[f8968,f927])).
fof(f8968,plain,(
  index_57 != s(s(s(s(s(index_147)))))),
  inference(forward_demodulation,[],[f8967,f923])).
fof(f8967,plain,(
  index_57 != s(s(s(s(s(s(index_141))))))),
  inference(forward_demodulation,[],[f8966,f2148])).
fof(f8966,plain,(
  index_57 != s(s(s(s(s(s(s(index_132)))))))),
  inference(forward_demodulation,[],[f8965,f919])).
fof(f8965,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_126))))))))),
  inference(forward_demodulation,[],[f8964,f915])).
fof(f8964,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_120)))))))))),
  inference(forward_demodulation,[],[f8963,f2139])).
fof(f8963,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_111))))))))))),
  inference(forward_demodulation,[],[f8962,f907])).
fof(f8962,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))),
  inference(forward_demodulation,[],[f8961,f903])).
fof(f8961,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))),
  inference(forward_demodulation,[],[f8960,f2517])).
fof(f8960,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))),
  inference(forward_demodulation,[],[f8959,f1033])).
fof(f8959,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))),
  inference(forward_demodulation,[],[f8958,f1023])).
fof(f8958,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))),
  inference(forward_demodulation,[],[f8401,f2508])).
fof(f8401,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))),
  inference(superposition,[],[f77,f1019])).
fof(f77,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86578,plain,(
  elem_352 = select(earray_184,index_57) | index_192 = index_57),
  inference(superposition,[],[f86577,f4501])).
fof(f4501,plain,(
  ( ! [X17] : (select(earray_184,X17) = select(earray_193,X17) | index_192 = X17) )),
  inference(superposition,[],[f2,f1402])).
fof(f1402,plain,(
  earray_193 = store(earray_184,index_192,e33)),
  inference(backward_demodulation,[],[f1400,f123])).
fof(f123,axiom,(
  earray_193 = store(earray_191,index_192,e33)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1400,plain,(
  earray_184 = earray_191),
  inference(forward_demodulation,[],[f1399,f796])).
fof(f796,plain,(
  earray_184 = rselect_seq(queue_185)),
  inference(superposition,[],[f3,f615])).
fof(f615,plain,(
  queue_185 = rstore_seq(queue_181,earray_184)),
  inference(definition_unfolding,[],[f362,f473])).
fof(f362,axiom,(
  queue_185 = rstore_seq(q32,earray_184)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1399,plain,(
  earray_191 = rselect_seq(queue_185)),
  inference(forward_demodulation,[],[f1313,f1287])).
fof(f1287,plain,(
  earray_191 = rselect_seq(queue_187)),
  inference(forward_demodulation,[],[f1270,f512])).
fof(f512,plain,(
  earray_191 = rselect_seq(queue_190)),
  inference(definition_unfolding,[],[f122,f474])).
fof(f122,axiom,(
  earray_191 = rselect_seq(q33)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1270,plain,(
  rselect_seq(queue_190) = rselect_seq(queue_187)),
  inference(superposition,[],[f6,f365])).
fof(f1313,plain,(
  rselect_seq(queue_185) = rselect_seq(queue_187)),
  inference(superposition,[],[f7,f363])).
fof(f86577,plain,(
  elem_352 = select(earray_193,index_57)),
  inference(subsumption_resolution,[],[f86575,f9700])).
fof(f9700,plain,(
  index_198 != index_57),
  inference(forward_demodulation,[],[f9699,f943])).
fof(f943,plain,(
  index_198 = s(index_192)),
  inference(backward_demodulation,[],[f940,f236])).
fof(f236,axiom,(
  index_195 = s(index_192)),
  file('SWV568-1.050.tptp',unknown)).
fof(f940,plain,(
  index_198 = index_195),
  inference(forward_demodulation,[],[f865,f564])).
fof(f564,plain,(
  index_198 = rselect_tail(queue_196)),
  inference(definition_unfolding,[],[f237,f475])).
fof(f475,axiom,(
  q34 = queue_196),
  file('SWV568-1.050.tptp',unknown)).
fof(f237,axiom,(
  index_198 = rselect_tail(q34)),
  file('SWV568-1.050.tptp',unknown)).
fof(f865,plain,(
  index_195 = rselect_tail(queue_196)),
  inference(superposition,[],[f5,f367])).
fof(f367,axiom,(
  queue_196 = rstore_tail(queue_194,index_195)),
  file('SWV568-1.050.tptp',unknown)).
fof(f9699,plain,(
  index_57 != s(index_192)),
  inference(forward_demodulation,[],[f9698,f2235])).
fof(f9698,plain,(
  index_57 != s(s(index_183))),
  inference(forward_demodulation,[],[f9697,f939])).
fof(f9697,plain,(
  index_57 != s(s(s(index_177)))),
  inference(forward_demodulation,[],[f9696,f935])).
fof(f9696,plain,(
  index_57 != s(s(s(s(index_171))))),
  inference(forward_demodulation,[],[f9695,f2226])).
fof(f9695,plain,(
  index_57 != s(s(s(s(s(index_162)))))),
  inference(forward_demodulation,[],[f9694,f927])).
fof(f9694,plain,(
  index_57 != s(s(s(s(s(s(index_147))))))),
  inference(forward_demodulation,[],[f9693,f923])).
fof(f9693,plain,(
  index_57 != s(s(s(s(s(s(s(index_141)))))))),
  inference(forward_demodulation,[],[f9692,f2148])).
fof(f9692,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_132))))))))),
  inference(forward_demodulation,[],[f9691,f919])).
fof(f9691,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_126)))))))))),
  inference(forward_demodulation,[],[f9690,f915])).
fof(f9690,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_120))))))))))),
  inference(forward_demodulation,[],[f9689,f2139])).
fof(f9689,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_111)))))))))))),
  inference(forward_demodulation,[],[f9688,f907])).
fof(f9688,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_105))))))))))))),
  inference(forward_demodulation,[],[f9687,f903])).
fof(f9687,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))))),
  inference(forward_demodulation,[],[f9686,f2517])).
fof(f9686,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))))),
  inference(forward_demodulation,[],[f9685,f1033])).
fof(f9685,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))))),
  inference(forward_demodulation,[],[f9684,f1023])).
fof(f9684,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))))),
  inference(forward_demodulation,[],[f9098,f2508])).
fof(f9098,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))))),
  inference(superposition,[],[f76,f1019])).
fof(f76,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86575,plain,(
  elem_352 = select(earray_193,index_57) | index_198 = index_57),
  inference(superposition,[],[f86574,f4502])).
fof(f4502,plain,(
  ( ! [X18] : (select(earray_193,X18) = select(earray_199,X18) | index_198 = X18) )),
  inference(superposition,[],[f2,f1406])).
fof(f1406,plain,(
  earray_199 = store(earray_193,index_198,e34)),
  inference(backward_demodulation,[],[f1405,f125])).
fof(f125,axiom,(
  earray_199 = store(earray_197,index_198,e34)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1405,plain,(
  earray_193 = earray_197),
  inference(backward_demodulation,[],[f1404,f513])).
fof(f513,plain,(
  earray_197 = rselect_seq(queue_196)),
  inference(definition_unfolding,[],[f124,f475])).
fof(f124,axiom,(
  earray_197 = rselect_seq(q34)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1404,plain,(
  earray_193 = rselect_seq(queue_196)),
  inference(forward_demodulation,[],[f1314,f798])).
fof(f798,plain,(
  earray_193 = rselect_seq(queue_194)),
  inference(superposition,[],[f3,f616])).
fof(f616,plain,(
  queue_194 = rstore_seq(queue_190,earray_193)),
  inference(definition_unfolding,[],[f366,f474])).
fof(f366,axiom,(
  queue_194 = rstore_seq(q33,earray_193)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1314,plain,(
  rselect_seq(queue_196) = rselect_seq(queue_194)),
  inference(superposition,[],[f7,f942])).
fof(f942,plain,(
  queue_196 = rstore_tail(queue_194,index_198)),
  inference(backward_demodulation,[],[f940,f367])).
fof(f86574,plain,(
  elem_352 = select(earray_199,index_57)),
  inference(subsumption_resolution,[],[f86572,f10459])).
fof(f10459,plain,(
  index_204 != index_57),
  inference(forward_demodulation,[],[f10458,f947])).
fof(f947,plain,(
  index_204 = s(index_198)),
  inference(backward_demodulation,[],[f944,f238])).
fof(f238,axiom,(
  index_201 = s(index_198)),
  file('SWV568-1.050.tptp',unknown)).
fof(f944,plain,(
  index_204 = index_201),
  inference(forward_demodulation,[],[f866,f565])).
fof(f565,plain,(
  index_204 = rselect_tail(queue_202)),
  inference(definition_unfolding,[],[f239,f476])).
fof(f476,axiom,(
  q35 = queue_202),
  file('SWV568-1.050.tptp',unknown)).
fof(f239,axiom,(
  index_204 = rselect_tail(q35)),
  file('SWV568-1.050.tptp',unknown)).
fof(f866,plain,(
  index_201 = rselect_tail(queue_202)),
  inference(superposition,[],[f5,f369])).
fof(f369,axiom,(
  queue_202 = rstore_tail(queue_200,index_201)),
  file('SWV568-1.050.tptp',unknown)).
fof(f10458,plain,(
  index_57 != s(index_198)),
  inference(forward_demodulation,[],[f10457,f943])).
fof(f10457,plain,(
  index_57 != s(s(index_192))),
  inference(forward_demodulation,[],[f10456,f2235])).
fof(f10456,plain,(
  index_57 != s(s(s(index_183)))),
  inference(forward_demodulation,[],[f10455,f939])).
fof(f10455,plain,(
  index_57 != s(s(s(s(index_177))))),
  inference(forward_demodulation,[],[f10454,f935])).
fof(f10454,plain,(
  index_57 != s(s(s(s(s(index_171)))))),
  inference(forward_demodulation,[],[f10453,f2226])).
fof(f10453,plain,(
  index_57 != s(s(s(s(s(s(index_162))))))),
  inference(forward_demodulation,[],[f10452,f927])).
fof(f10452,plain,(
  index_57 != s(s(s(s(s(s(s(index_147)))))))),
  inference(forward_demodulation,[],[f10451,f923])).
fof(f10451,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_141))))))))),
  inference(forward_demodulation,[],[f10450,f2148])).
fof(f10450,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_132)))))))))),
  inference(forward_demodulation,[],[f10449,f919])).
fof(f10449,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_126))))))))))),
  inference(forward_demodulation,[],[f10448,f915])).
fof(f10448,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_120)))))))))))),
  inference(forward_demodulation,[],[f10447,f2139])).
fof(f10447,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_111))))))))))))),
  inference(forward_demodulation,[],[f10446,f907])).
fof(f10446,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))))),
  inference(forward_demodulation,[],[f10445,f903])).
fof(f10445,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))))),
  inference(forward_demodulation,[],[f10444,f2517])).
fof(f10444,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))))),
  inference(forward_demodulation,[],[f10443,f1033])).
fof(f10443,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))))),
  inference(forward_demodulation,[],[f10442,f1023])).
fof(f10442,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))))),
  inference(forward_demodulation,[],[f9829,f2508])).
fof(f9829,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))))),
  inference(superposition,[],[f75,f1019])).
fof(f75,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86572,plain,(
  elem_352 = select(earray_199,index_57) | index_204 = index_57),
  inference(superposition,[],[f86571,f4503])).
fof(f4503,plain,(
  ( ! [X19] : (select(earray_199,X19) = select(earray_205,X19) | index_204 = X19) )),
  inference(superposition,[],[f2,f1409])).
fof(f1409,plain,(
  earray_205 = store(earray_199,index_204,e35)),
  inference(backward_demodulation,[],[f1408,f129])).
fof(f129,axiom,(
  earray_205 = store(earray_203,index_204,e35)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1408,plain,(
  earray_199 = earray_203),
  inference(backward_demodulation,[],[f1407,f516])).
fof(f516,plain,(
  earray_203 = rselect_seq(queue_202)),
  inference(definition_unfolding,[],[f128,f476])).
fof(f128,axiom,(
  earray_203 = rselect_seq(q35)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1407,plain,(
  earray_199 = rselect_seq(queue_202)),
  inference(forward_demodulation,[],[f1315,f799])).
fof(f799,plain,(
  earray_199 = rselect_seq(queue_200)),
  inference(superposition,[],[f3,f617])).
fof(f617,plain,(
  queue_200 = rstore_seq(queue_196,earray_199)),
  inference(definition_unfolding,[],[f368,f475])).
fof(f368,axiom,(
  queue_200 = rstore_seq(q34,earray_199)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1315,plain,(
  rselect_seq(queue_202) = rselect_seq(queue_200)),
  inference(superposition,[],[f7,f946])).
fof(f946,plain,(
  queue_202 = rstore_tail(queue_200,index_204)),
  inference(backward_demodulation,[],[f944,f369])).
fof(f86571,plain,(
  elem_352 = select(earray_205,index_57)),
  inference(subsumption_resolution,[],[f86569,f11250])).
fof(f11250,plain,(
  index_213 != index_57),
  inference(forward_demodulation,[],[f11249,f2244])).
fof(f2244,plain,(
  index_213 = s(index_204)),
  inference(backward_demodulation,[],[f2243,f240])).
fof(f240,axiom,(
  index_207 = s(index_204)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2243,plain,(
  index_213 = index_207),
  inference(forward_demodulation,[],[f2242,f567])).
fof(f567,plain,(
  index_213 = rselect_tail(queue_211)),
  inference(definition_unfolding,[],[f244,f477])).
fof(f477,axiom,(
  q36 = queue_211),
  file('SWV568-1.050.tptp',unknown)).
fof(f244,axiom,(
  index_213 = rselect_tail(q36)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2242,plain,(
  index_207 = rselect_tail(queue_211)),
  inference(forward_demodulation,[],[f2061,f867])).
fof(f867,plain,(
  index_207 = rselect_tail(queue_208)),
  inference(superposition,[],[f5,f371])).
fof(f371,axiom,(
  queue_208 = rstore_tail(queue_206,index_207)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2061,plain,(
  rselect_tail(queue_211) = rselect_tail(queue_208)),
  inference(superposition,[],[f11,f372])).
fof(f372,axiom,(
  queue_211 = rstore_head(queue_208,index_210)),
  file('SWV568-1.050.tptp',unknown)).
fof(f11249,plain,(
  index_57 != s(index_204)),
  inference(forward_demodulation,[],[f11248,f947])).
fof(f11248,plain,(
  index_57 != s(s(index_198))),
  inference(forward_demodulation,[],[f11247,f943])).
fof(f11247,plain,(
  index_57 != s(s(s(index_192)))),
  inference(forward_demodulation,[],[f11246,f2235])).
fof(f11246,plain,(
  index_57 != s(s(s(s(index_183))))),
  inference(forward_demodulation,[],[f11245,f939])).
fof(f11245,plain,(
  index_57 != s(s(s(s(s(index_177)))))),
  inference(forward_demodulation,[],[f11244,f935])).
fof(f11244,plain,(
  index_57 != s(s(s(s(s(s(index_171))))))),
  inference(forward_demodulation,[],[f11243,f2226])).
fof(f11243,plain,(
  index_57 != s(s(s(s(s(s(s(index_162)))))))),
  inference(forward_demodulation,[],[f11242,f927])).
fof(f11242,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_147))))))))),
  inference(forward_demodulation,[],[f11241,f923])).
fof(f11241,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_141)))))))))),
  inference(forward_demodulation,[],[f11240,f2148])).
fof(f11240,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_132))))))))))),
  inference(forward_demodulation,[],[f11239,f919])).
fof(f11239,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_126)))))))))))),
  inference(forward_demodulation,[],[f11238,f915])).
fof(f11238,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_120))))))))))))),
  inference(forward_demodulation,[],[f11237,f2139])).
fof(f11237,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_111)))))))))))))),
  inference(forward_demodulation,[],[f11236,f907])).
fof(f11236,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105))))))))))))))),
  inference(forward_demodulation,[],[f11235,f903])).
fof(f11235,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))))))),
  inference(forward_demodulation,[],[f11234,f2517])).
fof(f11234,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))))))),
  inference(forward_demodulation,[],[f11233,f1033])).
fof(f11233,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))))))),
  inference(forward_demodulation,[],[f11232,f1023])).
fof(f11232,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))))))),
  inference(forward_demodulation,[],[f10593,f2508])).
fof(f10593,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))))))),
  inference(superposition,[],[f74,f1019])).
fof(f74,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86569,plain,(
  elem_352 = select(earray_205,index_57) | index_213 = index_57),
  inference(superposition,[],[f86568,f4505])).
fof(f4505,plain,(
  ( ! [X21] : (select(earray_205,X21) = select(earray_214,X21) | index_213 = X21) )),
  inference(superposition,[],[f2,f1413])).
fof(f1413,plain,(
  earray_214 = store(earray_205,index_213,e36)),
  inference(backward_demodulation,[],[f1411,f131])).
fof(f131,axiom,(
  earray_214 = store(earray_212,index_213,e36)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1411,plain,(
  earray_205 = earray_212),
  inference(forward_demodulation,[],[f1410,f800])).
fof(f800,plain,(
  earray_205 = rselect_seq(queue_206)),
  inference(superposition,[],[f3,f618])).
fof(f618,plain,(
  queue_206 = rstore_seq(queue_202,earray_205)),
  inference(definition_unfolding,[],[f370,f476])).
fof(f370,axiom,(
  queue_206 = rstore_seq(q35,earray_205)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1410,plain,(
  earray_212 = rselect_seq(queue_206)),
  inference(forward_demodulation,[],[f1316,f1288])).
fof(f1288,plain,(
  earray_212 = rselect_seq(queue_208)),
  inference(forward_demodulation,[],[f1271,f517])).
fof(f517,plain,(
  earray_212 = rselect_seq(queue_211)),
  inference(definition_unfolding,[],[f130,f477])).
fof(f130,axiom,(
  earray_212 = rselect_seq(q36)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1271,plain,(
  rselect_seq(queue_211) = rselect_seq(queue_208)),
  inference(superposition,[],[f6,f372])).
fof(f1316,plain,(
  rselect_seq(queue_206) = rselect_seq(queue_208)),
  inference(superposition,[],[f7,f371])).
fof(f86568,plain,(
  elem_352 = select(earray_214,index_57)),
  inference(subsumption_resolution,[],[f86566,f12074])).
fof(f12074,plain,(
  index_219 != index_57),
  inference(forward_demodulation,[],[f12073,f951])).
fof(f951,plain,(
  index_219 = s(index_213)),
  inference(backward_demodulation,[],[f948,f245])).
fof(f245,axiom,(
  index_216 = s(index_213)),
  file('SWV568-1.050.tptp',unknown)).
fof(f948,plain,(
  index_219 = index_216),
  inference(forward_demodulation,[],[f868,f568])).
fof(f568,plain,(
  index_219 = rselect_tail(queue_217)),
  inference(definition_unfolding,[],[f246,f478])).
fof(f478,axiom,(
  q37 = queue_217),
  file('SWV568-1.050.tptp',unknown)).
fof(f246,axiom,(
  index_219 = rselect_tail(q37)),
  file('SWV568-1.050.tptp',unknown)).
fof(f868,plain,(
  index_216 = rselect_tail(queue_217)),
  inference(superposition,[],[f5,f374])).
fof(f374,axiom,(
  queue_217 = rstore_tail(queue_215,index_216)),
  file('SWV568-1.050.tptp',unknown)).
fof(f12073,plain,(
  index_57 != s(index_213)),
  inference(forward_demodulation,[],[f12072,f2244])).
fof(f12072,plain,(
  index_57 != s(s(index_204))),
  inference(forward_demodulation,[],[f12071,f947])).
fof(f12071,plain,(
  index_57 != s(s(s(index_198)))),
  inference(forward_demodulation,[],[f12070,f943])).
fof(f12070,plain,(
  index_57 != s(s(s(s(index_192))))),
  inference(forward_demodulation,[],[f12069,f2235])).
fof(f12069,plain,(
  index_57 != s(s(s(s(s(index_183)))))),
  inference(forward_demodulation,[],[f12068,f939])).
fof(f12068,plain,(
  index_57 != s(s(s(s(s(s(index_177))))))),
  inference(forward_demodulation,[],[f12067,f935])).
fof(f12067,plain,(
  index_57 != s(s(s(s(s(s(s(index_171)))))))),
  inference(forward_demodulation,[],[f12066,f2226])).
fof(f12066,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_162))))))))),
  inference(forward_demodulation,[],[f12065,f927])).
fof(f12065,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_147)))))))))),
  inference(forward_demodulation,[],[f12064,f923])).
fof(f12064,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_141))))))))))),
  inference(forward_demodulation,[],[f12063,f2148])).
fof(f12063,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_132)))))))))))),
  inference(forward_demodulation,[],[f12062,f919])).
fof(f12062,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_126))))))))))))),
  inference(forward_demodulation,[],[f12061,f915])).
fof(f12061,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_120)))))))))))))),
  inference(forward_demodulation,[],[f12060,f2139])).
fof(f12060,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111))))))))))))))),
  inference(forward_demodulation,[],[f12059,f907])).
fof(f12059,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))))))),
  inference(forward_demodulation,[],[f12058,f903])).
fof(f12058,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))))))),
  inference(forward_demodulation,[],[f12057,f2517])).
fof(f12057,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))))))),
  inference(forward_demodulation,[],[f12056,f1033])).
fof(f12056,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))))))),
  inference(forward_demodulation,[],[f12055,f1023])).
fof(f12055,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))))))),
  inference(forward_demodulation,[],[f11390,f2508])).
fof(f11390,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))))))),
  inference(superposition,[],[f73,f1019])).
fof(f73,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86566,plain,(
  elem_352 = select(earray_214,index_57) | index_219 = index_57),
  inference(superposition,[],[f86565,f4506])).
fof(f4506,plain,(
  ( ! [X22] : (select(earray_214,X22) = select(earray_220,X22) | index_219 = X22) )),
  inference(superposition,[],[f2,f1417])).
fof(f1417,plain,(
  earray_220 = store(earray_214,index_219,e37)),
  inference(backward_demodulation,[],[f1416,f134])).
fof(f134,axiom,(
  earray_220 = store(earray_218,index_219,e37)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1416,plain,(
  earray_214 = earray_218),
  inference(backward_demodulation,[],[f1415,f518])).
fof(f518,plain,(
  earray_218 = rselect_seq(queue_217)),
  inference(definition_unfolding,[],[f132,f478])).
fof(f132,axiom,(
  earray_218 = rselect_seq(q37)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1415,plain,(
  earray_214 = rselect_seq(queue_217)),
  inference(forward_demodulation,[],[f1317,f801])).
fof(f801,plain,(
  earray_214 = rselect_seq(queue_215)),
  inference(superposition,[],[f3,f619])).
fof(f619,plain,(
  queue_215 = rstore_seq(queue_211,earray_214)),
  inference(definition_unfolding,[],[f373,f477])).
fof(f373,axiom,(
  queue_215 = rstore_seq(q36,earray_214)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1317,plain,(
  rselect_seq(queue_217) = rselect_seq(queue_215)),
  inference(superposition,[],[f7,f950])).
fof(f950,plain,(
  queue_217 = rstore_tail(queue_215,index_219)),
  inference(backward_demodulation,[],[f948,f374])).
fof(f86565,plain,(
  elem_352 = select(earray_220,index_57)),
  inference(subsumption_resolution,[],[f86563,f12928])).
fof(f12928,plain,(
  index_225 != index_57),
  inference(forward_demodulation,[],[f12927,f955])).
fof(f955,plain,(
  index_225 = s(index_219)),
  inference(backward_demodulation,[],[f952,f247])).
fof(f247,axiom,(
  index_222 = s(index_219)),
  file('SWV568-1.050.tptp',unknown)).
fof(f952,plain,(
  index_225 = index_222),
  inference(forward_demodulation,[],[f869,f569])).
fof(f569,plain,(
  index_225 = rselect_tail(queue_223)),
  inference(definition_unfolding,[],[f248,f479])).
fof(f479,axiom,(
  q38 = queue_223),
  file('SWV568-1.050.tptp',unknown)).
fof(f248,axiom,(
  index_225 = rselect_tail(q38)),
  file('SWV568-1.050.tptp',unknown)).
fof(f869,plain,(
  index_222 = rselect_tail(queue_223)),
  inference(superposition,[],[f5,f376])).
fof(f376,axiom,(
  queue_223 = rstore_tail(queue_221,index_222)),
  file('SWV568-1.050.tptp',unknown)).
fof(f12927,plain,(
  index_57 != s(index_219)),
  inference(forward_demodulation,[],[f12926,f951])).
fof(f12926,plain,(
  index_57 != s(s(index_213))),
  inference(forward_demodulation,[],[f12925,f2244])).
fof(f12925,plain,(
  index_57 != s(s(s(index_204)))),
  inference(forward_demodulation,[],[f12924,f947])).
fof(f12924,plain,(
  index_57 != s(s(s(s(index_198))))),
  inference(forward_demodulation,[],[f12923,f943])).
fof(f12923,plain,(
  index_57 != s(s(s(s(s(index_192)))))),
  inference(forward_demodulation,[],[f12922,f2235])).
fof(f12922,plain,(
  index_57 != s(s(s(s(s(s(index_183))))))),
  inference(forward_demodulation,[],[f12921,f939])).
fof(f12921,plain,(
  index_57 != s(s(s(s(s(s(s(index_177)))))))),
  inference(forward_demodulation,[],[f12920,f935])).
fof(f12920,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_171))))))))),
  inference(forward_demodulation,[],[f12919,f2226])).
fof(f12919,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_162)))))))))),
  inference(forward_demodulation,[],[f12918,f927])).
fof(f12918,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_147))))))))))),
  inference(forward_demodulation,[],[f12917,f923])).
fof(f12917,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_141)))))))))))),
  inference(forward_demodulation,[],[f12916,f2148])).
fof(f12916,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_132))))))))))))),
  inference(forward_demodulation,[],[f12915,f919])).
fof(f12915,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_126)))))))))))))),
  inference(forward_demodulation,[],[f12914,f915])).
fof(f12914,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120))))))))))))))),
  inference(forward_demodulation,[],[f12913,f2139])).
fof(f12913,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111)))))))))))))))),
  inference(forward_demodulation,[],[f12912,f907])).
fof(f12912,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105))))))))))))))))),
  inference(forward_demodulation,[],[f12911,f903])).
fof(f12911,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))))))))),
  inference(forward_demodulation,[],[f12910,f2517])).
fof(f12910,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))))))))),
  inference(forward_demodulation,[],[f12909,f1033])).
fof(f12909,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))))))))),
  inference(forward_demodulation,[],[f12908,f1023])).
fof(f12908,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))))))))),
  inference(forward_demodulation,[],[f12219,f2508])).
fof(f12219,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))))))))),
  inference(superposition,[],[f72,f1019])).
fof(f72,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86563,plain,(
  elem_352 = select(earray_220,index_57) | index_225 = index_57),
  inference(superposition,[],[f86554,f4508])).
fof(f4508,plain,(
  ( ! [X24] : (select(earray_220,X24) = select(earray_226,X24) | index_225 = X24) )),
  inference(superposition,[],[f2,f1420])).
fof(f1420,plain,(
  earray_226 = store(earray_220,index_225,e38)),
  inference(backward_demodulation,[],[f1419,f136])).
fof(f136,axiom,(
  earray_226 = store(earray_224,index_225,e38)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1419,plain,(
  earray_220 = earray_224),
  inference(backward_demodulation,[],[f1418,f519])).
fof(f519,plain,(
  earray_224 = rselect_seq(queue_223)),
  inference(definition_unfolding,[],[f135,f479])).
fof(f135,axiom,(
  earray_224 = rselect_seq(q38)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1418,plain,(
  earray_220 = rselect_seq(queue_223)),
  inference(forward_demodulation,[],[f1318,f802])).
fof(f802,plain,(
  earray_220 = rselect_seq(queue_221)),
  inference(superposition,[],[f3,f620])).
fof(f620,plain,(
  queue_221 = rstore_seq(queue_217,earray_220)),
  inference(definition_unfolding,[],[f375,f478])).
fof(f375,axiom,(
  queue_221 = rstore_seq(q37,earray_220)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1318,plain,(
  rselect_seq(queue_223) = rselect_seq(queue_221)),
  inference(superposition,[],[f7,f954])).
fof(f954,plain,(
  queue_223 = rstore_tail(queue_221,index_225)),
  inference(backward_demodulation,[],[f952,f376])).
fof(f86554,plain,(
  elem_352 = select(earray_226,index_57)),
  inference(subsumption_resolution,[],[f86552,f13811])).
fof(f13811,plain,(
  index_240 != index_57),
  inference(forward_demodulation,[],[f13810,f2253])).
fof(f2253,plain,(
  index_240 = s(index_225)),
  inference(backward_demodulation,[],[f2252,f249])).
fof(f249,axiom,(
  index_228 = s(index_225)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2252,plain,(
  index_240 = index_228),
  inference(forward_demodulation,[],[f2251,f571])).
fof(f571,plain,(
  index_240 = rselect_tail(queue_232)),
  inference(definition_unfolding,[],[f255,f480])).
fof(f480,axiom,(
  q39 = queue_232),
  file('SWV568-1.050.tptp',unknown)).
fof(f255,axiom,(
  index_240 = rselect_tail(q39)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2251,plain,(
  index_228 = rselect_tail(queue_232)),
  inference(forward_demodulation,[],[f2062,f870])).
fof(f870,plain,(
  index_228 = rselect_tail(queue_229)),
  inference(superposition,[],[f5,f378])).
fof(f378,axiom,(
  queue_229 = rstore_tail(queue_227,index_228)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2062,plain,(
  rselect_tail(queue_232) = rselect_tail(queue_229)),
  inference(superposition,[],[f11,f380])).
fof(f380,axiom,(
  queue_232 = rstore_head(queue_229,index_231)),
  file('SWV568-1.050.tptp',unknown)).
fof(f13810,plain,(
  index_57 != s(index_225)),
  inference(forward_demodulation,[],[f13809,f955])).
fof(f13809,plain,(
  index_57 != s(s(index_219))),
  inference(forward_demodulation,[],[f13808,f951])).
fof(f13808,plain,(
  index_57 != s(s(s(index_213)))),
  inference(forward_demodulation,[],[f13807,f2244])).
fof(f13807,plain,(
  index_57 != s(s(s(s(index_204))))),
  inference(forward_demodulation,[],[f13806,f947])).
fof(f13806,plain,(
  index_57 != s(s(s(s(s(index_198)))))),
  inference(forward_demodulation,[],[f13805,f943])).
fof(f13805,plain,(
  index_57 != s(s(s(s(s(s(index_192))))))),
  inference(forward_demodulation,[],[f13804,f2235])).
fof(f13804,plain,(
  index_57 != s(s(s(s(s(s(s(index_183)))))))),
  inference(forward_demodulation,[],[f13803,f939])).
fof(f13803,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_177))))))))),
  inference(forward_demodulation,[],[f13802,f935])).
fof(f13802,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_171)))))))))),
  inference(forward_demodulation,[],[f13801,f2226])).
fof(f13801,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_162))))))))))),
  inference(forward_demodulation,[],[f13800,f927])).
fof(f13800,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_147)))))))))))),
  inference(forward_demodulation,[],[f13799,f923])).
fof(f13799,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_141))))))))))))),
  inference(forward_demodulation,[],[f13798,f2148])).
fof(f13798,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_132)))))))))))))),
  inference(forward_demodulation,[],[f13797,f919])).
fof(f13797,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126))))))))))))))),
  inference(forward_demodulation,[],[f13796,f915])).
fof(f13796,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120)))))))))))))))),
  inference(forward_demodulation,[],[f13795,f2139])).
fof(f13795,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111))))))))))))))))),
  inference(forward_demodulation,[],[f13794,f907])).
fof(f13794,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))))))))),
  inference(forward_demodulation,[],[f13793,f903])).
fof(f13793,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))))))))),
  inference(forward_demodulation,[],[f13792,f2517])).
fof(f13792,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))))))))),
  inference(forward_demodulation,[],[f13791,f1033])).
fof(f13791,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))))))))),
  inference(forward_demodulation,[],[f13790,f1023])).
fof(f13790,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))))))))),
  inference(forward_demodulation,[],[f13078,f2508])).
fof(f13078,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))))))))),
  inference(superposition,[],[f71,f1019])).
fof(f71,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86552,plain,(
  elem_352 = select(earray_226,index_57) | index_240 = index_57),
  inference(superposition,[],[f86549,f4509])).
fof(f4509,plain,(
  ( ! [X25] : (select(earray_226,X25) = select(earray_241,X25) | index_240 = X25) )),
  inference(superposition,[],[f2,f1424])).
fof(f1424,plain,(
  earray_241 = store(earray_226,index_240,e39)),
  inference(backward_demodulation,[],[f1422,f140])).
fof(f140,axiom,(
  earray_241 = store(earray_239,index_240,e39)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1422,plain,(
  earray_226 = earray_239),
  inference(forward_demodulation,[],[f1421,f803])).
fof(f803,plain,(
  earray_226 = rselect_seq(queue_227)),
  inference(superposition,[],[f3,f621])).
fof(f621,plain,(
  queue_227 = rstore_seq(queue_223,earray_226)),
  inference(definition_unfolding,[],[f377,f479])).
fof(f377,axiom,(
  queue_227 = rstore_seq(q38,earray_226)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1421,plain,(
  earray_239 = rselect_seq(queue_227)),
  inference(forward_demodulation,[],[f1319,f1289])).
fof(f1289,plain,(
  earray_239 = rselect_seq(queue_229)),
  inference(forward_demodulation,[],[f1272,f521])).
fof(f521,plain,(
  earray_239 = rselect_seq(queue_232)),
  inference(definition_unfolding,[],[f139,f480])).
fof(f139,axiom,(
  earray_239 = rselect_seq(q39)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1272,plain,(
  rselect_seq(queue_232) = rselect_seq(queue_229)),
  inference(superposition,[],[f6,f380])).
fof(f1319,plain,(
  rselect_seq(queue_227) = rselect_seq(queue_229)),
  inference(superposition,[],[f7,f378])).
fof(f86549,plain,(
  elem_352 = select(earray_241,index_57)),
  inference(subsumption_resolution,[],[f86547,f14722])).
fof(f14722,plain,(
  index_246 != index_57),
  inference(forward_demodulation,[],[f14721,f963])).
fof(f963,plain,(
  index_246 = s(index_240)),
  inference(backward_demodulation,[],[f960,f256])).
fof(f256,axiom,(
  index_243 = s(index_240)),
  file('SWV568-1.050.tptp',unknown)).
fof(f960,plain,(
  index_246 = index_243),
  inference(forward_demodulation,[],[f873,f572])).
fof(f572,plain,(
  index_246 = rselect_tail(queue_244)),
  inference(definition_unfolding,[],[f257,f482])).
fof(f482,axiom,(
  q40 = queue_244),
  file('SWV568-1.050.tptp',unknown)).
fof(f257,axiom,(
  index_246 = rselect_tail(q40)),
  file('SWV568-1.050.tptp',unknown)).
fof(f873,plain,(
  index_243 = rselect_tail(queue_244)),
  inference(superposition,[],[f5,f384])).
fof(f384,axiom,(
  queue_244 = rstore_tail(queue_242,index_243)),
  file('SWV568-1.050.tptp',unknown)).
fof(f14721,plain,(
  index_57 != s(index_240)),
  inference(forward_demodulation,[],[f14720,f2253])).
fof(f14720,plain,(
  index_57 != s(s(index_225))),
  inference(forward_demodulation,[],[f14719,f955])).
fof(f14719,plain,(
  index_57 != s(s(s(index_219)))),
  inference(forward_demodulation,[],[f14718,f951])).
fof(f14718,plain,(
  index_57 != s(s(s(s(index_213))))),
  inference(forward_demodulation,[],[f14717,f2244])).
fof(f14717,plain,(
  index_57 != s(s(s(s(s(index_204)))))),
  inference(forward_demodulation,[],[f14716,f947])).
fof(f14716,plain,(
  index_57 != s(s(s(s(s(s(index_198))))))),
  inference(forward_demodulation,[],[f14715,f943])).
fof(f14715,plain,(
  index_57 != s(s(s(s(s(s(s(index_192)))))))),
  inference(forward_demodulation,[],[f14714,f2235])).
fof(f14714,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_183))))))))),
  inference(forward_demodulation,[],[f14713,f939])).
fof(f14713,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_177)))))))))),
  inference(forward_demodulation,[],[f14712,f935])).
fof(f14712,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_171))))))))))),
  inference(forward_demodulation,[],[f14711,f2226])).
fof(f14711,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_162)))))))))))),
  inference(forward_demodulation,[],[f14710,f927])).
fof(f14710,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_147))))))))))))),
  inference(forward_demodulation,[],[f14709,f923])).
fof(f14709,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_141)))))))))))))),
  inference(forward_demodulation,[],[f14708,f2148])).
fof(f14708,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132))))))))))))))),
  inference(forward_demodulation,[],[f14707,f919])).
fof(f14707,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126)))))))))))))))),
  inference(forward_demodulation,[],[f14706,f915])).
fof(f14706,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120))))))))))))))))),
  inference(forward_demodulation,[],[f14705,f2139])).
fof(f14705,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111)))))))))))))))))),
  inference(forward_demodulation,[],[f14704,f907])).
fof(f14704,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105))))))))))))))))))),
  inference(forward_demodulation,[],[f14703,f903])).
fof(f14703,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))))))))))),
  inference(forward_demodulation,[],[f14702,f2517])).
fof(f14702,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))))))))))),
  inference(forward_demodulation,[],[f14701,f1033])).
fof(f14701,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))))))))))),
  inference(forward_demodulation,[],[f14700,f1023])).
fof(f14700,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))))))))))),
  inference(forward_demodulation,[],[f13966,f2508])).
fof(f13966,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))))))))))),
  inference(superposition,[],[f70,f1019])).
fof(f70,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86547,plain,(
  elem_352 = select(earray_241,index_57) | index_246 = index_57),
  inference(superposition,[],[f86544,f4511])).
fof(f4511,plain,(
  ( ! [X27] : (select(earray_241,X27) = select(earray_247,X27) | index_246 = X27) )),
  inference(superposition,[],[f2,f1436])).
fof(f1436,plain,(
  earray_247 = store(earray_241,index_246,e40)),
  inference(backward_demodulation,[],[f1435,f142])).
fof(f142,axiom,(
  earray_247 = store(earray_245,index_246,e40)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1435,plain,(
  earray_241 = earray_245),
  inference(backward_demodulation,[],[f1434,f522])).
fof(f522,plain,(
  earray_245 = rselect_seq(queue_244)),
  inference(definition_unfolding,[],[f141,f482])).
fof(f141,axiom,(
  earray_245 = rselect_seq(q40)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1434,plain,(
  earray_241 = rselect_seq(queue_244)),
  inference(forward_demodulation,[],[f1322,f804])).
fof(f804,plain,(
  earray_241 = rselect_seq(queue_242)),
  inference(superposition,[],[f3,f624])).
fof(f624,plain,(
  queue_242 = rstore_seq(queue_232,earray_241)),
  inference(definition_unfolding,[],[f383,f480])).
fof(f383,axiom,(
  queue_242 = rstore_seq(q39,earray_241)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1322,plain,(
  rselect_seq(queue_244) = rselect_seq(queue_242)),
  inference(superposition,[],[f7,f962])).
fof(f962,plain,(
  queue_244 = rstore_tail(queue_242,index_246)),
  inference(backward_demodulation,[],[f960,f384])).
fof(f86544,plain,(
  elem_352 = select(earray_247,index_57)),
  inference(subsumption_resolution,[],[f86542,f15660])).
fof(f15660,plain,(
  index_252 != index_57),
  inference(forward_demodulation,[],[f15659,f967])).
fof(f967,plain,(
  index_252 = s(index_246)),
  inference(backward_demodulation,[],[f964,f258])).
fof(f258,axiom,(
  index_249 = s(index_246)),
  file('SWV568-1.050.tptp',unknown)).
fof(f964,plain,(
  index_252 = index_249),
  inference(forward_demodulation,[],[f874,f573])).
fof(f573,plain,(
  index_252 = rselect_tail(queue_250)),
  inference(definition_unfolding,[],[f259,f483])).
fof(f483,axiom,(
  q41 = queue_250),
  file('SWV568-1.050.tptp',unknown)).
fof(f259,axiom,(
  index_252 = rselect_tail(q41)),
  file('SWV568-1.050.tptp',unknown)).
fof(f874,plain,(
  index_249 = rselect_tail(queue_250)),
  inference(superposition,[],[f5,f387])).
fof(f387,axiom,(
  queue_250 = rstore_tail(queue_248,index_249)),
  file('SWV568-1.050.tptp',unknown)).
fof(f15659,plain,(
  index_57 != s(index_246)),
  inference(forward_demodulation,[],[f15658,f963])).
fof(f15658,plain,(
  index_57 != s(s(index_240))),
  inference(forward_demodulation,[],[f15657,f2253])).
fof(f15657,plain,(
  index_57 != s(s(s(index_225)))),
  inference(forward_demodulation,[],[f15656,f955])).
fof(f15656,plain,(
  index_57 != s(s(s(s(index_219))))),
  inference(forward_demodulation,[],[f15655,f951])).
fof(f15655,plain,(
  index_57 != s(s(s(s(s(index_213)))))),
  inference(forward_demodulation,[],[f15654,f2244])).
fof(f15654,plain,(
  index_57 != s(s(s(s(s(s(index_204))))))),
  inference(forward_demodulation,[],[f15653,f947])).
fof(f15653,plain,(
  index_57 != s(s(s(s(s(s(s(index_198)))))))),
  inference(forward_demodulation,[],[f15652,f943])).
fof(f15652,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_192))))))))),
  inference(forward_demodulation,[],[f15651,f2235])).
fof(f15651,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_183)))))))))),
  inference(forward_demodulation,[],[f15650,f939])).
fof(f15650,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_177))))))))))),
  inference(forward_demodulation,[],[f15649,f935])).
fof(f15649,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_171)))))))))))),
  inference(forward_demodulation,[],[f15648,f2226])).
fof(f15648,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_162))))))))))))),
  inference(forward_demodulation,[],[f15647,f927])).
fof(f15647,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_147)))))))))))))),
  inference(forward_demodulation,[],[f15646,f923])).
fof(f15646,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141))))))))))))))),
  inference(forward_demodulation,[],[f15645,f2148])).
fof(f15645,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132)))))))))))))))),
  inference(forward_demodulation,[],[f15644,f919])).
fof(f15644,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126))))))))))))))))),
  inference(forward_demodulation,[],[f15643,f915])).
fof(f15643,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120)))))))))))))))))),
  inference(forward_demodulation,[],[f15642,f2139])).
fof(f15642,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111))))))))))))))))))),
  inference(forward_demodulation,[],[f15641,f907])).
fof(f15641,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))))))))))),
  inference(forward_demodulation,[],[f15640,f903])).
fof(f15640,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))))))))))),
  inference(forward_demodulation,[],[f15639,f2517])).
fof(f15639,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))))))))))),
  inference(forward_demodulation,[],[f15638,f1033])).
fof(f15638,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))))))))))),
  inference(forward_demodulation,[],[f15637,f1023])).
fof(f15637,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f14882,f2508])).
fof(f14882,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))))))))))),
  inference(superposition,[],[f69,f1019])).
fof(f69,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86542,plain,(
  elem_352 = select(earray_247,index_57) | index_252 = index_57),
  inference(superposition,[],[f86539,f4512])).
fof(f4512,plain,(
  ( ! [X28] : (select(earray_247,X28) = select(earray_253,X28) | index_252 = X28) )),
  inference(superposition,[],[f2,f1439])).
fof(f1439,plain,(
  earray_253 = store(earray_247,index_252,e41)),
  inference(backward_demodulation,[],[f1438,f144])).
fof(f144,axiom,(
  earray_253 = store(earray_251,index_252,e41)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1438,plain,(
  earray_247 = earray_251),
  inference(backward_demodulation,[],[f1437,f523])).
fof(f523,plain,(
  earray_251 = rselect_seq(queue_250)),
  inference(definition_unfolding,[],[f143,f483])).
fof(f143,axiom,(
  earray_251 = rselect_seq(q41)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1437,plain,(
  earray_247 = rselect_seq(queue_250)),
  inference(forward_demodulation,[],[f1323,f806])).
fof(f806,plain,(
  earray_247 = rselect_seq(queue_248)),
  inference(superposition,[],[f3,f625])).
fof(f625,plain,(
  queue_248 = rstore_seq(queue_244,earray_247)),
  inference(definition_unfolding,[],[f385,f482])).
fof(f385,axiom,(
  queue_248 = rstore_seq(q40,earray_247)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1323,plain,(
  rselect_seq(queue_250) = rselect_seq(queue_248)),
  inference(superposition,[],[f7,f966])).
fof(f966,plain,(
  queue_250 = rstore_tail(queue_248,index_252)),
  inference(backward_demodulation,[],[f964,f387])).
fof(f86539,plain,(
  elem_352 = select(earray_253,index_57)),
  inference(subsumption_resolution,[],[f86537,f16623])).
fof(f16623,plain,(
  index_261 != index_57),
  inference(forward_demodulation,[],[f16622,f2262])).
fof(f2262,plain,(
  index_261 = s(index_252)),
  inference(backward_demodulation,[],[f2261,f260])).
fof(f260,axiom,(
  index_255 = s(index_252)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2261,plain,(
  index_261 = index_255),
  inference(forward_demodulation,[],[f2260,f574])).
fof(f574,plain,(
  index_261 = rselect_tail(queue_259)),
  inference(definition_unfolding,[],[f264,f484])).
fof(f484,axiom,(
  q42 = queue_259),
  file('SWV568-1.050.tptp',unknown)).
fof(f264,axiom,(
  index_261 = rselect_tail(q42)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2260,plain,(
  index_255 = rselect_tail(queue_259)),
  inference(forward_demodulation,[],[f2063,f875])).
fof(f875,plain,(
  index_255 = rselect_tail(queue_256)),
  inference(superposition,[],[f5,f389])).
fof(f389,axiom,(
  queue_256 = rstore_tail(queue_254,index_255)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2063,plain,(
  rselect_tail(queue_259) = rselect_tail(queue_256)),
  inference(superposition,[],[f11,f390])).
fof(f390,axiom,(
  queue_259 = rstore_head(queue_256,index_258)),
  file('SWV568-1.050.tptp',unknown)).
fof(f16622,plain,(
  index_57 != s(index_252)),
  inference(forward_demodulation,[],[f16621,f967])).
fof(f16621,plain,(
  index_57 != s(s(index_246))),
  inference(forward_demodulation,[],[f16620,f963])).
fof(f16620,plain,(
  index_57 != s(s(s(index_240)))),
  inference(forward_demodulation,[],[f16619,f2253])).
fof(f16619,plain,(
  index_57 != s(s(s(s(index_225))))),
  inference(forward_demodulation,[],[f16618,f955])).
fof(f16618,plain,(
  index_57 != s(s(s(s(s(index_219)))))),
  inference(forward_demodulation,[],[f16617,f951])).
fof(f16617,plain,(
  index_57 != s(s(s(s(s(s(index_213))))))),
  inference(forward_demodulation,[],[f16616,f2244])).
fof(f16616,plain,(
  index_57 != s(s(s(s(s(s(s(index_204)))))))),
  inference(forward_demodulation,[],[f16615,f947])).
fof(f16615,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_198))))))))),
  inference(forward_demodulation,[],[f16614,f943])).
fof(f16614,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_192)))))))))),
  inference(forward_demodulation,[],[f16613,f2235])).
fof(f16613,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_183))))))))))),
  inference(forward_demodulation,[],[f16612,f939])).
fof(f16612,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_177)))))))))))),
  inference(forward_demodulation,[],[f16611,f935])).
fof(f16611,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_171))))))))))))),
  inference(forward_demodulation,[],[f16610,f2226])).
fof(f16610,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_162)))))))))))))),
  inference(forward_demodulation,[],[f16609,f927])).
fof(f16609,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_147))))))))))))))),
  inference(forward_demodulation,[],[f16608,f923])).
fof(f16608,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141)))))))))))))))),
  inference(forward_demodulation,[],[f16607,f2148])).
fof(f16607,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132))))))))))))))))),
  inference(forward_demodulation,[],[f16606,f919])).
fof(f16606,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126)))))))))))))))))),
  inference(forward_demodulation,[],[f16605,f915])).
fof(f16605,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120))))))))))))))))))),
  inference(forward_demodulation,[],[f16604,f2139])).
fof(f16604,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111)))))))))))))))))))),
  inference(forward_demodulation,[],[f16603,f907])).
fof(f16603,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105))))))))))))))))))))),
  inference(forward_demodulation,[],[f16602,f903])).
fof(f16602,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))))))))))))),
  inference(forward_demodulation,[],[f16601,f2517])).
fof(f16601,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))))))))))))),
  inference(forward_demodulation,[],[f16600,f1033])).
fof(f16600,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f16599,f1023])).
fof(f16599,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f15825,f2508])).
fof(f15825,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))))))))))))),
  inference(superposition,[],[f68,f1019])).
fof(f68,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86537,plain,(
  elem_352 = select(earray_253,index_57) | index_261 = index_57),
  inference(superposition,[],[f86534,f4513])).
fof(f4513,plain,(
  ( ! [X29] : (select(earray_253,X29) = select(earray_262,X29) | index_261 = X29) )),
  inference(superposition,[],[f2,f1443])).
fof(f1443,plain,(
  earray_262 = store(earray_253,index_261,e42)),
  inference(backward_demodulation,[],[f1441,f146])).
fof(f146,axiom,(
  earray_262 = store(earray_260,index_261,e42)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1441,plain,(
  earray_253 = earray_260),
  inference(forward_demodulation,[],[f1440,f807])).
fof(f807,plain,(
  earray_253 = rselect_seq(queue_254)),
  inference(superposition,[],[f3,f626])).
fof(f626,plain,(
  queue_254 = rstore_seq(queue_250,earray_253)),
  inference(definition_unfolding,[],[f388,f483])).
fof(f388,axiom,(
  queue_254 = rstore_seq(q41,earray_253)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1440,plain,(
  earray_260 = rselect_seq(queue_254)),
  inference(forward_demodulation,[],[f1324,f1290])).
fof(f1290,plain,(
  earray_260 = rselect_seq(queue_256)),
  inference(forward_demodulation,[],[f1273,f524])).
fof(f524,plain,(
  earray_260 = rselect_seq(queue_259)),
  inference(definition_unfolding,[],[f145,f484])).
fof(f145,axiom,(
  earray_260 = rselect_seq(q42)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1273,plain,(
  rselect_seq(queue_259) = rselect_seq(queue_256)),
  inference(superposition,[],[f6,f390])).
fof(f1324,plain,(
  rselect_seq(queue_254) = rselect_seq(queue_256)),
  inference(superposition,[],[f7,f389])).
fof(f86534,plain,(
  elem_352 = select(earray_262,index_57)),
  inference(subsumption_resolution,[],[f86532,f17610])).
fof(f17610,plain,(
  index_267 != index_57),
  inference(forward_demodulation,[],[f17609,f971])).
fof(f971,plain,(
  index_267 = s(index_261)),
  inference(backward_demodulation,[],[f968,f265])).
fof(f265,axiom,(
  index_264 = s(index_261)),
  file('SWV568-1.050.tptp',unknown)).
fof(f968,plain,(
  index_267 = index_264),
  inference(forward_demodulation,[],[f876,f575])).
fof(f575,plain,(
  index_267 = rselect_tail(queue_265)),
  inference(definition_unfolding,[],[f266,f485])).
fof(f485,axiom,(
  q43 = queue_265),
  file('SWV568-1.050.tptp',unknown)).
fof(f266,axiom,(
  index_267 = rselect_tail(q43)),
  file('SWV568-1.050.tptp',unknown)).
fof(f876,plain,(
  index_264 = rselect_tail(queue_265)),
  inference(superposition,[],[f5,f392])).
fof(f392,axiom,(
  queue_265 = rstore_tail(queue_263,index_264)),
  file('SWV568-1.050.tptp',unknown)).
fof(f17609,plain,(
  index_57 != s(index_261)),
  inference(forward_demodulation,[],[f17608,f2262])).
fof(f17608,plain,(
  index_57 != s(s(index_252))),
  inference(forward_demodulation,[],[f17607,f967])).
fof(f17607,plain,(
  index_57 != s(s(s(index_246)))),
  inference(forward_demodulation,[],[f17606,f963])).
fof(f17606,plain,(
  index_57 != s(s(s(s(index_240))))),
  inference(forward_demodulation,[],[f17605,f2253])).
fof(f17605,plain,(
  index_57 != s(s(s(s(s(index_225)))))),
  inference(forward_demodulation,[],[f17604,f955])).
fof(f17604,plain,(
  index_57 != s(s(s(s(s(s(index_219))))))),
  inference(forward_demodulation,[],[f17603,f951])).
fof(f17603,plain,(
  index_57 != s(s(s(s(s(s(s(index_213)))))))),
  inference(forward_demodulation,[],[f17602,f2244])).
fof(f17602,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_204))))))))),
  inference(forward_demodulation,[],[f17601,f947])).
fof(f17601,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_198)))))))))),
  inference(forward_demodulation,[],[f17600,f943])).
fof(f17600,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_192))))))))))),
  inference(forward_demodulation,[],[f17599,f2235])).
fof(f17599,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_183)))))))))))),
  inference(forward_demodulation,[],[f17598,f939])).
fof(f17598,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_177))))))))))))),
  inference(forward_demodulation,[],[f17597,f935])).
fof(f17597,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_171)))))))))))))),
  inference(forward_demodulation,[],[f17596,f2226])).
fof(f17596,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_162))))))))))))))),
  inference(forward_demodulation,[],[f17595,f927])).
fof(f17595,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_147)))))))))))))))),
  inference(forward_demodulation,[],[f17594,f923])).
fof(f17594,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141))))))))))))))))),
  inference(forward_demodulation,[],[f17593,f2148])).
fof(f17593,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132)))))))))))))))))),
  inference(forward_demodulation,[],[f17592,f919])).
fof(f17592,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126))))))))))))))))))),
  inference(forward_demodulation,[],[f17591,f915])).
fof(f17591,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120)))))))))))))))))))),
  inference(forward_demodulation,[],[f17590,f2139])).
fof(f17590,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111))))))))))))))))))))),
  inference(forward_demodulation,[],[f17589,f907])).
fof(f17589,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))))))))))))),
  inference(forward_demodulation,[],[f17588,f903])).
fof(f17588,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))))))))))))),
  inference(forward_demodulation,[],[f17587,f2517])).
fof(f17587,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f17586,f1033])).
fof(f17586,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f17585,f1023])).
fof(f17585,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f16793,f2508])).
fof(f16793,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))))))))))))),
  inference(superposition,[],[f67,f1019])).
fof(f67,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86532,plain,(
  elem_352 = select(earray_262,index_57) | index_267 = index_57),
  inference(superposition,[],[f86514,f4514])).
fof(f4514,plain,(
  ( ! [X30] : (select(earray_262,X30) = select(earray_268,X30) | index_267 = X30) )),
  inference(superposition,[],[f2,f1447])).
fof(f1447,plain,(
  earray_268 = store(earray_262,index_267,e43)),
  inference(backward_demodulation,[],[f1446,f148])).
fof(f148,axiom,(
  earray_268 = store(earray_266,index_267,e43)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1446,plain,(
  earray_262 = earray_266),
  inference(backward_demodulation,[],[f1445,f525])).
fof(f525,plain,(
  earray_266 = rselect_seq(queue_265)),
  inference(definition_unfolding,[],[f147,f485])).
fof(f147,axiom,(
  earray_266 = rselect_seq(q43)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1445,plain,(
  earray_262 = rselect_seq(queue_265)),
  inference(forward_demodulation,[],[f1325,f808])).
fof(f808,plain,(
  earray_262 = rselect_seq(queue_263)),
  inference(superposition,[],[f3,f627])).
fof(f627,plain,(
  queue_263 = rstore_seq(queue_259,earray_262)),
  inference(definition_unfolding,[],[f391,f484])).
fof(f391,axiom,(
  queue_263 = rstore_seq(q42,earray_262)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1325,plain,(
  rselect_seq(queue_265) = rselect_seq(queue_263)),
  inference(superposition,[],[f7,f970])).
fof(f970,plain,(
  queue_265 = rstore_tail(queue_263,index_267)),
  inference(backward_demodulation,[],[f968,f392])).
fof(f86514,plain,(
  elem_352 = select(earray_268,index_57)),
  inference(subsumption_resolution,[],[f86512,f18621])).
fof(f18621,plain,(
  index_273 != index_57),
  inference(forward_demodulation,[],[f18620,f975])).
fof(f975,plain,(
  index_273 = s(index_267)),
  inference(backward_demodulation,[],[f972,f268])).
fof(f268,axiom,(
  index_270 = s(index_267)),
  file('SWV568-1.050.tptp',unknown)).
fof(f972,plain,(
  index_273 = index_270),
  inference(forward_demodulation,[],[f877,f576])).
fof(f576,plain,(
  index_273 = rselect_tail(queue_271)),
  inference(definition_unfolding,[],[f269,f486])).
fof(f486,axiom,(
  q44 = queue_271),
  file('SWV568-1.050.tptp',unknown)).
fof(f269,axiom,(
  index_273 = rselect_tail(q44)),
  file('SWV568-1.050.tptp',unknown)).
fof(f877,plain,(
  index_270 = rselect_tail(queue_271)),
  inference(superposition,[],[f5,f394])).
fof(f394,axiom,(
  queue_271 = rstore_tail(queue_269,index_270)),
  file('SWV568-1.050.tptp',unknown)).
fof(f18620,plain,(
  index_57 != s(index_267)),
  inference(forward_demodulation,[],[f18619,f971])).
fof(f18619,plain,(
  index_57 != s(s(index_261))),
  inference(forward_demodulation,[],[f18618,f2262])).
fof(f18618,plain,(
  index_57 != s(s(s(index_252)))),
  inference(forward_demodulation,[],[f18617,f967])).
fof(f18617,plain,(
  index_57 != s(s(s(s(index_246))))),
  inference(forward_demodulation,[],[f18616,f963])).
fof(f18616,plain,(
  index_57 != s(s(s(s(s(index_240)))))),
  inference(forward_demodulation,[],[f18615,f2253])).
fof(f18615,plain,(
  index_57 != s(s(s(s(s(s(index_225))))))),
  inference(forward_demodulation,[],[f18614,f955])).
fof(f18614,plain,(
  index_57 != s(s(s(s(s(s(s(index_219)))))))),
  inference(forward_demodulation,[],[f18613,f951])).
fof(f18613,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_213))))))))),
  inference(forward_demodulation,[],[f18612,f2244])).
fof(f18612,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_204)))))))))),
  inference(forward_demodulation,[],[f18611,f947])).
fof(f18611,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_198))))))))))),
  inference(forward_demodulation,[],[f18610,f943])).
fof(f18610,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_192)))))))))))),
  inference(forward_demodulation,[],[f18609,f2235])).
fof(f18609,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_183))))))))))))),
  inference(forward_demodulation,[],[f18608,f939])).
fof(f18608,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_177)))))))))))))),
  inference(forward_demodulation,[],[f18607,f935])).
fof(f18607,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_171))))))))))))))),
  inference(forward_demodulation,[],[f18606,f2226])).
fof(f18606,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_162)))))))))))))))),
  inference(forward_demodulation,[],[f18605,f927])).
fof(f18605,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_147))))))))))))))))),
  inference(forward_demodulation,[],[f18604,f923])).
fof(f18604,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141)))))))))))))))))),
  inference(forward_demodulation,[],[f18603,f2148])).
fof(f18603,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132))))))))))))))))))),
  inference(forward_demodulation,[],[f18602,f919])).
fof(f18602,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126)))))))))))))))))))),
  inference(forward_demodulation,[],[f18601,f915])).
fof(f18601,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120))))))))))))))))))))),
  inference(forward_demodulation,[],[f18600,f2139])).
fof(f18600,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111)))))))))))))))))))))),
  inference(forward_demodulation,[],[f18599,f907])).
fof(f18599,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105))))))))))))))))))))))),
  inference(forward_demodulation,[],[f18598,f903])).
fof(f18598,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f18597,f2517])).
fof(f18597,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f18596,f1033])).
fof(f18596,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f18595,f1023])).
fof(f18595,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f17785,f2508])).
fof(f17785,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))))))))))))))),
  inference(superposition,[],[f66,f1019])).
fof(f66,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86512,plain,(
  elem_352 = select(earray_268,index_57) | index_273 = index_57),
  inference(superposition,[],[f86509,f4515])).
fof(f4515,plain,(
  ( ! [X31] : (select(earray_268,X31) = select(earray_274,X31) | index_273 = X31) )),
  inference(superposition,[],[f2,f1450])).
fof(f1450,plain,(
  earray_274 = store(earray_268,index_273,e44)),
  inference(backward_demodulation,[],[f1449,f150])).
fof(f150,axiom,(
  earray_274 = store(earray_272,index_273,e44)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1449,plain,(
  earray_268 = earray_272),
  inference(backward_demodulation,[],[f1448,f526])).
fof(f526,plain,(
  earray_272 = rselect_seq(queue_271)),
  inference(definition_unfolding,[],[f149,f486])).
fof(f149,axiom,(
  earray_272 = rselect_seq(q44)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1448,plain,(
  earray_268 = rselect_seq(queue_271)),
  inference(forward_demodulation,[],[f1326,f809])).
fof(f809,plain,(
  earray_268 = rselect_seq(queue_269)),
  inference(superposition,[],[f3,f628])).
fof(f628,plain,(
  queue_269 = rstore_seq(queue_265,earray_268)),
  inference(definition_unfolding,[],[f393,f485])).
fof(f393,axiom,(
  queue_269 = rstore_seq(q43,earray_268)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1326,plain,(
  rselect_seq(queue_271) = rselect_seq(queue_269)),
  inference(superposition,[],[f7,f974])).
fof(f974,plain,(
  queue_271 = rstore_tail(queue_269,index_273)),
  inference(backward_demodulation,[],[f972,f394])).
fof(f86509,plain,(
  elem_352 = select(earray_274,index_57)),
  inference(subsumption_resolution,[],[f86507,f19655])).
fof(f19655,plain,(
  index_282 != index_57),
  inference(forward_demodulation,[],[f19654,f2280])).
fof(f2280,plain,(
  index_282 = s(index_273)),
  inference(backward_demodulation,[],[f2279,f270])).
fof(f270,axiom,(
  index_276 = s(index_273)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2279,plain,(
  index_282 = index_276),
  inference(forward_demodulation,[],[f2278,f577])).
fof(f577,plain,(
  index_282 = rselect_tail(queue_280)),
  inference(definition_unfolding,[],[f273,f487])).
fof(f487,axiom,(
  q45 = queue_280),
  file('SWV568-1.050.tptp',unknown)).
fof(f273,axiom,(
  index_282 = rselect_tail(q45)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2278,plain,(
  index_276 = rselect_tail(queue_280)),
  inference(forward_demodulation,[],[f2065,f878])).
fof(f878,plain,(
  index_276 = rselect_tail(queue_277)),
  inference(superposition,[],[f5,f396])).
fof(f396,axiom,(
  queue_277 = rstore_tail(queue_275,index_276)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2065,plain,(
  rselect_tail(queue_280) = rselect_tail(queue_277)),
  inference(superposition,[],[f11,f398])).
fof(f398,axiom,(
  queue_280 = rstore_head(queue_277,index_279)),
  file('SWV568-1.050.tptp',unknown)).
fof(f19654,plain,(
  index_57 != s(index_273)),
  inference(forward_demodulation,[],[f19653,f975])).
fof(f19653,plain,(
  index_57 != s(s(index_267))),
  inference(forward_demodulation,[],[f19652,f971])).
fof(f19652,plain,(
  index_57 != s(s(s(index_261)))),
  inference(forward_demodulation,[],[f19651,f2262])).
fof(f19651,plain,(
  index_57 != s(s(s(s(index_252))))),
  inference(forward_demodulation,[],[f19650,f967])).
fof(f19650,plain,(
  index_57 != s(s(s(s(s(index_246)))))),
  inference(forward_demodulation,[],[f19649,f963])).
fof(f19649,plain,(
  index_57 != s(s(s(s(s(s(index_240))))))),
  inference(forward_demodulation,[],[f19648,f2253])).
fof(f19648,plain,(
  index_57 != s(s(s(s(s(s(s(index_225)))))))),
  inference(forward_demodulation,[],[f19647,f955])).
fof(f19647,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_219))))))))),
  inference(forward_demodulation,[],[f19646,f951])).
fof(f19646,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_213)))))))))),
  inference(forward_demodulation,[],[f19645,f2244])).
fof(f19645,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_204))))))))))),
  inference(forward_demodulation,[],[f19644,f947])).
fof(f19644,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_198)))))))))))),
  inference(forward_demodulation,[],[f19643,f943])).
fof(f19643,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_192))))))))))))),
  inference(forward_demodulation,[],[f19642,f2235])).
fof(f19642,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_183)))))))))))))),
  inference(forward_demodulation,[],[f19641,f939])).
fof(f19641,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_177))))))))))))))),
  inference(forward_demodulation,[],[f19640,f935])).
fof(f19640,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_171)))))))))))))))),
  inference(forward_demodulation,[],[f19639,f2226])).
fof(f19639,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_162))))))))))))))))),
  inference(forward_demodulation,[],[f19638,f927])).
fof(f19638,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_147)))))))))))))))))),
  inference(forward_demodulation,[],[f19637,f923])).
fof(f19637,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141))))))))))))))))))),
  inference(forward_demodulation,[],[f19636,f2148])).
fof(f19636,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132)))))))))))))))))))),
  inference(forward_demodulation,[],[f19635,f919])).
fof(f19635,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126))))))))))))))))))))),
  inference(forward_demodulation,[],[f19634,f915])).
fof(f19634,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120)))))))))))))))))))))),
  inference(forward_demodulation,[],[f19633,f2139])).
fof(f19633,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111))))))))))))))))))))))),
  inference(forward_demodulation,[],[f19632,f907])).
fof(f19632,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f19631,f903])).
fof(f19631,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f19630,f2517])).
fof(f19630,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f19629,f1033])).
fof(f19629,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f19628,f1023])).
fof(f19628,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f18801,f2508])).
fof(f18801,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))))))))))))))),
  inference(superposition,[],[f65,f1019])).
fof(f65,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86507,plain,(
  elem_352 = select(earray_274,index_57) | index_282 = index_57),
  inference(superposition,[],[f86504,f4516])).
fof(f4516,plain,(
  ( ! [X32] : (select(earray_274,X32) = select(earray_283,X32) | index_282 = X32) )),
  inference(superposition,[],[f2,f1454])).
fof(f1454,plain,(
  earray_283 = store(earray_274,index_282,e45)),
  inference(backward_demodulation,[],[f1452,f152])).
fof(f152,axiom,(
  earray_283 = store(earray_281,index_282,e45)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1452,plain,(
  earray_274 = earray_281),
  inference(forward_demodulation,[],[f1451,f810])).
fof(f810,plain,(
  earray_274 = rselect_seq(queue_275)),
  inference(superposition,[],[f3,f629])).
fof(f629,plain,(
  queue_275 = rstore_seq(queue_271,earray_274)),
  inference(definition_unfolding,[],[f395,f486])).
fof(f395,axiom,(
  queue_275 = rstore_seq(q44,earray_274)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1451,plain,(
  earray_281 = rselect_seq(queue_275)),
  inference(forward_demodulation,[],[f1327,f1292])).
fof(f1292,plain,(
  earray_281 = rselect_seq(queue_277)),
  inference(forward_demodulation,[],[f1275,f527])).
fof(f527,plain,(
  earray_281 = rselect_seq(queue_280)),
  inference(definition_unfolding,[],[f151,f487])).
fof(f151,axiom,(
  earray_281 = rselect_seq(q45)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1275,plain,(
  rselect_seq(queue_280) = rselect_seq(queue_277)),
  inference(superposition,[],[f6,f398])).
fof(f1327,plain,(
  rselect_seq(queue_275) = rselect_seq(queue_277)),
  inference(superposition,[],[f7,f396])).
fof(f86504,plain,(
  elem_352 = select(earray_283,index_57)),
  inference(subsumption_resolution,[],[f86502,f20711])).
fof(f20711,plain,(
  index_288 != index_57),
  inference(forward_demodulation,[],[f20710,f979])).
fof(f979,plain,(
  index_288 = s(index_282)),
  inference(backward_demodulation,[],[f976,f274])).
fof(f274,axiom,(
  index_285 = s(index_282)),
  file('SWV568-1.050.tptp',unknown)).
fof(f976,plain,(
  index_288 = index_285),
  inference(forward_demodulation,[],[f879,f578])).
fof(f578,plain,(
  index_288 = rselect_tail(queue_286)),
  inference(definition_unfolding,[],[f275,f488])).
fof(f488,axiom,(
  q46 = queue_286),
  file('SWV568-1.050.tptp',unknown)).
fof(f275,axiom,(
  index_288 = rselect_tail(q46)),
  file('SWV568-1.050.tptp',unknown)).
fof(f879,plain,(
  index_285 = rselect_tail(queue_286)),
  inference(superposition,[],[f5,f400])).
fof(f400,axiom,(
  queue_286 = rstore_tail(queue_284,index_285)),
  file('SWV568-1.050.tptp',unknown)).
fof(f20710,plain,(
  index_57 != s(index_282)),
  inference(forward_demodulation,[],[f20709,f2280])).
fof(f20709,plain,(
  index_57 != s(s(index_273))),
  inference(forward_demodulation,[],[f20708,f975])).
fof(f20708,plain,(
  index_57 != s(s(s(index_267)))),
  inference(forward_demodulation,[],[f20707,f971])).
fof(f20707,plain,(
  index_57 != s(s(s(s(index_261))))),
  inference(forward_demodulation,[],[f20706,f2262])).
fof(f20706,plain,(
  index_57 != s(s(s(s(s(index_252)))))),
  inference(forward_demodulation,[],[f20705,f967])).
fof(f20705,plain,(
  index_57 != s(s(s(s(s(s(index_246))))))),
  inference(forward_demodulation,[],[f20704,f963])).
fof(f20704,plain,(
  index_57 != s(s(s(s(s(s(s(index_240)))))))),
  inference(forward_demodulation,[],[f20703,f2253])).
fof(f20703,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_225))))))))),
  inference(forward_demodulation,[],[f20702,f955])).
fof(f20702,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_219)))))))))),
  inference(forward_demodulation,[],[f20701,f951])).
fof(f20701,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_213))))))))))),
  inference(forward_demodulation,[],[f20700,f2244])).
fof(f20700,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_204)))))))))))),
  inference(forward_demodulation,[],[f20699,f947])).
fof(f20699,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_198))))))))))))),
  inference(forward_demodulation,[],[f20698,f943])).
fof(f20698,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_192)))))))))))))),
  inference(forward_demodulation,[],[f20697,f2235])).
fof(f20697,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_183))))))))))))))),
  inference(forward_demodulation,[],[f20696,f939])).
fof(f20696,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_177)))))))))))))))),
  inference(forward_demodulation,[],[f20695,f935])).
fof(f20695,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_171))))))))))))))))),
  inference(forward_demodulation,[],[f20694,f2226])).
fof(f20694,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_162)))))))))))))))))),
  inference(forward_demodulation,[],[f20693,f927])).
fof(f20693,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_147))))))))))))))))))),
  inference(forward_demodulation,[],[f20692,f923])).
fof(f20692,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141)))))))))))))))))))),
  inference(forward_demodulation,[],[f20691,f2148])).
fof(f20691,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132))))))))))))))))))))),
  inference(forward_demodulation,[],[f20690,f919])).
fof(f20690,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126)))))))))))))))))))))),
  inference(forward_demodulation,[],[f20689,f915])).
fof(f20689,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120))))))))))))))))))))))),
  inference(forward_demodulation,[],[f20688,f2139])).
fof(f20688,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f20687,f907])).
fof(f20687,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f20686,f903])).
fof(f20686,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f20685,f2517])).
fof(f20685,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f20684,f1033])).
fof(f20684,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f20683,f1023])).
fof(f20683,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f19840,f2508])).
fof(f19840,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))))))))))))))))),
  inference(superposition,[],[f64,f1019])).
fof(f64,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86502,plain,(
  elem_352 = select(earray_283,index_57) | index_288 = index_57),
  inference(superposition,[],[f86499,f4517])).
fof(f4517,plain,(
  ( ! [X33] : (select(earray_283,X33) = select(earray_289,X33) | index_288 = X33) )),
  inference(superposition,[],[f2,f1458])).
fof(f1458,plain,(
  earray_289 = store(earray_283,index_288,e46)),
  inference(backward_demodulation,[],[f1457,f154])).
fof(f154,axiom,(
  earray_289 = store(earray_287,index_288,e46)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1457,plain,(
  earray_283 = earray_287),
  inference(backward_demodulation,[],[f1456,f528])).
fof(f528,plain,(
  earray_287 = rselect_seq(queue_286)),
  inference(definition_unfolding,[],[f153,f488])).
fof(f153,axiom,(
  earray_287 = rselect_seq(q46)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1456,plain,(
  earray_283 = rselect_seq(queue_286)),
  inference(forward_demodulation,[],[f1328,f812])).
fof(f812,plain,(
  earray_283 = rselect_seq(queue_284)),
  inference(superposition,[],[f3,f630])).
fof(f630,plain,(
  queue_284 = rstore_seq(queue_280,earray_283)),
  inference(definition_unfolding,[],[f399,f487])).
fof(f399,axiom,(
  queue_284 = rstore_seq(q45,earray_283)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1328,plain,(
  rselect_seq(queue_286) = rselect_seq(queue_284)),
  inference(superposition,[],[f7,f978])).
fof(f978,plain,(
  queue_286 = rstore_tail(queue_284,index_288)),
  inference(backward_demodulation,[],[f976,f400])).
fof(f86499,plain,(
  elem_352 = select(earray_289,index_57)),
  inference(subsumption_resolution,[],[f86497,f21787])).
fof(f21787,plain,(
  index_294 != index_57),
  inference(forward_demodulation,[],[f21786,f983])).
fof(f983,plain,(
  index_294 = s(index_288)),
  inference(backward_demodulation,[],[f980,f276])).
fof(f276,axiom,(
  index_291 = s(index_288)),
  file('SWV568-1.050.tptp',unknown)).
fof(f980,plain,(
  index_294 = index_291),
  inference(forward_demodulation,[],[f880,f579])).
fof(f579,plain,(
  index_294 = rselect_tail(queue_292)),
  inference(definition_unfolding,[],[f277,f489])).
fof(f489,axiom,(
  q47 = queue_292),
  file('SWV568-1.050.tptp',unknown)).
fof(f277,axiom,(
  index_294 = rselect_tail(q47)),
  file('SWV568-1.050.tptp',unknown)).
fof(f880,plain,(
  index_291 = rselect_tail(queue_292)),
  inference(superposition,[],[f5,f402])).
fof(f402,axiom,(
  queue_292 = rstore_tail(queue_290,index_291)),
  file('SWV568-1.050.tptp',unknown)).
fof(f21786,plain,(
  index_57 != s(index_288)),
  inference(forward_demodulation,[],[f21785,f979])).
fof(f21785,plain,(
  index_57 != s(s(index_282))),
  inference(forward_demodulation,[],[f21784,f2280])).
fof(f21784,plain,(
  index_57 != s(s(s(index_273)))),
  inference(forward_demodulation,[],[f21783,f975])).
fof(f21783,plain,(
  index_57 != s(s(s(s(index_267))))),
  inference(forward_demodulation,[],[f21782,f971])).
fof(f21782,plain,(
  index_57 != s(s(s(s(s(index_261)))))),
  inference(forward_demodulation,[],[f21781,f2262])).
fof(f21781,plain,(
  index_57 != s(s(s(s(s(s(index_252))))))),
  inference(forward_demodulation,[],[f21780,f967])).
fof(f21780,plain,(
  index_57 != s(s(s(s(s(s(s(index_246)))))))),
  inference(forward_demodulation,[],[f21779,f963])).
fof(f21779,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_240))))))))),
  inference(forward_demodulation,[],[f21778,f2253])).
fof(f21778,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_225)))))))))),
  inference(forward_demodulation,[],[f21777,f955])).
fof(f21777,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_219))))))))))),
  inference(forward_demodulation,[],[f21776,f951])).
fof(f21776,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_213)))))))))))),
  inference(forward_demodulation,[],[f21775,f2244])).
fof(f21775,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_204))))))))))))),
  inference(forward_demodulation,[],[f21774,f947])).
fof(f21774,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_198)))))))))))))),
  inference(forward_demodulation,[],[f21773,f943])).
fof(f21773,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_192))))))))))))))),
  inference(forward_demodulation,[],[f21772,f2235])).
fof(f21772,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_183)))))))))))))))),
  inference(forward_demodulation,[],[f21771,f939])).
fof(f21771,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_177))))))))))))))))),
  inference(forward_demodulation,[],[f21770,f935])).
fof(f21770,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_171)))))))))))))))))),
  inference(forward_demodulation,[],[f21769,f2226])).
fof(f21769,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_162))))))))))))))))))),
  inference(forward_demodulation,[],[f21768,f927])).
fof(f21768,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_147)))))))))))))))))))),
  inference(forward_demodulation,[],[f21767,f923])).
fof(f21767,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141))))))))))))))))))))),
  inference(forward_demodulation,[],[f21766,f2148])).
fof(f21766,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132)))))))))))))))))))))),
  inference(forward_demodulation,[],[f21765,f919])).
fof(f21765,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126))))))))))))))))))))))),
  inference(forward_demodulation,[],[f21764,f915])).
fof(f21764,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f21763,f2139])).
fof(f21763,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f21762,f907])).
fof(f21762,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f21761,f903])).
fof(f21761,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f21760,f2517])).
fof(f21760,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f21759,f1033])).
fof(f21759,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f21758,f1023])).
fof(f21758,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f20900,f2508])).
fof(f20900,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))))))))))))))))),
  inference(superposition,[],[f63,f1019])).
fof(f63,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86497,plain,(
  elem_352 = select(earray_289,index_57) | index_294 = index_57),
  inference(superposition,[],[f86494,f4518])).
fof(f4518,plain,(
  ( ! [X34] : (select(earray_289,X34) = select(earray_295,X34) | index_294 = X34) )),
  inference(superposition,[],[f2,f1461])).
fof(f1461,plain,(
  earray_295 = store(earray_289,index_294,e47)),
  inference(backward_demodulation,[],[f1460,f157])).
fof(f157,axiom,(
  earray_295 = store(earray_293,index_294,e47)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1460,plain,(
  earray_289 = earray_293),
  inference(backward_demodulation,[],[f1459,f530])).
fof(f530,plain,(
  earray_293 = rselect_seq(queue_292)),
  inference(definition_unfolding,[],[f156,f489])).
fof(f156,axiom,(
  earray_293 = rselect_seq(q47)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1459,plain,(
  earray_289 = rselect_seq(queue_292)),
  inference(forward_demodulation,[],[f1329,f813])).
fof(f813,plain,(
  earray_289 = rselect_seq(queue_290)),
  inference(superposition,[],[f3,f631])).
fof(f631,plain,(
  queue_290 = rstore_seq(queue_286,earray_289)),
  inference(definition_unfolding,[],[f401,f488])).
fof(f401,axiom,(
  queue_290 = rstore_seq(q46,earray_289)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1329,plain,(
  rselect_seq(queue_292) = rselect_seq(queue_290)),
  inference(superposition,[],[f7,f982])).
fof(f982,plain,(
  queue_292 = rstore_tail(queue_290,index_294)),
  inference(backward_demodulation,[],[f980,f402])).
fof(f86494,plain,(
  elem_352 = select(earray_295,index_57)),
  inference(subsumption_resolution,[],[f86492,f22883])).
fof(f22883,plain,(
  index_303 != index_57),
  inference(forward_demodulation,[],[f22882,f2289])).
fof(f2289,plain,(
  index_303 = s(index_294)),
  inference(backward_demodulation,[],[f2288,f278])).
fof(f278,axiom,(
  index_297 = s(index_294)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2288,plain,(
  index_303 = index_297),
  inference(forward_demodulation,[],[f2287,f582])).
fof(f582,plain,(
  index_303 = rselect_tail(queue_301)),
  inference(definition_unfolding,[],[f283,f490])).
fof(f490,axiom,(
  q48 = queue_301),
  file('SWV568-1.050.tptp',unknown)).
fof(f283,axiom,(
  index_303 = rselect_tail(q48)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2287,plain,(
  index_297 = rselect_tail(queue_301)),
  inference(forward_demodulation,[],[f2066,f881])).
fof(f881,plain,(
  index_297 = rselect_tail(queue_298)),
  inference(superposition,[],[f5,f404])).
fof(f404,axiom,(
  queue_298 = rstore_tail(queue_296,index_297)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2066,plain,(
  rselect_tail(queue_301) = rselect_tail(queue_298)),
  inference(superposition,[],[f11,f1779])).
fof(f1779,plain,(
  queue_301 = rstore_head(queue_298,index_351)),
  inference(backward_demodulation,[],[f1778,f405])).
fof(f405,axiom,(
  queue_301 = rstore_head(queue_298,index_300)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1778,plain,(
  index_351 = index_300),
  inference(backward_demodulation,[],[f1777,f1776])).
fof(f1776,plain,(
  index_300 = rselect_head(queue_307)),
  inference(backward_demodulation,[],[f1775,f1566])).
fof(f1566,plain,(
  rselect_head(queue_305) = rselect_head(queue_307)),
  inference(superposition,[],[f8,f986])).
fof(f986,plain,(
  queue_307 = rstore_tail(queue_305,index_315)),
  inference(backward_demodulation,[],[f984,f407])).
fof(f407,axiom,(
  queue_307 = rstore_tail(queue_305,index_306)),
  file('SWV568-1.050.tptp',unknown)).
fof(f984,plain,(
  index_315 = index_306),
  inference(forward_demodulation,[],[f882,f584])).
fof(f584,plain,(
  index_315 = rselect_tail(queue_307)),
  inference(definition_unfolding,[],[f287,f491])).
fof(f491,axiom,(
  q49 = queue_307),
  file('SWV568-1.050.tptp',unknown)).
fof(f287,axiom,(
  index_315 = rselect_tail(q49)),
  file('SWV568-1.050.tptp',unknown)).
fof(f882,plain,(
  index_306 = rselect_tail(queue_307)),
  inference(superposition,[],[f5,f407])).
fof(f1775,plain,(
  index_300 = rselect_head(queue_305)),
  inference(forward_demodulation,[],[f1633,f844])).
fof(f844,plain,(
  index_300 = rselect_head(queue_301)),
  inference(superposition,[],[f4,f405])).
fof(f1633,plain,(
  rselect_head(queue_301) = rselect_head(queue_305)),
  inference(superposition,[],[f9,f633])).
fof(f633,plain,(
  queue_305 = rstore_seq(queue_301,earray_304)),
  inference(definition_unfolding,[],[f406,f490])).
fof(f406,axiom,(
  queue_305 = rstore_seq(q48,earray_304)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1777,plain,(
  index_351 = rselect_head(queue_307)),
  inference(forward_demodulation,[],[f1634,f1595])).
fof(f1595,plain,(
  index_351 = rselect_head(queue_317)),
  inference(forward_demodulation,[],[f1568,f589])).
fof(f589,plain,(
  index_351 = rselect_head(queue_319)),
  inference(definition_unfolding,[],[f302,f493])).
fof(f493,axiom,(
  q50 = queue_319),
  file('SWV568-1.050.tptp',unknown)).
fof(f302,axiom,(
  index_351 = rselect_head(q50)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1568,plain,(
  rselect_head(queue_319) = rselect_head(queue_317)),
  inference(superposition,[],[f8,f411])).
fof(f411,axiom,(
  queue_319 = rstore_tail(queue_317,index_318)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1634,plain,(
  rselect_head(queue_307) = rselect_head(queue_317)),
  inference(superposition,[],[f9,f635])).
fof(f635,plain,(
  queue_317 = rstore_seq(queue_307,earray_316)),
  inference(definition_unfolding,[],[f410,f491])).
fof(f410,axiom,(
  queue_317 = rstore_seq(q49,earray_316)),
  file('SWV568-1.050.tptp',unknown)).
fof(f22882,plain,(
  index_57 != s(index_294)),
  inference(forward_demodulation,[],[f22881,f983])).
fof(f22881,plain,(
  index_57 != s(s(index_288))),
  inference(forward_demodulation,[],[f22880,f979])).
fof(f22880,plain,(
  index_57 != s(s(s(index_282)))),
  inference(forward_demodulation,[],[f22879,f2280])).
fof(f22879,plain,(
  index_57 != s(s(s(s(index_273))))),
  inference(forward_demodulation,[],[f22878,f975])).
fof(f22878,plain,(
  index_57 != s(s(s(s(s(index_267)))))),
  inference(forward_demodulation,[],[f22877,f971])).
fof(f22877,plain,(
  index_57 != s(s(s(s(s(s(index_261))))))),
  inference(forward_demodulation,[],[f22876,f2262])).
fof(f22876,plain,(
  index_57 != s(s(s(s(s(s(s(index_252)))))))),
  inference(forward_demodulation,[],[f22875,f967])).
fof(f22875,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_246))))))))),
  inference(forward_demodulation,[],[f22874,f963])).
fof(f22874,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_240)))))))))),
  inference(forward_demodulation,[],[f22873,f2253])).
fof(f22873,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_225))))))))))),
  inference(forward_demodulation,[],[f22872,f955])).
fof(f22872,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_219)))))))))))),
  inference(forward_demodulation,[],[f22871,f951])).
fof(f22871,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_213))))))))))))),
  inference(forward_demodulation,[],[f22870,f2244])).
fof(f22870,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_204)))))))))))))),
  inference(forward_demodulation,[],[f22869,f947])).
fof(f22869,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_198))))))))))))))),
  inference(forward_demodulation,[],[f22868,f943])).
fof(f22868,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_192)))))))))))))))),
  inference(forward_demodulation,[],[f22867,f2235])).
fof(f22867,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_183))))))))))))))))),
  inference(forward_demodulation,[],[f22866,f939])).
fof(f22866,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_177)))))))))))))))))),
  inference(forward_demodulation,[],[f22865,f935])).
fof(f22865,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_171))))))))))))))))))),
  inference(forward_demodulation,[],[f22864,f2226])).
fof(f22864,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_162)))))))))))))))))))),
  inference(forward_demodulation,[],[f22863,f927])).
fof(f22863,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_147))))))))))))))))))))),
  inference(forward_demodulation,[],[f22862,f923])).
fof(f22862,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141)))))))))))))))))))))),
  inference(forward_demodulation,[],[f22861,f2148])).
fof(f22861,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132))))))))))))))))))))))),
  inference(forward_demodulation,[],[f22860,f919])).
fof(f22860,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f22859,f915])).
fof(f22859,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f22858,f2139])).
fof(f22858,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111)))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f22857,f907])).
fof(f22857,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f22856,f903])).
fof(f22856,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99)))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f22855,f2517])).
fof(f22855,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f22854,f1033])).
fof(f22854,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84)))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f22853,f1023])).
fof(f22853,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72))))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f21980,f2508])).
fof(f21980,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63)))))))))))))))))))))))))))))))),
  inference(superposition,[],[f62,f1019])).
fof(f62,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4)))))))))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86492,plain,(
  elem_352 = select(earray_295,index_57) | index_303 = index_57),
  inference(superposition,[],[f86491,f4519])).
fof(f4519,plain,(
  ( ! [X35] : (select(earray_295,X35) = select(earray_304,X35) | index_303 = X35) )),
  inference(superposition,[],[f2,f1465])).
fof(f1465,plain,(
  earray_304 = store(earray_295,index_303,e48)),
  inference(backward_demodulation,[],[f1463,f159])).
fof(f159,axiom,(
  earray_304 = store(earray_302,index_303,e48)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1463,plain,(
  earray_295 = earray_302),
  inference(forward_demodulation,[],[f1462,f814])).
fof(f814,plain,(
  earray_295 = rselect_seq(queue_296)),
  inference(superposition,[],[f3,f632])).
fof(f632,plain,(
  queue_296 = rstore_seq(queue_292,earray_295)),
  inference(definition_unfolding,[],[f403,f489])).
fof(f403,axiom,(
  queue_296 = rstore_seq(q47,earray_295)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1462,plain,(
  earray_302 = rselect_seq(queue_296)),
  inference(forward_demodulation,[],[f1330,f1293])).
fof(f1293,plain,(
  earray_302 = rselect_seq(queue_298)),
  inference(forward_demodulation,[],[f1276,f531])).
fof(f531,plain,(
  earray_302 = rselect_seq(queue_301)),
  inference(definition_unfolding,[],[f158,f490])).
fof(f158,axiom,(
  earray_302 = rselect_seq(q48)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1276,plain,(
  rselect_seq(queue_301) = rselect_seq(queue_298)),
  inference(superposition,[],[f6,f405])).
fof(f1330,plain,(
  rselect_seq(queue_296) = rselect_seq(queue_298)),
  inference(superposition,[],[f7,f404])).
fof(f86491,plain,(
  elem_352 = select(earray_304,index_57)),
  inference(subsumption_resolution,[],[f86487,f23998])).
fof(f23998,plain,(
  index_315 != index_57),
  inference(forward_demodulation,[],[f23997,f987])).
fof(f987,plain,(
  index_315 = s(index_303)),
  inference(backward_demodulation,[],[f984,f284])).
fof(f284,axiom,(
  index_306 = s(index_303)),
  file('SWV568-1.050.tptp',unknown)).
fof(f23997,plain,(
  index_57 != s(index_303)),
  inference(forward_demodulation,[],[f23996,f2289])).
fof(f23996,plain,(
  index_57 != s(s(index_294))),
  inference(forward_demodulation,[],[f23995,f983])).
fof(f23995,plain,(
  index_57 != s(s(s(index_288)))),
  inference(forward_demodulation,[],[f23994,f979])).
fof(f23994,plain,(
  index_57 != s(s(s(s(index_282))))),
  inference(forward_demodulation,[],[f23993,f2280])).
fof(f23993,plain,(
  index_57 != s(s(s(s(s(index_273)))))),
  inference(forward_demodulation,[],[f23992,f975])).
fof(f23992,plain,(
  index_57 != s(s(s(s(s(s(index_267))))))),
  inference(forward_demodulation,[],[f23991,f971])).
fof(f23991,plain,(
  index_57 != s(s(s(s(s(s(s(index_261)))))))),
  inference(forward_demodulation,[],[f23990,f2262])).
fof(f23990,plain,(
  index_57 != s(s(s(s(s(s(s(s(index_252))))))))),
  inference(forward_demodulation,[],[f23989,f967])).
fof(f23989,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(index_246)))))))))),
  inference(forward_demodulation,[],[f23988,f963])).
fof(f23988,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(index_240))))))))))),
  inference(forward_demodulation,[],[f23987,f2253])).
fof(f23987,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(index_225)))))))))))),
  inference(forward_demodulation,[],[f23986,f955])).
fof(f23986,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(index_219))))))))))))),
  inference(forward_demodulation,[],[f23985,f951])).
fof(f23985,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(index_213)))))))))))))),
  inference(forward_demodulation,[],[f23984,f2244])).
fof(f23984,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_204))))))))))))))),
  inference(forward_demodulation,[],[f23983,f947])).
fof(f23983,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_198)))))))))))))))),
  inference(forward_demodulation,[],[f23982,f943])).
fof(f23982,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_192))))))))))))))))),
  inference(forward_demodulation,[],[f23981,f2235])).
fof(f23981,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_183)))))))))))))))))),
  inference(forward_demodulation,[],[f23980,f939])).
fof(f23980,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_177))))))))))))))))))),
  inference(forward_demodulation,[],[f23979,f935])).
fof(f23979,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_171)))))))))))))))))))),
  inference(forward_demodulation,[],[f23978,f2226])).
fof(f23978,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_162))))))))))))))))))))),
  inference(forward_demodulation,[],[f23977,f927])).
fof(f23977,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_147)))))))))))))))))))))),
  inference(forward_demodulation,[],[f23976,f923])).
fof(f23976,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_141))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23975,f2148])).
fof(f23975,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_132)))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23974,f919])).
fof(f23974,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_126))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23973,f915])).
fof(f23973,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_120)))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23972,f2139])).
fof(f23972,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_111))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23971,f907])).
fof(f23971,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_105)))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23970,f903])).
fof(f23970,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_99))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23969,f2517])).
fof(f23969,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_90)))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23968,f1033])).
fof(f23968,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_84))))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23967,f1023])).
fof(f23967,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_72)))))))))))))))))))))))))))))))),
  inference(forward_demodulation,[],[f23079,f2508])).
fof(f23079,plain,(
  index_57 != s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(index_63))))))))))))))))))))))))))))))))),
  inference(superposition,[],[f61,f1019])).
fof(f61,axiom,(
  ( ! [X4] : (s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(s(X4))))))))))))))))))))))))))))))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86487,plain,(
  elem_352 = select(earray_304,index_57) | index_315 = index_57),
  inference(superposition,[],[f4520,f2495])).
fof(f2495,plain,(
  elem_352 = select(earray_316,index_57)),
  inference(backward_demodulation,[],[f2493,f1475])).
fof(f1475,plain,(
  elem_352 = select(earray_316,index_351)),
  inference(backward_demodulation,[],[f1474,f195])).
fof(f195,axiom,(
  elem_352 = select(earray_350,index_351)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1474,plain,(
  earray_316 = earray_350),
  inference(backward_demodulation,[],[f1473,f539])).
fof(f539,plain,(
  earray_350 = rselect_seq(queue_319)),
  inference(definition_unfolding,[],[f174,f493])).
fof(f174,axiom,(
  earray_350 = rselect_seq(q50)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1473,plain,(
  earray_316 = rselect_seq(queue_319)),
  inference(forward_demodulation,[],[f1333,f816])).
fof(f816,plain,(
  earray_316 = rselect_seq(queue_317)),
  inference(superposition,[],[f3,f635])).
fof(f1333,plain,(
  rselect_seq(queue_319) = rselect_seq(queue_317)),
  inference(superposition,[],[f7,f411])).
fof(f2493,plain,(
  index_57 = index_351),
  inference(forward_demodulation,[],[f2486,f1015])).
fof(f1015,plain,(
  index_57 = s(index_51)),
  inference(backward_demodulation,[],[f1012,f310])).
fof(f310,axiom,(
  index_54 = s(index_51)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1012,plain,(
  index_57 = index_54),
  inference(forward_demodulation,[],[f893,f593])).
fof(f593,plain,(
  index_57 = rselect_tail(queue_55)),
  inference(definition_unfolding,[],[f311,f455])).
fof(f455,axiom,(
  q16 = queue_55),
  file('SWV568-1.050.tptp',unknown)).
fof(f311,axiom,(
  index_57 = rselect_tail(q16)),
  file('SWV568-1.050.tptp',unknown)).
fof(f893,plain,(
  index_54 = rselect_tail(queue_55)),
  inference(superposition,[],[f5,f431])).
fof(f431,axiom,(
  queue_55 = rstore_tail(queue_53,index_54)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2486,plain,(
  index_351 = s(index_51)),
  inference(backward_demodulation,[],[f2473,f1784])).
fof(f1784,plain,(
  index_351 = s(index_279)),
  inference(backward_demodulation,[],[f1778,f1772])).
fof(f1772,plain,(
  index_300 = s(index_279)),
  inference(backward_demodulation,[],[f1767,f282])).
fof(f282,axiom,(
  index_300 = s(index_299)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1767,plain,(
  index_279 = index_299),
  inference(forward_demodulation,[],[f1766,f1765])).
fof(f1765,plain,(
  index_279 = rselect_head(queue_292)),
  inference(backward_demodulation,[],[f1764,f1564])).
fof(f1564,plain,(
  rselect_head(queue_290) = rselect_head(queue_292)),
  inference(superposition,[],[f8,f982])).
fof(f1764,plain,(
  index_279 = rselect_head(queue_290)),
  inference(forward_demodulation,[],[f1631,f1763])).
fof(f1763,plain,(
  index_279 = rselect_head(queue_286)),
  inference(backward_demodulation,[],[f1762,f1563])).
fof(f1563,plain,(
  rselect_head(queue_284) = rselect_head(queue_286)),
  inference(superposition,[],[f8,f978])).
fof(f1762,plain,(
  index_279 = rselect_head(queue_284)),
  inference(forward_demodulation,[],[f1630,f843])).
fof(f843,plain,(
  index_279 = rselect_head(queue_280)),
  inference(superposition,[],[f4,f398])).
fof(f1630,plain,(
  rselect_head(queue_280) = rselect_head(queue_284)),
  inference(superposition,[],[f9,f630])).
fof(f1631,plain,(
  rselect_head(queue_286) = rselect_head(queue_290)),
  inference(superposition,[],[f9,f631])).
fof(f1766,plain,(
  index_299 = rselect_head(queue_292)),
  inference(forward_demodulation,[],[f1632,f1594])).
fof(f1594,plain,(
  index_299 = rselect_head(queue_296)),
  inference(forward_demodulation,[],[f1565,f279])).
fof(f279,axiom,(
  index_299 = rselect_head(queue_298)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1565,plain,(
  rselect_head(queue_298) = rselect_head(queue_296)),
  inference(superposition,[],[f8,f404])).
fof(f1632,plain,(
  rselect_head(queue_292) = rselect_head(queue_296)),
  inference(superposition,[],[f9,f632])).
fof(f2473,plain,(
  index_51 = index_279),
  inference(forward_demodulation,[],[f2472,f592])).
fof(f592,plain,(
  index_51 = rselect_tail(queue_49)),
  inference(definition_unfolding,[],[f309,f454])).
fof(f454,axiom,(
  q15 = queue_49),
  file('SWV568-1.050.tptp',unknown)).
fof(f309,axiom,(
  index_51 = rselect_tail(q15)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2472,plain,(
  index_279 = rselect_tail(queue_49)),
  inference(forward_demodulation,[],[f2069,f2458])).
fof(f2458,plain,(
  index_279 = rselect_tail(queue_46)),
  inference(backward_demodulation,[],[f2455,f891])).
fof(f891,plain,(
  index_45 = rselect_tail(queue_46)),
  inference(superposition,[],[f5,f427])).
fof(f427,axiom,(
  queue_46 = rstore_tail(queue_44,index_45)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2455,plain,(
  index_279 = index_45),
  inference(backward_demodulation,[],[f2450,f306])).
fof(f306,axiom,(
  index_45 = s(index_42)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2450,plain,(
  index_279 = s(index_42)),
  inference(backward_demodulation,[],[f2440,f1757])).
fof(f1757,plain,(
  index_279 = s(index_258)),
  inference(backward_demodulation,[],[f1752,f272])).
fof(f272,axiom,(
  index_279 = s(index_278)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1752,plain,(
  index_258 = index_278),
  inference(forward_demodulation,[],[f1751,f1750])).
fof(f1750,plain,(
  index_258 = rselect_head(queue_271)),
  inference(backward_demodulation,[],[f1749,f1561])).
fof(f1561,plain,(
  rselect_head(queue_269) = rselect_head(queue_271)),
  inference(superposition,[],[f8,f974])).
fof(f1749,plain,(
  index_258 = rselect_head(queue_269)),
  inference(forward_demodulation,[],[f1627,f1748])).
fof(f1748,plain,(
  index_258 = rselect_head(queue_265)),
  inference(backward_demodulation,[],[f1747,f1560])).
fof(f1560,plain,(
  rselect_head(queue_263) = rselect_head(queue_265)),
  inference(superposition,[],[f8,f970])).
fof(f1747,plain,(
  index_258 = rselect_head(queue_263)),
  inference(forward_demodulation,[],[f1626,f841])).
fof(f841,plain,(
  index_258 = rselect_head(queue_259)),
  inference(superposition,[],[f4,f390])).
fof(f1626,plain,(
  rselect_head(queue_259) = rselect_head(queue_263)),
  inference(superposition,[],[f9,f627])).
fof(f1627,plain,(
  rselect_head(queue_265) = rselect_head(queue_269)),
  inference(superposition,[],[f9,f628])).
fof(f1751,plain,(
  index_278 = rselect_head(queue_271)),
  inference(forward_demodulation,[],[f1628,f1593])).
fof(f1593,plain,(
  index_278 = rselect_head(queue_275)),
  inference(forward_demodulation,[],[f1562,f271])).
fof(f271,axiom,(
  index_278 = rselect_head(queue_277)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1562,plain,(
  rselect_head(queue_277) = rselect_head(queue_275)),
  inference(superposition,[],[f8,f396])).
fof(f1628,plain,(
  rselect_head(queue_271) = rselect_head(queue_275)),
  inference(superposition,[],[f9,f629])).
fof(f2440,plain,(
  index_42 = index_258),
  inference(forward_demodulation,[],[f2436,f1007])).
fof(f1007,plain,(
  index_42 = s(index_36)),
  inference(backward_demodulation,[],[f1004,f304])).
fof(f304,axiom,(
  index_39 = s(index_36)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1004,plain,(
  index_42 = index_39),
  inference(forward_demodulation,[],[f890,f591])).
fof(f591,plain,(
  index_42 = rselect_tail(queue_40)),
  inference(definition_unfolding,[],[f305,f453])).
fof(f453,axiom,(
  q14 = queue_40),
  file('SWV568-1.050.tptp',unknown)).
fof(f305,axiom,(
  index_42 = rselect_tail(q14)),
  file('SWV568-1.050.tptp',unknown)).
fof(f890,plain,(
  index_39 = rselect_tail(queue_40)),
  inference(superposition,[],[f5,f425])).
fof(f425,axiom,(
  queue_40 = rstore_tail(queue_38,index_39)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2436,plain,(
  index_258 = s(index_36)),
  inference(backward_demodulation,[],[f2425,f1744])).
fof(f1744,plain,(
  index_258 = s(index_231)),
  inference(backward_demodulation,[],[f1739,f262])).
fof(f262,axiom,(
  index_258 = s(index_257)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1739,plain,(
  index_231 = index_257),
  inference(forward_demodulation,[],[f1738,f1737])).
fof(f1737,plain,(
  index_231 = rselect_head(queue_250)),
  inference(backward_demodulation,[],[f1736,f1558])).
fof(f1558,plain,(
  rselect_head(queue_248) = rselect_head(queue_250)),
  inference(superposition,[],[f8,f966])).
fof(f1736,plain,(
  index_231 = rselect_head(queue_248)),
  inference(forward_demodulation,[],[f1624,f1733])).
fof(f1733,plain,(
  index_231 = rselect_head(queue_244)),
  inference(backward_demodulation,[],[f1732,f1557])).
fof(f1557,plain,(
  rselect_head(queue_242) = rselect_head(queue_244)),
  inference(superposition,[],[f8,f962])).
fof(f1732,plain,(
  index_231 = rselect_head(queue_242)),
  inference(forward_demodulation,[],[f1622,f840])).
fof(f840,plain,(
  index_231 = rselect_head(queue_232)),
  inference(superposition,[],[f4,f380])).
fof(f1622,plain,(
  rselect_head(queue_232) = rselect_head(queue_242)),
  inference(superposition,[],[f9,f624])).
fof(f1624,plain,(
  rselect_head(queue_244) = rselect_head(queue_248)),
  inference(superposition,[],[f9,f625])).
fof(f1738,plain,(
  index_257 = rselect_head(queue_250)),
  inference(forward_demodulation,[],[f1625,f1592])).
fof(f1592,plain,(
  index_257 = rselect_head(queue_254)),
  inference(forward_demodulation,[],[f1559,f261])).
fof(f261,axiom,(
  index_257 = rselect_head(queue_256)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1559,plain,(
  rselect_head(queue_256) = rselect_head(queue_254)),
  inference(superposition,[],[f8,f389])).
fof(f1625,plain,(
  rselect_head(queue_250) = rselect_head(queue_254)),
  inference(superposition,[],[f9,f626])).
fof(f2425,plain,(
  index_36 = index_231),
  inference(forward_demodulation,[],[f2421,f995])).
fof(f995,plain,(
  index_36 = s(index_30)),
  inference(backward_demodulation,[],[f992,f293])).
fof(f293,axiom,(
  index_33 = s(index_30)),
  file('SWV568-1.050.tptp',unknown)).
fof(f992,plain,(
  index_36 = index_33),
  inference(forward_demodulation,[],[f885,f590])).
fof(f590,plain,(
  index_36 = rselect_tail(queue_34)),
  inference(definition_unfolding,[],[f303,f452])).
fof(f452,axiom,(
  q13 = queue_34),
  file('SWV568-1.050.tptp',unknown)).
fof(f303,axiom,(
  index_36 = rselect_tail(q13)),
  file('SWV568-1.050.tptp',unknown)).
fof(f885,plain,(
  index_33 = rselect_tail(queue_34)),
  inference(superposition,[],[f5,f419])).
fof(f419,axiom,(
  queue_34 = rstore_tail(queue_32,index_33)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2421,plain,(
  index_231 = s(index_30)),
  inference(backward_demodulation,[],[f2410,f1729])).
fof(f1729,plain,(
  index_231 = s(index_210)),
  inference(backward_demodulation,[],[f1724,f251])).
fof(f251,axiom,(
  index_231 = s(index_230)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1724,plain,(
  index_210 = index_230),
  inference(forward_demodulation,[],[f1723,f1722])).
fof(f1722,plain,(
  index_210 = rselect_head(queue_223)),
  inference(backward_demodulation,[],[f1721,f1553])).
fof(f1553,plain,(
  rselect_head(queue_221) = rselect_head(queue_223)),
  inference(superposition,[],[f8,f954])).
fof(f1721,plain,(
  index_210 = rselect_head(queue_221)),
  inference(forward_demodulation,[],[f1620,f1720])).
fof(f1720,plain,(
  index_210 = rselect_head(queue_217)),
  inference(backward_demodulation,[],[f1719,f1552])).
fof(f1552,plain,(
  rselect_head(queue_215) = rselect_head(queue_217)),
  inference(superposition,[],[f8,f950])).
fof(f1719,plain,(
  index_210 = rselect_head(queue_215)),
  inference(forward_demodulation,[],[f1619,f839])).
fof(f839,plain,(
  index_210 = rselect_head(queue_211)),
  inference(superposition,[],[f4,f372])).
fof(f1619,plain,(
  rselect_head(queue_211) = rselect_head(queue_215)),
  inference(superposition,[],[f9,f619])).
fof(f1620,plain,(
  rselect_head(queue_217) = rselect_head(queue_221)),
  inference(superposition,[],[f9,f620])).
fof(f1723,plain,(
  index_230 = rselect_head(queue_223)),
  inference(forward_demodulation,[],[f1621,f1590])).
fof(f1590,plain,(
  index_230 = rselect_head(queue_227)),
  inference(forward_demodulation,[],[f1554,f250])).
fof(f250,axiom,(
  index_230 = rselect_head(queue_229)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1554,plain,(
  rselect_head(queue_229) = rselect_head(queue_227)),
  inference(superposition,[],[f8,f378])).
fof(f1621,plain,(
  rselect_head(queue_223) = rselect_head(queue_227)),
  inference(superposition,[],[f9,f621])).
fof(f2410,plain,(
  index_30 = index_210),
  inference(forward_demodulation,[],[f2406,f2271])).
fof(f2271,plain,(
  index_30 = s(index_21)),
  inference(backward_demodulation,[],[f2270,f254])).
fof(f254,axiom,(
  index_24 = s(index_21)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2270,plain,(
  index_30 = index_24),
  inference(forward_demodulation,[],[f2269,f581])).
fof(f581,plain,(
  index_30 = rselect_tail(queue_28)),
  inference(definition_unfolding,[],[f281,f451])).
fof(f451,axiom,(
  q12 = queue_28),
  file('SWV568-1.050.tptp',unknown)).
fof(f281,axiom,(
  index_30 = rselect_tail(q12)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2269,plain,(
  index_24 = rselect_tail(queue_28)),
  inference(forward_demodulation,[],[f2064,f871])).
fof(f871,plain,(
  index_24 = rselect_tail(queue_25)),
  inference(superposition,[],[f5,f386])).
fof(f386,axiom,(
  queue_25 = rstore_tail(queue_23,index_24)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2064,plain,(
  rselect_tail(queue_28) = rselect_tail(queue_25)),
  inference(superposition,[],[f11,f397])).
fof(f397,axiom,(
  queue_28 = rstore_head(queue_25,index_27)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2406,plain,(
  index_210 = s(index_21)),
  inference(backward_demodulation,[],[f2395,f1716])).
fof(f1716,plain,(
  index_210 = s(index_189)),
  inference(backward_demodulation,[],[f1711,f243])).
fof(f243,axiom,(
  index_210 = s(index_209)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1711,plain,(
  index_189 = index_209),
  inference(forward_demodulation,[],[f1710,f1709])).
fof(f1709,plain,(
  index_189 = rselect_head(queue_202)),
  inference(backward_demodulation,[],[f1708,f1550])).
fof(f1550,plain,(
  rselect_head(queue_200) = rselect_head(queue_202)),
  inference(superposition,[],[f8,f946])).
fof(f1708,plain,(
  index_189 = rselect_head(queue_200)),
  inference(forward_demodulation,[],[f1617,f1707])).
fof(f1707,plain,(
  index_189 = rselect_head(queue_196)),
  inference(backward_demodulation,[],[f1706,f1549])).
fof(f1549,plain,(
  rselect_head(queue_194) = rselect_head(queue_196)),
  inference(superposition,[],[f8,f942])).
fof(f1706,plain,(
  index_189 = rselect_head(queue_194)),
  inference(forward_demodulation,[],[f1616,f838])).
fof(f838,plain,(
  index_189 = rselect_head(queue_190)),
  inference(superposition,[],[f4,f365])).
fof(f1616,plain,(
  rselect_head(queue_190) = rselect_head(queue_194)),
  inference(superposition,[],[f9,f616])).
fof(f1617,plain,(
  rselect_head(queue_196) = rselect_head(queue_200)),
  inference(superposition,[],[f9,f617])).
fof(f1710,plain,(
  index_209 = rselect_head(queue_202)),
  inference(forward_demodulation,[],[f1618,f1589])).
fof(f1589,plain,(
  index_209 = rselect_head(queue_206)),
  inference(forward_demodulation,[],[f1551,f241])).
fof(f241,axiom,(
  index_209 = rselect_head(queue_208)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1551,plain,(
  rselect_head(queue_208) = rselect_head(queue_206)),
  inference(superposition,[],[f8,f371])).
fof(f1618,plain,(
  rselect_head(queue_202) = rselect_head(queue_206)),
  inference(superposition,[],[f9,f618])).
fof(f2395,plain,(
  index_21 = index_189),
  inference(forward_demodulation,[],[f2390,f931])).
fof(f931,plain,(
  index_21 = s(index_15)),
  inference(backward_demodulation,[],[f928,f229])).
fof(f229,axiom,(
  index_18 = s(index_15)),
  file('SWV568-1.050.tptp',unknown)).
fof(f928,plain,(
  index_21 = index_18),
  inference(forward_demodulation,[],[f861,f566])).
fof(f566,plain,(
  index_21 = rselect_tail(queue_19)),
  inference(definition_unfolding,[],[f242,f450])).
fof(f450,axiom,(
  q11 = queue_19),
  file('SWV568-1.050.tptp',unknown)).
fof(f242,axiom,(
  index_21 = rselect_tail(q11)),
  file('SWV568-1.050.tptp',unknown)).
fof(f861,plain,(
  index_18 = rselect_tail(queue_19)),
  inference(superposition,[],[f5,f364])).
fof(f364,axiom,(
  queue_19 = rstore_tail(queue_17,index_18)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2390,plain,(
  s(index_15) = index_189),
  inference(backward_demodulation,[],[f2379,f1699])).
fof(f1699,plain,(
  index_189 = s(index_168)),
  inference(backward_demodulation,[],[f1694,f234])).
fof(f234,axiom,(
  index_189 = s(index_188)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1694,plain,(
  index_168 = index_188),
  inference(forward_demodulation,[],[f1693,f1692])).
fof(f1692,plain,(
  index_168 = rselect_head(queue_181)),
  inference(backward_demodulation,[],[f1691,f1547])).
fof(f1547,plain,(
  rselect_head(queue_179) = rselect_head(queue_181)),
  inference(superposition,[],[f8,f938])).
fof(f1691,plain,(
  index_168 = rselect_head(queue_179)),
  inference(forward_demodulation,[],[f1613,f1690])).
fof(f1690,plain,(
  index_168 = rselect_head(queue_175)),
  inference(backward_demodulation,[],[f1689,f1546])).
fof(f1546,plain,(
  rselect_head(queue_173) = rselect_head(queue_175)),
  inference(superposition,[],[f8,f934])).
fof(f1689,plain,(
  index_168 = rselect_head(queue_173)),
  inference(forward_demodulation,[],[f1612,f837])).
fof(f837,plain,(
  index_168 = rselect_head(queue_169)),
  inference(superposition,[],[f4,f356])).
fof(f1612,plain,(
  rselect_head(queue_169) = rselect_head(queue_173)),
  inference(superposition,[],[f9,f613])).
fof(f1613,plain,(
  rselect_head(queue_175) = rselect_head(queue_179)),
  inference(superposition,[],[f9,f614])).
fof(f1693,plain,(
  index_188 = rselect_head(queue_181)),
  inference(forward_demodulation,[],[f1614,f1588])).
fof(f1588,plain,(
  index_188 = rselect_head(queue_185)),
  inference(forward_demodulation,[],[f1548,f233])).
fof(f233,axiom,(
  index_188 = rselect_head(queue_187)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1548,plain,(
  rselect_head(queue_187) = rselect_head(queue_185)),
  inference(superposition,[],[f8,f363])).
fof(f1614,plain,(
  rselect_head(queue_181) = rselect_head(queue_185)),
  inference(superposition,[],[f9,f615])).
fof(f2379,plain,(
  index_15 = index_168),
  inference(forward_demodulation,[],[f2367,f911])).
fof(f911,plain,(
  index_15 = s(index_9)),
  inference(backward_demodulation,[],[f908,f204])).
fof(f204,axiom,(
  index_12 = s(index_9)),
  file('SWV568-1.050.tptp',unknown)).
fof(f908,plain,(
  index_15 = index_12),
  inference(forward_demodulation,[],[f852,f557])).
fof(f557,plain,(
  index_15 = rselect_tail(queue_13)),
  inference(definition_unfolding,[],[f216,f449])).
fof(f449,axiom,(
  q10 = queue_13),
  file('SWV568-1.050.tptp',unknown)).
fof(f216,axiom,(
  index_15 = rselect_tail(q10)),
  file('SWV568-1.050.tptp',unknown)).
fof(f852,plain,(
  index_12 = rselect_tail(queue_13)),
  inference(superposition,[],[f5,f342])).
fof(f342,axiom,(
  queue_13 = rstore_tail(queue_11,index_12)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2367,plain,(
  s(index_9) = index_168),
  inference(backward_demodulation,[],[f2356,f1684])).
fof(f1684,plain,(
  index_168 = s(index_138)),
  inference(backward_demodulation,[],[f1679,f225])).
fof(f225,axiom,(
  index_168 = s(index_167)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1679,plain,(
  index_138 = index_167),
  inference(forward_demodulation,[],[f1678,f1677])).
fof(f1677,plain,(
  index_138 = rselect_head(queue_151)),
  inference(backward_demodulation,[],[f1676,f1542])).
fof(f1542,plain,(
  rselect_head(queue_149) = rselect_head(queue_151)),
  inference(superposition,[],[f8,f926])).
fof(f1676,plain,(
  index_138 = rselect_head(queue_149)),
  inference(forward_demodulation,[],[f1609,f1675])).
fof(f1675,plain,(
  index_138 = rselect_head(queue_145)),
  inference(backward_demodulation,[],[f1674,f1541])).
fof(f1541,plain,(
  rselect_head(queue_143) = rselect_head(queue_145)),
  inference(superposition,[],[f8,f922])).
fof(f1674,plain,(
  index_138 = rselect_head(queue_143)),
  inference(forward_demodulation,[],[f1608,f835])).
fof(f835,plain,(
  index_138 = rselect_head(queue_139)),
  inference(superposition,[],[f4,f346])).
fof(f1608,plain,(
  rselect_head(queue_139) = rselect_head(queue_143)),
  inference(superposition,[],[f9,f608])).
fof(f1609,plain,(
  rselect_head(queue_145) = rselect_head(queue_149)),
  inference(superposition,[],[f9,f609])).
fof(f1678,plain,(
  index_167 = rselect_head(queue_151)),
  inference(forward_demodulation,[],[f1610,f1587])).
fof(f1587,plain,(
  index_167 = rselect_head(queue_164)),
  inference(forward_demodulation,[],[f1544,f224])).
fof(f224,axiom,(
  index_167 = rselect_head(queue_166)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1544,plain,(
  rselect_head(queue_166) = rselect_head(queue_164)),
  inference(superposition,[],[f8,f355])).
fof(f1610,plain,(
  rselect_head(queue_151) = rselect_head(queue_164)),
  inference(superposition,[],[f9,f611])).
fof(f2356,plain,(
  index_9 = index_138),
  inference(forward_demodulation,[],[f2355,f598])).
fof(f598,plain,(
  index_9 = rselect_tail(queue_349)),
  inference(definition_unfolding,[],[f324,f497])).
fof(f497,axiom,(
  q9 = queue_349),
  file('SWV568-1.050.tptp',unknown)).
fof(f324,axiom,(
  index_9 = rselect_tail(q9)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2355,plain,(
  index_138 = rselect_tail(queue_349)),
  inference(forward_demodulation,[],[f2068,f2350])).
fof(f2350,plain,(
  index_138 = rselect_tail(queue_346)),
  inference(backward_demodulation,[],[f2347,f889])).
fof(f889,plain,(
  index_345 = rselect_tail(queue_346)),
  inference(superposition,[],[f5,f422])).
fof(f422,axiom,(
  queue_346 = rstore_tail(queue_344,index_345)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2347,plain,(
  index_138 = index_345),
  inference(backward_demodulation,[],[f2344,f299])).
fof(f299,axiom,(
  index_345 = s(index_342)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2344,plain,(
  index_138 = s(index_342)),
  inference(backward_demodulation,[],[f2333,f1671])).
fof(f1671,plain,(
  index_138 = s(index_117)),
  inference(backward_demodulation,[],[f1666,f212])).
fof(f212,axiom,(
  index_138 = s(index_137)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1666,plain,(
  index_117 = index_137),
  inference(forward_demodulation,[],[f1665,f1662])).
fof(f1662,plain,(
  index_117 = rselect_head(queue_130)),
  inference(backward_demodulation,[],[f1661,f1539])).
fof(f1539,plain,(
  rselect_head(queue_128) = rselect_head(queue_130)),
  inference(superposition,[],[f8,f918])).
fof(f1661,plain,(
  index_117 = rselect_head(queue_128)),
  inference(forward_demodulation,[],[f1605,f1660])).
fof(f1660,plain,(
  index_117 = rselect_head(queue_124)),
  inference(backward_demodulation,[],[f1659,f1538])).
fof(f1538,plain,(
  rselect_head(queue_122) = rselect_head(queue_124)),
  inference(superposition,[],[f8,f914])).
fof(f1659,plain,(
  index_117 = rselect_head(queue_122)),
  inference(forward_demodulation,[],[f1604,f834])).
fof(f834,plain,(
  index_117 = rselect_head(queue_118)),
  inference(superposition,[],[f4,f338])).
fof(f1604,plain,(
  rselect_head(queue_118) = rselect_head(queue_122)),
  inference(superposition,[],[f9,f605])).
fof(f1605,plain,(
  rselect_head(queue_124) = rselect_head(queue_128)),
  inference(superposition,[],[f9,f606])).
fof(f1665,plain,(
  index_137 = rselect_head(queue_130)),
  inference(forward_demodulation,[],[f1607,f1585])).
fof(f1585,plain,(
  index_137 = rselect_head(queue_134)),
  inference(forward_demodulation,[],[f1540,f211])).
fof(f211,axiom,(
  index_137 = rselect_head(queue_136)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1540,plain,(
  rselect_head(queue_136) = rselect_head(queue_134)),
  inference(superposition,[],[f8,f345])).
fof(f1607,plain,(
  rselect_head(queue_130) = rselect_head(queue_134)),
  inference(superposition,[],[f9,f607])).
fof(f2333,plain,(
  index_342 = index_117),
  inference(forward_demodulation,[],[f2319,f1003])).
fof(f1003,plain,(
  index_342 = s(index_336)),
  inference(backward_demodulation,[],[f1000,f297])).
fof(f297,axiom,(
  index_339 = s(index_336)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1000,plain,(
  index_342 = index_339),
  inference(forward_demodulation,[],[f888,f588])).
fof(f588,plain,(
  index_342 = rselect_tail(queue_340)),
  inference(definition_unfolding,[],[f298,f496])).
fof(f496,axiom,(
  q8 = queue_340),
  file('SWV568-1.050.tptp',unknown)).
fof(f298,axiom,(
  index_342 = rselect_tail(q8)),
  file('SWV568-1.050.tptp',unknown)).
fof(f888,plain,(
  index_339 = rselect_tail(queue_340)),
  inference(superposition,[],[f5,f420])).
fof(f420,axiom,(
  queue_340 = rstore_tail(queue_338,index_339)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2319,plain,(
  index_117 = s(index_336)),
  inference(backward_demodulation,[],[f2317,f203])).
fof(f203,axiom,(
  index_117 = s(index_116)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2317,plain,(
  index_336 = index_116),
  inference(forward_demodulation,[],[f2313,f999])).
fof(f999,plain,(
  index_336 = s(index_330)),
  inference(backward_demodulation,[],[f996,f295])).
fof(f295,axiom,(
  index_333 = s(index_330)),
  file('SWV568-1.050.tptp',unknown)).
fof(f996,plain,(
  index_336 = index_333),
  inference(forward_demodulation,[],[f887,f587])).
fof(f587,plain,(
  index_336 = rselect_tail(queue_334)),
  inference(definition_unfolding,[],[f296,f495])).
fof(f495,axiom,(
  q7 = queue_334),
  file('SWV568-1.050.tptp',unknown)).
fof(f296,axiom,(
  index_336 = rselect_tail(q7)),
  file('SWV568-1.050.tptp',unknown)).
fof(f887,plain,(
  index_333 = rselect_tail(queue_334)),
  inference(superposition,[],[f5,f417])).
fof(f417,axiom,(
  queue_334 = rstore_tail(queue_332,index_333)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2313,plain,(
  index_116 = s(index_330)),
  inference(backward_demodulation,[],[f2297,f2220])).
fof(f2220,plain,(
  index_116 = s(index_324)),
  inference(backward_demodulation,[],[f2210,f1873])).
fof(f1873,plain,(
  index_116 = s(index_69)),
  inference(backward_demodulation,[],[f1868,f1864])).
fof(f1864,plain,(
  index_96 = s(index_69)),
  inference(backward_demodulation,[],[f1859,f328])).
fof(f328,axiom,(
  index_96 = s(index_95)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1859,plain,(
  index_69 = index_95),
  inference(forward_demodulation,[],[f1858,f1850])).
fof(f1850,plain,(
  index_69 = rselect_head(queue_88)),
  inference(backward_demodulation,[],[f1849,f1582])).
fof(f1582,plain,(
  rselect_head(queue_86) = rselect_head(queue_88)),
  inference(superposition,[],[f8,f1032])).
fof(f1032,plain,(
  queue_88 = rstore_tail(queue_86,index_90)),
  inference(backward_demodulation,[],[f1030,f443])).
fof(f1849,plain,(
  index_69 = rselect_head(queue_86)),
  inference(forward_demodulation,[],[f1647,f1848])).
fof(f1848,plain,(
  index_69 = rselect_head(queue_76)),
  inference(backward_demodulation,[],[f1847,f1580])).
fof(f1580,plain,(
  rselect_head(queue_74) = rselect_head(queue_76)),
  inference(superposition,[],[f8,f1022])).
fof(f1022,plain,(
  queue_76 = rstore_tail(queue_74,index_84)),
  inference(backward_demodulation,[],[f1020,f439])).
fof(f1847,plain,(
  index_69 = rselect_head(queue_74)),
  inference(forward_demodulation,[],[f1646,f848])).
fof(f848,plain,(
  index_69 = rselect_head(queue_70)),
  inference(superposition,[],[f4,f437])).
fof(f1646,plain,(
  rselect_head(queue_70) = rselect_head(queue_74)),
  inference(superposition,[],[f9,f647])).
fof(f647,plain,(
  queue_74 = rstore_seq(queue_70,earray_73)),
  inference(definition_unfolding,[],[f438,f457])).
fof(f438,axiom,(
  queue_74 = rstore_seq(q18,earray_73)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1647,plain,(
  rselect_head(queue_76) = rselect_head(queue_86)),
  inference(superposition,[],[f9,f649])).
fof(f649,plain,(
  queue_86 = rstore_seq(queue_76,earray_85)),
  inference(definition_unfolding,[],[f442,f458])).
fof(f442,axiom,(
  queue_86 = rstore_seq(q19,earray_85)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1858,plain,(
  index_95 = rselect_head(queue_88)),
  inference(forward_demodulation,[],[f1649,f1600])).
fof(f1600,plain,(
  index_95 = rselect_head(queue_92)),
  inference(forward_demodulation,[],[f1583,f327])).
fof(f327,axiom,(
  index_95 = rselect_head(queue_94)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1583,plain,(
  rselect_head(queue_94) = rselect_head(queue_92)),
  inference(superposition,[],[f8,f445])).
fof(f1649,plain,(
  rselect_head(queue_88) = rselect_head(queue_92)),
  inference(superposition,[],[f9,f1530])).
fof(f1530,plain,(
  queue_92 = rstore_seq(queue_88,earray_98)),
  inference(backward_demodulation,[],[f1528,f650])).
fof(f650,plain,(
  queue_92 = rstore_seq(queue_88,earray_91)),
  inference(definition_unfolding,[],[f444,f460])).
fof(f444,axiom,(
  queue_92 = rstore_seq(q20,earray_91)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1528,plain,(
  earray_98 = earray_91),
  inference(backward_demodulation,[],[f1527,f831])).
fof(f831,plain,(
  earray_91 = rselect_seq(queue_92)),
  inference(superposition,[],[f3,f650])).
fof(f1527,plain,(
  earray_98 = rselect_seq(queue_92)),
  inference(forward_demodulation,[],[f1348,f1298])).
fof(f1298,plain,(
  earray_98 = rselect_seq(queue_94)),
  inference(forward_demodulation,[],[f1281,f549])).
fof(f549,plain,(
  earray_98 = rselect_seq(queue_97)),
  inference(definition_unfolding,[],[f194,f461])).
fof(f194,axiom,(
  earray_98 = rselect_seq(q21)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1281,plain,(
  rselect_seq(queue_97) = rselect_seq(queue_94)),
  inference(superposition,[],[f6,f446])).
fof(f1348,plain,(
  rselect_seq(queue_92) = rselect_seq(queue_94)),
  inference(superposition,[],[f7,f445])).
fof(f2210,plain,(
  index_324 = index_69),
  inference(forward_demodulation,[],[f2205,f290])).
fof(f290,axiom,(
  index_324 = s(index_321)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2205,plain,(
  s(index_321) = index_69),
  inference(backward_demodulation,[],[f2195,f1842])).
fof(f1842,plain,(
  index_69 = s(index_48)),
  inference(backward_demodulation,[],[f1837,f317])).
fof(f317,axiom,(
  index_69 = s(index_68)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1837,plain,(
  index_48 = index_68),
  inference(forward_demodulation,[],[f1836,f1835])).
fof(f1835,plain,(
  index_48 = rselect_head(queue_61)),
  inference(backward_demodulation,[],[f1834,f1578])).
fof(f1578,plain,(
  rselect_head(queue_59) = rselect_head(queue_61)),
  inference(superposition,[],[f8,f1018])).
fof(f1018,plain,(
  queue_61 = rstore_tail(queue_59,index_63)),
  inference(backward_demodulation,[],[f1016,f433])).
fof(f1834,plain,(
  index_48 = rselect_head(queue_59)),
  inference(forward_demodulation,[],[f1643,f1833])).
fof(f1833,plain,(
  index_48 = rselect_head(queue_55)),
  inference(backward_demodulation,[],[f1832,f1577])).
fof(f1577,plain,(
  rselect_head(queue_53) = rselect_head(queue_55)),
  inference(superposition,[],[f8,f1014])).
fof(f1014,plain,(
  queue_55 = rstore_tail(queue_53,index_57)),
  inference(backward_demodulation,[],[f1012,f431])).
fof(f1832,plain,(
  index_48 = rselect_head(queue_53)),
  inference(forward_demodulation,[],[f1642,f847])).
fof(f847,plain,(
  index_48 = rselect_head(queue_49)),
  inference(superposition,[],[f4,f428])).
fof(f428,axiom,(
  queue_49 = rstore_head(queue_46,index_48)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1642,plain,(
  rselect_head(queue_49) = rselect_head(queue_53)),
  inference(superposition,[],[f9,f644])).
fof(f644,plain,(
  queue_53 = rstore_seq(queue_49,earray_52)),
  inference(definition_unfolding,[],[f430,f454])).
fof(f430,axiom,(
  queue_53 = rstore_seq(q15,earray_52)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1643,plain,(
  rselect_head(queue_55) = rselect_head(queue_59)),
  inference(superposition,[],[f9,f645])).
fof(f645,plain,(
  queue_59 = rstore_seq(queue_55,earray_58)),
  inference(definition_unfolding,[],[f432,f455])).
fof(f432,axiom,(
  queue_59 = rstore_seq(q16,earray_58)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1836,plain,(
  index_68 = rselect_head(queue_61)),
  inference(forward_demodulation,[],[f1644,f1599])).
fof(f1599,plain,(
  index_68 = rselect_head(queue_65)),
  inference(forward_demodulation,[],[f1579,f316])).
fof(f316,axiom,(
  index_68 = rselect_head(queue_67)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1579,plain,(
  rselect_head(queue_67) = rselect_head(queue_65)),
  inference(superposition,[],[f8,f435])).
fof(f1644,plain,(
  rselect_head(queue_61) = rselect_head(queue_65)),
  inference(superposition,[],[f9,f646])).
fof(f646,plain,(
  queue_65 = rstore_seq(queue_61,earray_64)),
  inference(definition_unfolding,[],[f434,f456])).
fof(f434,axiom,(
  queue_65 = rstore_seq(q17,earray_64)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2195,plain,(
  index_321 = index_48),
  inference(forward_demodulation,[],[f2191,f991])).
fof(f991,plain,(
  index_321 = s(index_309)),
  inference(backward_demodulation,[],[f988,f286])).
fof(f286,axiom,(
  index_312 = s(index_309)),
  file('SWV568-1.050.tptp',unknown)).
fof(f988,plain,(
  index_321 = index_312),
  inference(forward_demodulation,[],[f883,f585])).
fof(f585,plain,(
  index_321 = rselect_tail(queue_313)),
  inference(definition_unfolding,[],[f289,f492])).
fof(f492,axiom,(
  q5 = queue_313),
  file('SWV568-1.050.tptp',unknown)).
fof(f289,axiom,(
  index_321 = rselect_tail(q5)),
  file('SWV568-1.050.tptp',unknown)).
fof(f883,plain,(
  index_312 = rselect_tail(queue_313)),
  inference(superposition,[],[f5,f409])).
fof(f409,axiom,(
  queue_313 = rstore_tail(queue_311,index_312)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2191,plain,(
  s(index_309) = index_48),
  inference(backward_demodulation,[],[f2180,f1829])).
fof(f1829,plain,(
  index_48 = s(index_27)),
  inference(backward_demodulation,[],[f1824,f308])).
fof(f308,axiom,(
  index_48 = s(index_47)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1824,plain,(
  index_27 = index_47),
  inference(forward_demodulation,[],[f1823,f1803])).
fof(f1803,plain,(
  index_27 = rselect_head(queue_40)),
  inference(backward_demodulation,[],[f1802,f1574])).
fof(f1574,plain,(
  rselect_head(queue_38) = rselect_head(queue_40)),
  inference(superposition,[],[f8,f1006])).
fof(f1006,plain,(
  queue_40 = rstore_tail(queue_38,index_42)),
  inference(backward_demodulation,[],[f1004,f425])).
fof(f1802,plain,(
  index_27 = rselect_head(queue_38)),
  inference(forward_demodulation,[],[f1638,f1761])).
fof(f1761,plain,(
  index_27 = rselect_head(queue_34)),
  inference(backward_demodulation,[],[f1760,f1569])).
fof(f1569,plain,(
  rselect_head(queue_32) = rselect_head(queue_34)),
  inference(superposition,[],[f8,f994])).
fof(f994,plain,(
  queue_34 = rstore_tail(queue_32,index_36)),
  inference(backward_demodulation,[],[f992,f419])).
fof(f1760,plain,(
  index_27 = rselect_head(queue_32)),
  inference(forward_demodulation,[],[f1629,f842])).
fof(f842,plain,(
  index_27 = rselect_head(queue_28)),
  inference(superposition,[],[f4,f397])).
fof(f1629,plain,(
  rselect_head(queue_28) = rselect_head(queue_32)),
  inference(superposition,[],[f9,f636])).
fof(f636,plain,(
  queue_32 = rstore_seq(queue_28,earray_31)),
  inference(definition_unfolding,[],[f412,f451])).
fof(f412,axiom,(
  queue_32 = rstore_seq(q12,earray_31)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1638,plain,(
  rselect_head(queue_34) = rselect_head(queue_38)),
  inference(superposition,[],[f9,f641])).
fof(f641,plain,(
  queue_38 = rstore_seq(queue_34,earray_37)),
  inference(definition_unfolding,[],[f424,f452])).
fof(f424,axiom,(
  queue_38 = rstore_seq(q13,earray_37)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1823,plain,(
  index_47 = rselect_head(queue_40)),
  inference(forward_demodulation,[],[f1641,f1598])).
fof(f1598,plain,(
  index_47 = rselect_head(queue_44)),
  inference(forward_demodulation,[],[f1575,f307])).
fof(f307,axiom,(
  index_47 = rselect_head(queue_46)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1575,plain,(
  rselect_head(queue_46) = rselect_head(queue_44)),
  inference(superposition,[],[f8,f427])).
fof(f1641,plain,(
  rselect_head(queue_40) = rselect_head(queue_44)),
  inference(superposition,[],[f9,f642])).
fof(f642,plain,(
  queue_44 = rstore_seq(queue_40,earray_43)),
  inference(definition_unfolding,[],[f426,f453])).
fof(f426,axiom,(
  queue_44 = rstore_seq(q14,earray_43)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2180,plain,(
  index_309 = index_27),
  inference(forward_demodulation,[],[f2164,f959])).
fof(f959,plain,(
  index_309 = s(index_234)),
  inference(backward_demodulation,[],[f956,f253])).
fof(f253,axiom,(
  index_237 = s(index_234)),
  file('SWV568-1.050.tptp',unknown)).
fof(f956,plain,(
  index_309 = index_237),
  inference(forward_demodulation,[],[f872,f583])).
fof(f583,plain,(
  index_309 = rselect_tail(queue_238)),
  inference(definition_unfolding,[],[f285,f481])).
fof(f481,axiom,(
  q4 = queue_238),
  file('SWV568-1.050.tptp',unknown)).
fof(f285,axiom,(
  index_309 = rselect_tail(q4)),
  file('SWV568-1.050.tptp',unknown)).
fof(f872,plain,(
  index_237 = rselect_tail(queue_238)),
  inference(superposition,[],[f5,f382])).
fof(f382,axiom,(
  queue_238 = rstore_tail(queue_236,index_237)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2164,plain,(
  s(index_234) = index_27),
  inference(backward_demodulation,[],[f2156,f2121])).
fof(f2121,plain,(
  index_27 = s(index_156)),
  inference(backward_demodulation,[],[f2119,f267])).
fof(f267,axiom,(
  index_27 = s(index_26)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2119,plain,(
  index_156 = index_26),
  inference(forward_demodulation,[],[f2113,f219])).
fof(f219,axiom,(
  index_156 = s(index_153)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2113,plain,(
  s(index_153) = index_26),
  inference(backward_demodulation,[],[f2102,f1819])).
fof(f1819,plain,(
  index_26 = s(index_327)),
  inference(backward_demodulation,[],[f1814,f1810])).
fof(f1810,plain,(
  index_348 = s(index_327)),
  inference(backward_demodulation,[],[f1805,f301])).
fof(f301,axiom,(
  index_348 = s(index_347)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1805,plain,(
  index_327 = index_347),
  inference(forward_demodulation,[],[f1804,f1801])).
fof(f1801,plain,(
  index_327 = rselect_head(queue_340)),
  inference(backward_demodulation,[],[f1800,f1572])).
fof(f1572,plain,(
  rselect_head(queue_338) = rselect_head(queue_340)),
  inference(superposition,[],[f8,f1002])).
fof(f1002,plain,(
  queue_340 = rstore_tail(queue_338,index_342)),
  inference(backward_demodulation,[],[f1000,f420])).
fof(f1800,plain,(
  index_327 = rselect_head(queue_338)),
  inference(forward_demodulation,[],[f1637,f1799])).
fof(f1799,plain,(
  index_327 = rselect_head(queue_334)),
  inference(backward_demodulation,[],[f1798,f1571])).
fof(f1571,plain,(
  rselect_head(queue_332) = rselect_head(queue_334)),
  inference(superposition,[],[f8,f998])).
fof(f998,plain,(
  queue_334 = rstore_tail(queue_332,index_336)),
  inference(backward_demodulation,[],[f996,f417])).
fof(f1798,plain,(
  index_327 = rselect_head(queue_332)),
  inference(forward_demodulation,[],[f1636,f845])).
fof(f845,plain,(
  index_327 = rselect_head(queue_328)),
  inference(superposition,[],[f4,f415])).
fof(f415,axiom,(
  queue_328 = rstore_head(queue_325,index_327)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1636,plain,(
  rselect_head(queue_328) = rselect_head(queue_332)),
  inference(superposition,[],[f9,f638])).
fof(f638,plain,(
  queue_332 = rstore_seq(queue_328,earray_331)),
  inference(definition_unfolding,[],[f416,f494])).
fof(f494,axiom,(
  q6 = queue_328),
  file('SWV568-1.050.tptp',unknown)).
fof(f416,axiom,(
  queue_332 = rstore_seq(q6,earray_331)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1637,plain,(
  rselect_head(queue_334) = rselect_head(queue_338)),
  inference(superposition,[],[f9,f639])).
fof(f639,plain,(
  queue_338 = rstore_seq(queue_334,earray_337)),
  inference(definition_unfolding,[],[f418,f495])).
fof(f418,axiom,(
  queue_338 = rstore_seq(q7,earray_337)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1804,plain,(
  index_347 = rselect_head(queue_340)),
  inference(forward_demodulation,[],[f1639,f1597])).
fof(f1597,plain,(
  index_347 = rselect_head(queue_344)),
  inference(forward_demodulation,[],[f1573,f300])).
fof(f300,axiom,(
  index_347 = rselect_head(queue_346)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1573,plain,(
  rselect_head(queue_346) = rselect_head(queue_344)),
  inference(superposition,[],[f8,f422])).
fof(f1639,plain,(
  rselect_head(queue_340) = rselect_head(queue_344)),
  inference(superposition,[],[f9,f1493])).
fof(f1493,plain,(
  queue_344 = rstore_seq(queue_340,earray_8)),
  inference(backward_demodulation,[],[f1491,f640])).
fof(f640,plain,(
  queue_344 = rstore_seq(queue_340,earray_343)),
  inference(definition_unfolding,[],[f421,f496])).
fof(f421,axiom,(
  queue_344 = rstore_seq(q8,earray_343)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1491,plain,(
  earray_8 = earray_343),
  inference(backward_demodulation,[],[f1490,f821])).
fof(f821,plain,(
  earray_343 = rselect_seq(queue_344)),
  inference(superposition,[],[f3,f640])).
fof(f1490,plain,(
  earray_8 = rselect_seq(queue_344)),
  inference(forward_demodulation,[],[f1338,f1295])).
fof(f1295,plain,(
  earray_8 = rselect_seq(queue_346)),
  inference(forward_demodulation,[],[f1278,f546])).
fof(f546,plain,(
  earray_8 = rselect_seq(queue_349)),
  inference(definition_unfolding,[],[f189,f497])).
fof(f189,axiom,(
  earray_8 = rselect_seq(q9)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1278,plain,(
  rselect_seq(queue_349) = rselect_seq(queue_346)),
  inference(superposition,[],[f6,f423])).
fof(f423,axiom,(
  queue_349 = rstore_head(queue_346,index_348)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1338,plain,(
  rselect_seq(queue_344) = rselect_seq(queue_346)),
  inference(superposition,[],[f7,f422])).
fof(f1814,plain,(
  index_26 = index_348),
  inference(backward_demodulation,[],[f1813,f846])).
fof(f846,plain,(
  index_348 = rselect_head(queue_349)),
  inference(superposition,[],[f4,f423])).
fof(f1813,plain,(
  index_26 = rselect_head(queue_349)),
  inference(forward_demodulation,[],[f1640,f1703])).
fof(f1703,plain,(
  index_26 = rselect_head(queue_11)),
  inference(backward_demodulation,[],[f1702,f1664])).
fof(f1664,plain,(
  rselect_head(queue_11) = rselect_head(queue_19)),
  inference(backward_demodulation,[],[f1663,f1545])).
fof(f1545,plain,(
  rselect_head(queue_17) = rselect_head(queue_19)),
  inference(superposition,[],[f8,f930])).
fof(f930,plain,(
  queue_19 = rstore_tail(queue_17,index_21)),
  inference(backward_demodulation,[],[f928,f364])).
fof(f1663,plain,(
  rselect_head(queue_11) = rselect_head(queue_17)),
  inference(forward_demodulation,[],[f1606,f1536])).
fof(f1536,plain,(
  rselect_head(queue_11) = rselect_head(queue_13)),
  inference(superposition,[],[f8,f910])).
fof(f910,plain,(
  queue_13 = rstore_tail(queue_11,index_15)),
  inference(backward_demodulation,[],[f908,f342])).
fof(f1606,plain,(
  rselect_head(queue_13) = rselect_head(queue_17)),
  inference(superposition,[],[f9,f612])).
fof(f612,plain,(
  queue_17 = rstore_seq(queue_13,earray_16)),
  inference(definition_unfolding,[],[f357,f449])).
fof(f357,axiom,(
  queue_17 = rstore_seq(q10,earray_16)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1702,plain,(
  index_26 = rselect_head(queue_19)),
  inference(forward_demodulation,[],[f1615,f1591])).
fof(f1591,plain,(
  index_26 = rselect_head(queue_23)),
  inference(forward_demodulation,[],[f1555,f263])).
fof(f263,axiom,(
  index_26 = rselect_head(queue_25)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1555,plain,(
  rselect_head(queue_25) = rselect_head(queue_23)),
  inference(superposition,[],[f8,f386])).
fof(f1615,plain,(
  rselect_head(queue_19) = rselect_head(queue_23)),
  inference(superposition,[],[f9,f622])).
fof(f622,plain,(
  queue_23 = rstore_seq(queue_19,earray_22)),
  inference(definition_unfolding,[],[f379,f450])).
fof(f379,axiom,(
  queue_23 = rstore_seq(q11,earray_22)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1640,plain,(
  rselect_head(queue_349) = rselect_head(queue_11)),
  inference(superposition,[],[f9,f603])).
fof(f603,plain,(
  queue_11 = rstore_seq(queue_349,earray_10)),
  inference(definition_unfolding,[],[f335,f497])).
fof(f335,axiom,(
  queue_11 = rstore_seq(q9,earray_10)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2102,plain,(
  index_153 = index_327),
  inference(forward_demodulation,[],[f2098,f1028])).
fof(f1028,plain,(
  index_153 = s(index_78)),
  inference(backward_demodulation,[],[f1024,f321])).
fof(f321,axiom,(
  index_81 = s(index_78)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1024,plain,(
  index_153 = index_81),
  inference(forward_demodulation,[],[f897,f558])).
fof(f558,plain,(
  index_153 = rselect_tail(queue_82)),
  inference(definition_unfolding,[],[f218,f459])).
fof(f459,axiom,(
  q2 = queue_82),
  file('SWV568-1.050.tptp',unknown)).
fof(f218,axiom,(
  index_153 = rselect_tail(q2)),
  file('SWV568-1.050.tptp',unknown)).
fof(f897,plain,(
  index_81 = rselect_tail(queue_82)),
  inference(superposition,[],[f5,f441])).
fof(f441,axiom,(
  queue_82 = rstore_tail(queue_80,index_81)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2098,plain,(
  index_327 = s(index_78)),
  inference(backward_demodulation,[],[f2087,f1795])).
fof(f1795,plain,(
  index_327 = s(index_159)),
  inference(backward_demodulation,[],[f1790,f292])).
fof(f292,axiom,(
  index_327 = s(index_326)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1790,plain,(
  index_159 = index_326),
  inference(forward_demodulation,[],[f1789,f1735])).
fof(f1735,plain,(
  index_159 = rselect_head(queue_313)),
  inference(backward_demodulation,[],[f1734,f1567])).
fof(f1567,plain,(
  rselect_head(queue_311) = rselect_head(queue_313)),
  inference(superposition,[],[f8,f990])).
fof(f990,plain,(
  queue_313 = rstore_tail(queue_311,index_321)),
  inference(backward_demodulation,[],[f988,f409])).
fof(f1734,plain,(
  index_159 = rselect_head(queue_311)),
  inference(forward_demodulation,[],[f1623,f1688])).
fof(f1688,plain,(
  index_159 = rselect_head(queue_238)),
  inference(backward_demodulation,[],[f1687,f1556])).
fof(f1556,plain,(
  rselect_head(queue_236) = rselect_head(queue_238)),
  inference(superposition,[],[f8,f958])).
fof(f958,plain,(
  queue_238 = rstore_tail(queue_236,index_309)),
  inference(backward_demodulation,[],[f956,f382])).
fof(f1687,plain,(
  index_159 = rselect_head(queue_236)),
  inference(forward_demodulation,[],[f1611,f836])).
fof(f836,plain,(
  index_159 = rselect_head(queue_160)),
  inference(superposition,[],[f4,f353])).
fof(f353,axiom,(
  queue_160 = rstore_head(queue_157,index_159)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1611,plain,(
  rselect_head(queue_160) = rselect_head(queue_236)),
  inference(superposition,[],[f9,f623])).
fof(f623,plain,(
  queue_236 = rstore_seq(queue_160,earray_235)),
  inference(definition_unfolding,[],[f381,f470])).
fof(f470,axiom,(
  q3 = queue_160),
  file('SWV568-1.050.tptp',unknown)).
fof(f381,axiom,(
  queue_236 = rstore_seq(q3,earray_235)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1623,plain,(
  rselect_head(queue_238) = rselect_head(queue_311)),
  inference(superposition,[],[f9,f634])).
fof(f634,plain,(
  queue_311 = rstore_seq(queue_238,earray_310)),
  inference(definition_unfolding,[],[f408,f481])).
fof(f408,axiom,(
  queue_311 = rstore_seq(q4,earray_310)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1789,plain,(
  index_326 = rselect_head(queue_313)),
  inference(forward_demodulation,[],[f1635,f1596])).
fof(f1596,plain,(
  index_326 = rselect_head(queue_323)),
  inference(forward_demodulation,[],[f1570,f291])).
fof(f291,axiom,(
  index_326 = rselect_head(queue_325)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1570,plain,(
  rselect_head(queue_325) = rselect_head(queue_323)),
  inference(superposition,[],[f8,f414])).
fof(f414,axiom,(
  queue_325 = rstore_tail(queue_323,index_324)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1635,plain,(
  rselect_head(queue_313) = rselect_head(queue_323)),
  inference(superposition,[],[f9,f637])).
fof(f637,plain,(
  queue_323 = rstore_seq(queue_313,earray_322)),
  inference(definition_unfolding,[],[f413,f492])).
fof(f413,axiom,(
  queue_323 = rstore_seq(q5,earray_322)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2087,plain,(
  index_78 = index_159),
  inference(forward_demodulation,[],[f2085,f1011])).
fof(f1011,plain,(
  index_78 = s(index_3)),
  inference(backward_demodulation,[],[f1008,f312])).
fof(f312,axiom,(
  index_6 = s(index_3)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1008,plain,(
  index_78 = index_6),
  inference(forward_demodulation,[],[f892,f596])).
fof(f596,plain,(
  index_78 = rselect_tail(queue_7)),
  inference(definition_unfolding,[],[f320,f448])).
fof(f448,axiom,(
  q1 = queue_7),
  file('SWV568-1.050.tptp',unknown)).
fof(f320,axiom,(
  index_78 = rselect_tail(q1)),
  file('SWV568-1.050.tptp',unknown)).
fof(f892,plain,(
  index_6 = rselect_tail(queue_7)),
  inference(superposition,[],[f5,f436])).
fof(f436,axiom,(
  queue_7 = rstore_tail(queue_5,index_6)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2085,plain,(
  index_159 = s(index_3)),
  inference(backward_demodulation,[],[f2073,f1856])).
fof(f1856,plain,(
  index_159 = s(index_0)),
  inference(backward_demodulation,[],[f1852,f221])).
fof(f221,axiom,(
  index_159 = s(index_158)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1852,plain,(
  index_0 = index_158),
  inference(backward_demodulation,[],[f1851,f1586])).
fof(f1586,plain,(
  index_158 = rselect_head(queue_155)),
  inference(forward_demodulation,[],[f1543,f220])).
fof(f220,axiom,(
  index_158 = rselect_head(queue_157)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1543,plain,(
  rselect_head(queue_157) = rselect_head(queue_155)),
  inference(superposition,[],[f8,f352])).
fof(f352,axiom,(
  queue_157 = rstore_tail(queue_155,index_156)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1851,plain,(
  index_0 = rselect_head(queue_155)),
  inference(forward_demodulation,[],[f1648,f1846])).
fof(f1846,plain,(
  index_0 = rselect_head(queue_82)),
  inference(backward_demodulation,[],[f1845,f1581])).
fof(f1581,plain,(
  rselect_head(queue_80) = rselect_head(queue_82)),
  inference(superposition,[],[f8,f1027])).
fof(f1027,plain,(
  queue_82 = rstore_tail(queue_80,index_153)),
  inference(backward_demodulation,[],[f1024,f441])).
fof(f1845,plain,(
  index_0 = rselect_head(queue_80)),
  inference(forward_demodulation,[],[f1645,f1652])).
fof(f1652,plain,(
  index_0 = rselect_head(queue_7)),
  inference(backward_demodulation,[],[f1651,f1576])).
fof(f1576,plain,(
  rselect_head(queue_5) = rselect_head(queue_7)),
  inference(superposition,[],[f8,f1010])).
fof(f1010,plain,(
  queue_7 = rstore_tail(queue_5,index_78)),
  inference(backward_demodulation,[],[f1008,f436])).
fof(f1651,plain,(
  index_0 = rselect_head(queue_5)),
  inference(forward_demodulation,[],[f1601,f833])).
fof(f833,plain,(
  index_0 = rselect_head(queue_1)),
  inference(superposition,[],[f4,f330])).
fof(f330,axiom,(
  queue_1 = rstore_head(q,index_0)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1601,plain,(
  rselect_head(queue_1) = rselect_head(queue_5)),
  inference(superposition,[],[f9,f643])).
fof(f643,plain,(
  queue_5 = rstore_seq(queue_1,earray_4)),
  inference(definition_unfolding,[],[f429,f447])).
fof(f447,axiom,(
  q0 = queue_1),
  file('SWV568-1.050.tptp',unknown)).
fof(f429,axiom,(
  queue_5 = rstore_seq(q0,earray_4)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1645,plain,(
  rselect_head(queue_7) = rselect_head(queue_80)),
  inference(superposition,[],[f9,f1523])).
fof(f1523,plain,(
  queue_80 = rstore_seq(queue_7,earray_152)),
  inference(backward_demodulation,[],[f1520,f648])).
fof(f648,plain,(
  queue_80 = rstore_seq(queue_7,earray_79)),
  inference(definition_unfolding,[],[f440,f448])).
fof(f440,axiom,(
  queue_80 = rstore_seq(q1,earray_79)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1520,plain,(
  earray_152 = earray_79),
  inference(forward_demodulation,[],[f1519,f507])).
fof(f507,plain,(
  earray_152 = rselect_seq(queue_82)),
  inference(definition_unfolding,[],[f111,f459])).
fof(f111,axiom,(
  earray_152 = rselect_seq(q2)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1519,plain,(
  earray_79 = rselect_seq(queue_82)),
  inference(forward_demodulation,[],[f1346,f827])).
fof(f827,plain,(
  earray_79 = rselect_seq(queue_80)),
  inference(superposition,[],[f3,f648])).
fof(f1346,plain,(
  rselect_seq(queue_82) = rselect_seq(queue_80)),
  inference(superposition,[],[f7,f1027])).
fof(f1648,plain,(
  rselect_head(queue_155) = rselect_head(queue_82)),
  inference(superposition,[],[f9,f610])).
fof(f610,plain,(
  queue_155 = rstore_seq(queue_82,earray_154)),
  inference(definition_unfolding,[],[f351,f459])).
fof(f351,axiom,(
  queue_155 = rstore_seq(q2,earray_154)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2073,plain,(
  index_3 = index_0),
  inference(backward_demodulation,[],[f2072,f196])).
fof(f196,axiom,(
  index_0 = rselect_tail(q)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2072,plain,(
  index_3 = rselect_tail(q)),
  inference(forward_demodulation,[],[f2055,f580])).
fof(f580,plain,(
  index_3 = rselect_tail(queue_1)),
  inference(definition_unfolding,[],[f280,f447])).
fof(f280,axiom,(
  index_3 = rselect_tail(q0)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2055,plain,(
  rselect_tail(q) = rselect_tail(queue_1)),
  inference(superposition,[],[f11,f330])).
fof(f2156,plain,(
  index_234 = index_156),
  inference(forward_demodulation,[],[f2155,f570])).
fof(f570,plain,(
  index_234 = rselect_tail(queue_160)),
  inference(definition_unfolding,[],[f252,f470])).
fof(f252,axiom,(
  index_234 = rselect_tail(q3)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2155,plain,(
  index_156 = rselect_tail(queue_160)),
  inference(forward_demodulation,[],[f2058,f859])).
fof(f859,plain,(
  index_156 = rselect_tail(queue_157)),
  inference(superposition,[],[f5,f352])).
fof(f2058,plain,(
  rselect_tail(queue_160) = rselect_tail(queue_157)),
  inference(superposition,[],[f11,f353])).
fof(f2297,plain,(
  index_330 = index_324),
  inference(forward_demodulation,[],[f2296,f586])).
fof(f586,plain,(
  index_330 = rselect_tail(queue_328)),
  inference(definition_unfolding,[],[f294,f494])).
fof(f294,axiom,(
  index_330 = rselect_tail(q6)),
  file('SWV568-1.050.tptp',unknown)).
fof(f2296,plain,(
  index_324 = rselect_tail(queue_328)),
  inference(forward_demodulation,[],[f2067,f886])).
fof(f886,plain,(
  index_324 = rselect_tail(queue_325)),
  inference(superposition,[],[f5,f414])).
fof(f2067,plain,(
  rselect_tail(queue_328) = rselect_tail(queue_325)),
  inference(superposition,[],[f11,f415])).
fof(f2068,plain,(
  rselect_tail(queue_349) = rselect_tail(queue_346)),
  inference(superposition,[],[f11,f1815])).
fof(f1815,plain,(
  queue_349 = rstore_head(queue_346,index_26)),
  inference(backward_demodulation,[],[f1814,f423])).
fof(f2069,plain,(
  rselect_tail(queue_49) = rselect_tail(queue_46)),
  inference(superposition,[],[f11,f428])).
fof(f4520,plain,(
  ( ! [X36] : (select(earray_304,X36) = select(earray_316,X36) | index_315 = X36) )),
  inference(superposition,[],[f2,f1469])).
fof(f1469,plain,(
  earray_316 = store(earray_304,index_315,e49)),
  inference(backward_demodulation,[],[f1468,f164])).
fof(f164,axiom,(
  earray_316 = store(earray_314,index_315,e49)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1468,plain,(
  earray_304 = earray_314),
  inference(backward_demodulation,[],[f1467,f533])).
fof(f533,plain,(
  earray_314 = rselect_seq(queue_307)),
  inference(definition_unfolding,[],[f163,f491])).
fof(f163,axiom,(
  earray_314 = rselect_seq(q49)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1467,plain,(
  earray_304 = rselect_seq(queue_307)),
  inference(forward_demodulation,[],[f1331,f815])).
fof(f815,plain,(
  earray_304 = rselect_seq(queue_305)),
  inference(superposition,[],[f3,f633])).
fof(f1331,plain,(
  rselect_seq(queue_307) = rselect_seq(queue_305)),
  inference(superposition,[],[f7,f986])).
fof(f86673,plain,(
  e16 = select(earray_100,index_57)),
  inference(subsumption_resolution,[],[f86671,f3189])).
fof(f3189,plain,(
  index_99 != index_57),
  inference(forward_demodulation,[],[f3188,f2517])).
fof(f3188,plain,(
  index_57 != s(index_90)),
  inference(forward_demodulation,[],[f3187,f1033])).
fof(f3187,plain,(
  index_57 != s(s(index_84))),
  inference(forward_demodulation,[],[f3186,f1023])).
fof(f3186,plain,(
  index_57 != s(s(s(index_72)))),
  inference(forward_demodulation,[],[f3049,f2508])).
fof(f3049,plain,(
  index_57 != s(s(s(s(index_63))))),
  inference(superposition,[],[f89,f1019])).
fof(f89,axiom,(
  ( ! [X4] : (s(s(s(s(s(X4))))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86671,plain,(
  e16 = select(earray_100,index_57) | index_99 = index_57),
  inference(superposition,[],[f86670,f4487])).
fof(f4487,plain,(
  ( ! [X3] : (select(earray_98,X3) = select(earray_100,X3) | index_99 = X3) )),
  inference(superposition,[],[f2,f95])).
fof(f95,axiom,(
  earray_100 = store(earray_98,index_99,e21)),
  file('SWV568-1.050.tptp',unknown)).
fof(f86670,plain,(
  e16 = select(earray_98,index_57)),
  inference(subsumption_resolution,[],[f86669,f1122])).
fof(f1122,plain,(
  index_57 != index_63),
  inference(superposition,[],[f695,f1016])).
fof(f695,plain,(
  index_57 != index_60),
  inference(superposition,[],[f93,f313])).
fof(f93,axiom,(
  ( ! [X4] : (s(X4) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86669,plain,(
  e16 = select(earray_98,index_57) | index_57 = index_63),
  inference(subsumption_resolution,[],[f86668,f85777])).
fof(f85777,plain,(
  index_57 != index_72),
  inference(superposition,[],[f1017,f2508])).
fof(f1017,plain,(
  index_57 != s(index_63)),
  inference(backward_demodulation,[],[f1016,f761])).
fof(f761,plain,(
  index_57 != s(index_60)),
  inference(superposition,[],[f92,f313])).
fof(f92,axiom,(
  ( ! [X4] : (s(s(X4)) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86668,plain,(
  e16 = select(earray_98,index_57) | index_57 = index_72 | index_57 = index_63),
  inference(subsumption_resolution,[],[f86667,f85831])).
fof(f85831,plain,(
  index_57 != index_84),
  inference(superposition,[],[f2513,f1023])).
fof(f2513,plain,(
  index_57 != s(index_72)),
  inference(backward_demodulation,[],[f2507,f1261])).
fof(f1261,plain,(
  index_57 != s(index_66)),
  inference(forward_demodulation,[],[f1210,f315])).
fof(f1210,plain,(
  index_57 != s(s(index_63))),
  inference(superposition,[],[f91,f1019])).
fof(f91,axiom,(
  ( ! [X4] : (s(s(s(X4))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86667,plain,(
  e16 = select(earray_98,index_57) | index_57 = index_84 | index_57 = index_72 | index_57 = index_63),
  inference(subsumption_resolution,[],[f86663,f85869])).
fof(f85869,plain,(
  index_57 != index_90),
  inference(superposition,[],[f2939,f1033])).
fof(f2939,plain,(
  index_57 != s(index_84)),
  inference(forward_demodulation,[],[f2938,f1023])).
fof(f2938,plain,(
  index_57 != s(s(index_72))),
  inference(forward_demodulation,[],[f2842,f2508])).
fof(f2842,plain,(
  index_57 != s(s(s(index_63)))),
  inference(superposition,[],[f90,f1019])).
fof(f90,axiom,(
  ( ! [X4] : (s(s(s(s(X4)))) != X4) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86663,plain,(
  e16 = select(earray_98,index_57) | index_57 = index_90 | index_57 = index_84 | index_57 = index_72 | index_57 = index_63),
  inference(superposition,[],[f86643,f3000])).
fof(f3000,plain,(
  e16 = select(earray_58,index_57)),
  inference(superposition,[],[f1,f1507])).
fof(f1507,plain,(
  earray_58 = store(earray_52,index_57,e16)),
  inference(backward_demodulation,[],[f1506,f182])).
fof(f182,axiom,(
  earray_58 = store(earray_56,index_57,e16)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1506,plain,(
  earray_52 = earray_56),
  inference(backward_demodulation,[],[f1505,f542])).
fof(f542,plain,(
  earray_56 = rselect_seq(queue_55)),
  inference(definition_unfolding,[],[f181,f455])).
fof(f181,axiom,(
  earray_56 = rselect_seq(q16)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1505,plain,(
  earray_52 = rselect_seq(queue_55)),
  inference(forward_demodulation,[],[f1342,f824])).
fof(f824,plain,(
  earray_52 = rselect_seq(queue_53)),
  inference(superposition,[],[f3,f644])).
fof(f1342,plain,(
  rselect_seq(queue_55) = rselect_seq(queue_53)),
  inference(superposition,[],[f7,f1014])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (select(store(X0,X1,X2),X1) = X2) )),
  file('SWV568-1.050.tptp',unknown)).
fof(f86643,plain,(
  ( ! [X1] : (select(earray_98,X1) = select(earray_58,X1) | index_90 = X1 | index_84 = X1 | index_72 = X1 | index_63 = X1) )),
  inference(superposition,[],[f86624,f4530])).
fof(f4530,plain,(
  ( ! [X46] : (select(earray_58,X46) = select(earray_64,X46) | index_63 = X46) )),
  inference(superposition,[],[f2,f1510])).
fof(f1510,plain,(
  earray_64 = store(earray_58,index_63,e17)),
  inference(backward_demodulation,[],[f1509,f184])).
fof(f184,axiom,(
  earray_64 = store(earray_62,index_63,e17)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1509,plain,(
  earray_58 = earray_62),
  inference(backward_demodulation,[],[f1508,f543])).
fof(f543,plain,(
  earray_62 = rselect_seq(queue_61)),
  inference(definition_unfolding,[],[f183,f456])).
fof(f183,axiom,(
  earray_62 = rselect_seq(q17)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1508,plain,(
  earray_58 = rselect_seq(queue_61)),
  inference(forward_demodulation,[],[f1343,f825])).
fof(f825,plain,(
  earray_58 = rselect_seq(queue_59)),
  inference(superposition,[],[f3,f645])).
fof(f1343,plain,(
  rselect_seq(queue_61) = rselect_seq(queue_59)),
  inference(superposition,[],[f7,f1018])).
fof(f86624,plain,(
  ( ! [X1] : (select(earray_98,X1) = select(earray_64,X1) | index_90 = X1 | index_84 = X1 | index_72 = X1) )),
  inference(superposition,[],[f86556,f4531])).
fof(f4531,plain,(
  ( ! [X47] : (select(earray_64,X47) = select(earray_73,X47) | index_72 = X47) )),
  inference(superposition,[],[f2,f1514])).
fof(f1514,plain,(
  earray_73 = store(earray_64,index_72,e18)),
  inference(backward_demodulation,[],[f1512,f186])).
fof(f186,axiom,(
  earray_73 = store(earray_71,index_72,e18)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1512,plain,(
  earray_64 = earray_71),
  inference(forward_demodulation,[],[f1511,f826])).
fof(f826,plain,(
  earray_64 = rselect_seq(queue_65)),
  inference(superposition,[],[f3,f646])).
fof(f1511,plain,(
  earray_71 = rselect_seq(queue_65)),
  inference(forward_demodulation,[],[f1344,f1297])).
fof(f1297,plain,(
  earray_71 = rselect_seq(queue_67)),
  inference(forward_demodulation,[],[f1280,f544])).
fof(f544,plain,(
  earray_71 = rselect_seq(queue_70)),
  inference(definition_unfolding,[],[f185,f457])).
fof(f185,axiom,(
  earray_71 = rselect_seq(q18)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1280,plain,(
  rselect_seq(queue_70) = rselect_seq(queue_67)),
  inference(superposition,[],[f6,f437])).
fof(f1344,plain,(
  rselect_seq(queue_65) = rselect_seq(queue_67)),
  inference(superposition,[],[f7,f435])).
fof(f86556,plain,(
  ( ! [X1] : (select(earray_73,X1) = select(earray_98,X1) | index_90 = X1 | index_84 = X1) )),
  inference(superposition,[],[f4533,f4532])).
fof(f4532,plain,(
  ( ! [X48] : (select(earray_73,X48) = select(earray_85,X48) | index_84 = X48) )),
  inference(superposition,[],[f2,f1518])).
fof(f1518,plain,(
  earray_85 = store(earray_73,index_84,e19)),
  inference(backward_demodulation,[],[f1517,f191])).
fof(f191,axiom,(
  earray_85 = store(earray_83,index_84,e19)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1517,plain,(
  earray_73 = earray_83),
  inference(backward_demodulation,[],[f1516,f547])).
fof(f547,plain,(
  earray_83 = rselect_seq(queue_76)),
  inference(definition_unfolding,[],[f190,f458])).
fof(f190,axiom,(
  earray_83 = rselect_seq(q19)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1516,plain,(
  earray_73 = rselect_seq(queue_76)),
  inference(forward_demodulation,[],[f1345,f828])).
fof(f828,plain,(
  earray_73 = rselect_seq(queue_74)),
  inference(superposition,[],[f3,f647])).
fof(f1345,plain,(
  rselect_seq(queue_76) = rselect_seq(queue_74)),
  inference(superposition,[],[f7,f1022])).
fof(f4533,plain,(
  ( ! [X49] : (select(earray_85,X49) = select(earray_98,X49) | index_90 = X49) )),
  inference(superposition,[],[f2,f1529])).
fof(f1529,plain,(
  earray_98 = store(earray_85,index_90,e20)),
  inference(backward_demodulation,[],[f1528,f1526])).
fof(f1526,plain,(
  earray_91 = store(earray_85,index_90,e20)),
  inference(backward_demodulation,[],[f1525,f193])).
fof(f193,axiom,(
  earray_91 = store(earray_89,index_90,e20)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1525,plain,(
  earray_85 = earray_89),
  inference(backward_demodulation,[],[f1524,f548])).
fof(f548,plain,(
  earray_89 = rselect_seq(queue_88)),
  inference(definition_unfolding,[],[f192,f460])).
fof(f192,axiom,(
  earray_89 = rselect_seq(q20)),
  file('SWV568-1.050.tptp',unknown)).
fof(f1524,plain,(
  earray_85 = rselect_seq(queue_88)),
  inference(forward_demodulation,[],[f1347,f829])).
fof(f829,plain,(
  earray_85 = rselect_seq(queue_86)),
  inference(superposition,[],[f3,f649])).
fof(f1347,plain,(
  rselect_seq(queue_88) = rselect_seq(queue_86)),
  inference(superposition,[],[f7,f1032])).
% SZS output end Proof for SWV568-1.050
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 34668
% Time elapsed: 2.132 s
% ------------------------------
% ------------------------------
