% Refutation found. Thanks to Tanya!
% SZS status Theorem for SYN448+1
% SZS output start Proof for SYN448+1
fof(f4212,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f230,f252,f270,f279,f321,f326,f338,f349,f361,f368,f369,f373,f378,f382,f384,f392,f396,f401,f405,f406,f407,f411,f417,f423,f424,f432,f436,f440,f441,f449,f454,f458,f459,f471,f476,f477,f478,f479,f483,f484,f485,f486,f487,f492,f497,f502,f508,f513,f518,f524,f529,f534,f540,f545,f550,f572,f577,f582,f604,f609,f614,f620,f625,f630,f636,f641,f646,f652,f657,f662,f668,f673,f678,f684,f689,f694,f705,f710,f716,f721,f726,f748,f753,f758,f764,f769,f774,f775,f780,f785,f790,f796,f801,f806,f844,f849,f854,f860,f865,f870,f871,f892,f897,f902,f908,f913,f918,f924,f934,f940,f945,f950,f985,f1048,f1105,f1107,f1109,f1154,f1156,f1216,f1240,f1554,f1560,f1651,f1668,f1816,f1832,f2022,f2026,f2140,f2191,f2220,f2226,f2310,f2350,f2354,f2356,f2362,f2384,f2548,f2624,f2673,f2697,f2713,f2767,f2845,f2855,f2882,f2886,f2888,f2891,f2914,f2925,f2971,f2994,f3014,f3101,f3109,f3119,f3121,f3131,f3203,f3205,f3228,f3287,f3326,f3374,f3414,f3445,f3493,f3497,f3501,f3637,f3895,f3902,f4050,f4051,f4105,f4113,f4185,f4206])).
fof(f4206,plain,(
  ~spl0_44 | ~spl0_51 | spl0_136 | ~spl0_137),
  inference(avatar_contradiction_clause,[],[f4205])).
fof(f4205,plain,(
  $false | (~spl0_44 | ~spl0_51 | spl0_136 | ~spl0_137)),
  inference(subsumption_resolution,[],[f4189,f891])).
fof(f891,plain,(
  ~c0_1(a466) | spl0_136),
  inference(avatar_component_clause,[],[f889])).
fof(f889,plain,(
  spl0_136 <=> c0_1(a466)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_136])])).
fof(f4189,plain,(
  c0_1(a466) | (~spl0_44 | ~spl0_51 | ~spl0_137)),
  inference(resolution,[],[f4187,f896])).
fof(f896,plain,(
  c3_1(a466) | ~spl0_137),
  inference(avatar_component_clause,[],[f894])).
fof(f894,plain,(
  spl0_137 <=> c3_1(a466)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_137])])).
fof(f4187,plain,(
  ( ! [X52] : (~c3_1(X52) | c0_1(X52)) ) | (~spl0_44 | ~spl0_51)),
  inference(subsumption_resolution,[],[f439,f404])).
fof(f404,plain,(
  ( ! [X28] : (~c3_1(X28) | c0_1(X28) | ~c2_1(X28)) ) | ~spl0_44),
  inference(avatar_component_clause,[],[f403])).
fof(f403,plain,(
  spl0_44 <=> ! [X28] : (~c3_1(X28) | c0_1(X28) | ~c2_1(X28))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_44])])).
fof(f439,plain,(
  ( ! [X52] : (~c3_1(X52) | c0_1(X52) | c2_1(X52)) ) | ~spl0_51),
  inference(avatar_component_clause,[],[f438])).
fof(f438,plain,(
  spl0_51 <=> ! [X52] : (~c3_1(X52) | c0_1(X52) | c2_1(X52))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_51])])).
fof(f4185,plain,(
  ~spl0_59 | spl0_115 | ~spl0_116 | ~spl0_117),
  inference(avatar_contradiction_clause,[],[f4184])).
fof(f4184,plain,(
  $false | (~spl0_59 | spl0_115 | ~spl0_116 | ~spl0_117)),
  inference(subsumption_resolution,[],[f4183,f784])).
fof(f784,plain,(
  c2_1(a477) | ~spl0_116),
  inference(avatar_component_clause,[],[f782])).
fof(f782,plain,(
  spl0_116 <=> c2_1(a477)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_116])])).
fof(f4183,plain,(
  ~c2_1(a477) | (~spl0_59 | spl0_115 | ~spl0_117)),
  inference(subsumption_resolution,[],[f4174,f779])).
fof(f779,plain,(
  ~c3_1(a477) | spl0_115),
  inference(avatar_component_clause,[],[f777])).
fof(f777,plain,(
  spl0_115 <=> c3_1(a477)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_115])])).
fof(f4174,plain,(
  c3_1(a477) | ~c2_1(a477) | (~spl0_59 | ~spl0_117)),
  inference(resolution,[],[f475,f789])).
fof(f789,plain,(
  c1_1(a477) | ~spl0_117),
  inference(avatar_component_clause,[],[f787])).
fof(f787,plain,(
  spl0_117 <=> c1_1(a477)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_117])])).
fof(f475,plain,(
  ( ! [X73] : (~c1_1(X73) | c3_1(X73) | ~c2_1(X73)) ) | ~spl0_59),
  inference(avatar_component_clause,[],[f474])).
fof(f474,plain,(
  spl0_59 <=> ! [X73] : (~c2_1(X73) | c3_1(X73) | ~c1_1(X73))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_59])])).
fof(f4113,plain,(
  ~spl0_28 | ~spl0_64 | ~spl0_66 | spl0_160),
  inference(avatar_contradiction_clause,[],[f4112])).
fof(f4112,plain,(
  $false | (~spl0_28 | ~spl0_64 | ~spl0_66 | spl0_160)),
  inference(subsumption_resolution,[],[f4111,f517])).
fof(f517,plain,(
  c0_1(a473) | ~spl0_66),
  inference(avatar_component_clause,[],[f515])).
fof(f515,plain,(
  spl0_66 <=> c0_1(a473)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_66])])).
fof(f4111,plain,(
  ~c0_1(a473) | (~spl0_28 | ~spl0_64 | spl0_160)),
  inference(subsumption_resolution,[],[f4101,f2848])).
fof(f2848,plain,(
  ~c2_1(a473) | spl0_160),
  inference(avatar_component_clause,[],[f2847])).
fof(f2847,plain,(
  spl0_160 <=> c2_1(a473)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_160])])).
fof(f4101,plain,(
  c2_1(a473) | ~c0_1(a473) | (~spl0_28 | ~spl0_64)),
  inference(resolution,[],[f337,f507])).
fof(f507,plain,(
  c3_1(a473) | ~spl0_64),
  inference(avatar_component_clause,[],[f505])).
fof(f505,plain,(
  spl0_64 <=> c3_1(a473)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_64])])).
fof(f337,plain,(
  ( ! [X5] : (~c3_1(X5) | c2_1(X5) | ~c0_1(X5)) ) | ~spl0_28),
  inference(avatar_component_clause,[],[f336])).
fof(f336,plain,(
  spl0_28 <=> ! [X5] : (~c3_1(X5) | c2_1(X5) | ~c0_1(X5))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_28])])).
fof(f4105,plain,(
  ~spl0_28 | spl0_130 | ~spl0_131 | ~spl0_132),
  inference(avatar_contradiction_clause,[],[f4104])).
fof(f4104,plain,(
  $false | (~spl0_28 | spl0_130 | ~spl0_131 | ~spl0_132)),
  inference(subsumption_resolution,[],[f4103,f869])).
fof(f869,plain,(
  c0_1(a468) | ~spl0_132),
  inference(avatar_component_clause,[],[f867])).
fof(f867,plain,(
  spl0_132 <=> c0_1(a468)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_132])])).
fof(f4103,plain,(
  ~c0_1(a468) | (~spl0_28 | spl0_130 | ~spl0_131)),
  inference(subsumption_resolution,[],[f4090,f859])).
fof(f859,plain,(
  ~c2_1(a468) | spl0_130),
  inference(avatar_component_clause,[],[f857])).
fof(f857,plain,(
  spl0_130 <=> c2_1(a468)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_130])])).
fof(f4090,plain,(
  c2_1(a468) | ~c0_1(a468) | (~spl0_28 | ~spl0_131)),
  inference(resolution,[],[f337,f864])).
fof(f864,plain,(
  c3_1(a468) | ~spl0_131),
  inference(avatar_component_clause,[],[f862])).
fof(f862,plain,(
  spl0_131 <=> c3_1(a468)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_131])])).
fof(f4051,plain,(
  ~spl0_160 | ~spl0_32 | ~spl0_65 | ~spl0_66),
  inference(avatar_split_clause,[],[f4044,f515,f510,f351,f2847])).
fof(f351,plain,(
  spl0_32 <=> ! [X9] : (~c2_1(X9) | ~c0_1(X9) | ~c1_1(X9))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_32])])).
fof(f510,plain,(
  spl0_65 <=> c1_1(a473)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_65])])).
fof(f4044,plain,(
  ~c2_1(a473) | (~spl0_32 | ~spl0_65 | ~spl0_66)),
  inference(subsumption_resolution,[],[f4026,f517])).
fof(f4026,plain,(
  ~c0_1(a473) | ~c2_1(a473) | (~spl0_32 | ~spl0_65)),
  inference(resolution,[],[f352,f512])).
fof(f512,plain,(
  c1_1(a473) | ~spl0_65),
  inference(avatar_component_clause,[],[f510])).
fof(f352,plain,(
  ( ! [X9] : (~c1_1(X9) | ~c0_1(X9) | ~c2_1(X9)) ) | ~spl0_32),
  inference(avatar_component_clause,[],[f351])).
fof(f4050,plain,(
  ~spl0_61 | ~spl0_32 | ~spl0_62 | ~spl0_63),
  inference(avatar_split_clause,[],[f4047,f499,f494,f351,f489])).
fof(f489,plain,(
  spl0_61 <=> c2_1(a490)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_61])])).
fof(f494,plain,(
  spl0_62 <=> c1_1(a490)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_62])])).
fof(f499,plain,(
  spl0_63 <=> c0_1(a490)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_63])])).
fof(f4047,plain,(
  ~c2_1(a490) | (~spl0_32 | ~spl0_62 | ~spl0_63)),
  inference(subsumption_resolution,[],[f4027,f501])).
fof(f501,plain,(
  c0_1(a490) | ~spl0_63),
  inference(avatar_component_clause,[],[f499])).
fof(f4027,plain,(
  ~c0_1(a490) | ~c2_1(a490) | (~spl0_32 | ~spl0_62)),
  inference(resolution,[],[f352,f496])).
fof(f496,plain,(
  c1_1(a490) | ~spl0_62),
  inference(avatar_component_clause,[],[f494])).
fof(f3902,plain,(
  spl0_127 | ~spl0_37 | spl0_128 | ~spl0_129),
  inference(avatar_split_clause,[],[f3901,f851,f846,f371,f841])).
fof(f841,plain,(
  spl0_127 <=> c3_1(a471)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_127])])).
fof(f371,plain,(
  spl0_37 <=> ! [X16] : (~c2_1(X16) | c1_1(X16) | c3_1(X16))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_37])])).
fof(f846,plain,(
  spl0_128 <=> c1_1(a471)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_128])])).
fof(f851,plain,(
  spl0_129 <=> c2_1(a471)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_129])])).
fof(f3901,plain,(
  c3_1(a471) | (~spl0_37 | spl0_128 | ~spl0_129)),
  inference(subsumption_resolution,[],[f3809,f848])).
fof(f848,plain,(
  ~c1_1(a471) | spl0_128),
  inference(avatar_component_clause,[],[f846])).
fof(f3809,plain,(
  c1_1(a471) | c3_1(a471) | (~spl0_37 | ~spl0_129)),
  inference(resolution,[],[f372,f853])).
fof(f853,plain,(
  c2_1(a471) | ~spl0_129),
  inference(avatar_component_clause,[],[f851])).
fof(f372,plain,(
  ( ! [X16] : (~c2_1(X16) | c1_1(X16) | c3_1(X16)) ) | ~spl0_37),
  inference(avatar_component_clause,[],[f371])).
fof(f3895,plain,(
  spl0_113 | ~spl0_48 | spl0_112 | ~spl0_114),
  inference(avatar_split_clause,[],[f3894,f771,f761,f426,f766])).
fof(f766,plain,(
  spl0_113 <=> c0_1(a478)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_113])])).
fof(f426,plain,(
  spl0_48 <=> ! [X48] : (~c2_1(X48) | c0_1(X48) | c3_1(X48))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_48])])).
fof(f761,plain,(
  spl0_112 <=> c3_1(a478)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_112])])).
fof(f771,plain,(
  spl0_114 <=> c2_1(a478)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_114])])).
fof(f3894,plain,(
  c0_1(a478) | (~spl0_48 | spl0_112 | ~spl0_114)),
  inference(subsumption_resolution,[],[f3857,f763])).
fof(f763,plain,(
  ~c3_1(a478) | spl0_112),
  inference(avatar_component_clause,[],[f761])).
fof(f3857,plain,(
  c0_1(a478) | c3_1(a478) | (~spl0_48 | ~spl0_114)),
  inference(resolution,[],[f427,f773])).
fof(f773,plain,(
  c2_1(a478) | ~spl0_114),
  inference(avatar_component_clause,[],[f771])).
fof(f427,plain,(
  ( ! [X48] : (~c2_1(X48) | c0_1(X48) | c3_1(X48)) ) | ~spl0_48),
  inference(avatar_component_clause,[],[f426])).
fof(f3637,plain,(
  ~spl0_42 | ~spl0_51 | spl0_109 | spl0_110 | spl0_111),
  inference(avatar_contradiction_clause,[],[f3636])).
fof(f3636,plain,(
  $false | (~spl0_42 | ~spl0_51 | spl0_109 | spl0_110 | spl0_111)),
  inference(subsumption_resolution,[],[f3635,f747])).
fof(f747,plain,(
  ~c2_1(a480) | spl0_109),
  inference(avatar_component_clause,[],[f745])).
fof(f745,plain,(
  spl0_109 <=> c2_1(a480)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_109])])).
fof(f3635,plain,(
  c2_1(a480) | (~spl0_42 | ~spl0_51 | spl0_109 | spl0_110 | spl0_111)),
  inference(subsumption_resolution,[],[f3623,f757])).
fof(f757,plain,(
  ~c0_1(a480) | spl0_111),
  inference(avatar_component_clause,[],[f755])).
fof(f755,plain,(
  spl0_111 <=> c0_1(a480)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_111])])).
fof(f3623,plain,(
  c0_1(a480) | c2_1(a480) | (~spl0_42 | ~spl0_51 | spl0_109 | spl0_110)),
  inference(resolution,[],[f439,f3106])).
fof(f3106,plain,(
  c3_1(a480) | (~spl0_42 | spl0_109 | spl0_110)),
  inference(subsumption_resolution,[],[f3036,f747])).
fof(f3036,plain,(
  c3_1(a480) | c2_1(a480) | (~spl0_42 | spl0_110)),
  inference(resolution,[],[f395,f752])).
fof(f752,plain,(
  ~c1_1(a480) | spl0_110),
  inference(avatar_component_clause,[],[f750])).
fof(f750,plain,(
  spl0_110 <=> c1_1(a480)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_110])])).
fof(f395,plain,(
  ( ! [X24] : (c1_1(X24) | c3_1(X24) | c2_1(X24)) ) | ~spl0_42),
  inference(avatar_component_clause,[],[f394])).
fof(f394,plain,(
  spl0_42 <=> ! [X24] : (c3_1(X24) | c1_1(X24) | c2_1(X24))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_42])])).
fof(f3501,plain,(
  spl0_51 | ~spl0_45 | ~spl0_60),
  inference(avatar_split_clause,[],[f3475,f481,f409,f438])).
fof(f409,plain,(
  spl0_45 <=> ! [X36] : (~c3_1(X36) | c0_1(X36) | ~c1_1(X36))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_45])])).
fof(f481,plain,(
  spl0_60 <=> ! [X82] : (c2_1(X82) | c0_1(X82) | c1_1(X82))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_60])])).
fof(f3475,plain,(
  ( ! [X1] : (c0_1(X1) | c2_1(X1) | ~c3_1(X1)) ) | (~spl0_45 | ~spl0_60)),
  inference(duplicate_literal_removal,[],[f3457])).
fof(f3457,plain,(
  ( ! [X1] : (c0_1(X1) | c2_1(X1) | c0_1(X1) | ~c3_1(X1)) ) | (~spl0_45 | ~spl0_60)),
  inference(resolution,[],[f482,f410])).
fof(f410,plain,(
  ( ! [X36] : (~c1_1(X36) | c0_1(X36) | ~c3_1(X36)) ) | ~spl0_45),
  inference(avatar_component_clause,[],[f409])).
fof(f482,plain,(
  ( ! [X82] : (c1_1(X82) | c0_1(X82) | c2_1(X82)) ) | ~spl0_60),
  inference(avatar_component_clause,[],[f481])).
fof(f3497,plain,(
  ~spl0_60 | spl0_109 | spl0_110 | spl0_111),
  inference(avatar_contradiction_clause,[],[f3496])).
fof(f3496,plain,(
  $false | (~spl0_60 | spl0_109 | spl0_110 | spl0_111)),
  inference(subsumption_resolution,[],[f3495,f747])).
fof(f3495,plain,(
  c2_1(a480) | (~spl0_60 | spl0_110 | spl0_111)),
  inference(subsumption_resolution,[],[f3466,f757])).
fof(f3466,plain,(
  c0_1(a480) | c2_1(a480) | (~spl0_60 | spl0_110)),
  inference(resolution,[],[f482,f752])).
fof(f3493,plain,(
  ~spl0_36 | ~spl0_39 | spl0_118 | ~spl0_119 | ~spl0_120),
  inference(avatar_contradiction_clause,[],[f3492])).
fof(f3492,plain,(
  $false | (~spl0_36 | ~spl0_39 | spl0_118 | ~spl0_119 | ~spl0_120)),
  inference(subsumption_resolution,[],[f3491,f800])).
fof(f800,plain,(
  c2_1(a476) | ~spl0_119),
  inference(avatar_component_clause,[],[f798])).
fof(f798,plain,(
  spl0_119 <=> c2_1(a476)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_119])])).
fof(f3491,plain,(
  ~c2_1(a476) | (~spl0_36 | ~spl0_39 | spl0_118 | ~spl0_120)),
  inference(subsumption_resolution,[],[f3490,f805])).
fof(f805,plain,(
  c0_1(a476) | ~spl0_120),
  inference(avatar_component_clause,[],[f803])).
fof(f803,plain,(
  spl0_120 <=> c0_1(a476)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_120])])).
fof(f3490,plain,(
  ~c0_1(a476) | ~c2_1(a476) | (~spl0_36 | ~spl0_39 | spl0_118 | ~spl0_120)),
  inference(resolution,[],[f3412,f367])).
fof(f367,plain,(
  ( ! [X12] : (~c3_1(X12) | ~c0_1(X12) | ~c2_1(X12)) ) | ~spl0_36),
  inference(avatar_component_clause,[],[f366])).
fof(f366,plain,(
  spl0_36 <=> ! [X12] : (~c3_1(X12) | ~c0_1(X12) | ~c2_1(X12))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_36])])).
fof(f3412,plain,(
  c3_1(a476) | (~spl0_39 | spl0_118 | ~spl0_120)),
  inference(subsumption_resolution,[],[f3404,f795])).
fof(f795,plain,(
  ~c1_1(a476) | spl0_118),
  inference(avatar_component_clause,[],[f793])).
fof(f793,plain,(
  spl0_118 <=> c1_1(a476)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_118])])).
fof(f3404,plain,(
  c1_1(a476) | c3_1(a476) | (~spl0_39 | ~spl0_120)),
  inference(resolution,[],[f381,f805])).
fof(f381,plain,(
  ( ! [X18] : (~c0_1(X18) | c1_1(X18) | c3_1(X18)) ) | ~spl0_39),
  inference(avatar_component_clause,[],[f380])).
fof(f380,plain,(
  spl0_39 <=> ! [X18] : (~c0_1(X18) | c1_1(X18) | c3_1(X18))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_39])])).
fof(f3445,plain,(
  spl0_161 | ~spl0_46 | ~spl0_68 | ~spl0_69),
  inference(avatar_split_clause,[],[f3444,f531,f526,f415,f2852])).
fof(f2852,plain,(
  spl0_161 <=> c0_1(a470)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_161])])).
fof(f415,plain,(
  spl0_46 <=> ! [X41] : (~c2_1(X41) | c0_1(X41) | ~c1_1(X41))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_46])])).
fof(f526,plain,(
  spl0_68 <=> c2_1(a470)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_68])])).
fof(f531,plain,(
  spl0_69 <=> c1_1(a470)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_69])])).
fof(f3444,plain,(
  c0_1(a470) | (~spl0_46 | ~spl0_68 | ~spl0_69)),
  inference(subsumption_resolution,[],[f3422,f528])).
fof(f528,plain,(
  c2_1(a470) | ~spl0_68),
  inference(avatar_component_clause,[],[f526])).
fof(f3422,plain,(
  c0_1(a470) | ~c2_1(a470) | (~spl0_46 | ~spl0_69)),
  inference(resolution,[],[f416,f533])).
fof(f533,plain,(
  c1_1(a470) | ~spl0_69),
  inference(avatar_component_clause,[],[f531])).
fof(f416,plain,(
  ( ! [X41] : (~c1_1(X41) | c0_1(X41) | ~c2_1(X41)) ) | ~spl0_46),
  inference(avatar_component_clause,[],[f415])).
fof(f3414,plain,(
  spl0_77 | ~spl0_39 | spl0_76 | ~spl0_78),
  inference(avatar_split_clause,[],[f3413,f579,f569,f380,f574])).
fof(f574,plain,(
  spl0_77 <=> c1_1(a533)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_77])])).
fof(f569,plain,(
  spl0_76 <=> c3_1(a533)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_76])])).
fof(f579,plain,(
  spl0_78 <=> c0_1(a533)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_78])])).
fof(f3413,plain,(
  c1_1(a533) | (~spl0_39 | spl0_76 | ~spl0_78)),
  inference(subsumption_resolution,[],[f3409,f571])).
fof(f571,plain,(
  ~c3_1(a533) | spl0_76),
  inference(avatar_component_clause,[],[f569])).
fof(f3409,plain,(
  c1_1(a533) | c3_1(a533) | (~spl0_39 | ~spl0_78)),
  inference(resolution,[],[f381,f581])).
fof(f581,plain,(
  c0_1(a533) | ~spl0_78),
  inference(avatar_component_clause,[],[f579])).
fof(f3374,plain,(
  ~spl0_56 | spl0_94 | spl0_95 | spl0_96),
  inference(avatar_contradiction_clause,[],[f3373])).
fof(f3373,plain,(
  $false | (~spl0_56 | spl0_94 | spl0_95 | spl0_96)),
  inference(subsumption_resolution,[],[f3372,f667])).
fof(f667,plain,(
  ~c3_1(a494) | spl0_94),
  inference(avatar_component_clause,[],[f665])).
fof(f665,plain,(
  spl0_94 <=> c3_1(a494)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_94])])).
fof(f3372,plain,(
  c3_1(a494) | (~spl0_56 | spl0_95 | spl0_96)),
  inference(subsumption_resolution,[],[f3362,f677])).
fof(f677,plain,(
  ~c0_1(a494) | spl0_96),
  inference(avatar_component_clause,[],[f675])).
fof(f675,plain,(
  spl0_96 <=> c0_1(a494)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_96])])).
fof(f3362,plain,(
  c0_1(a494) | c3_1(a494) | (~spl0_56 | spl0_95)),
  inference(resolution,[],[f462,f672])).
fof(f672,plain,(
  ~c1_1(a494) | spl0_95),
  inference(avatar_component_clause,[],[f670])).
fof(f670,plain,(
  spl0_95 <=> c1_1(a494)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_95])])).
fof(f462,plain,(
  ( ! [X68] : (c1_1(X68) | c0_1(X68) | c3_1(X68)) ) | ~spl0_56),
  inference(avatar_component_clause,[],[f461])).
fof(f461,plain,(
  spl0_56 <=> ! [X68] : (c3_1(X68) | c0_1(X68) | c1_1(X68))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_56])])).
fof(f3326,plain,(
  ~spl0_159 | ~spl0_36 | ~spl0_86 | ~spl0_87),
  inference(avatar_split_clause,[],[f3325,f627,f622,f366,f2834])).
fof(f2834,plain,(
  spl0_159 <=> c0_1(a503)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_159])])).
fof(f622,plain,(
  spl0_86 <=> c3_1(a503)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_86])])).
fof(f627,plain,(
  spl0_87 <=> c2_1(a503)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_87])])).
fof(f3325,plain,(
  ~c0_1(a503) | (~spl0_36 | ~spl0_86 | ~spl0_87)),
  inference(subsumption_resolution,[],[f3314,f629])).
fof(f629,plain,(
  c2_1(a503) | ~spl0_87),
  inference(avatar_component_clause,[],[f627])).
fof(f3314,plain,(
  ~c0_1(a503) | ~c2_1(a503) | (~spl0_36 | ~spl0_86)),
  inference(resolution,[],[f367,f624])).
fof(f624,plain,(
  c3_1(a503) | ~spl0_86),
  inference(avatar_component_clause,[],[f622])).
fof(f3287,plain,(
  spl0_29 | ~spl0_23 | ~spl0_42),
  inference(avatar_split_clause,[],[f3286,f394,f315,f340])).
fof(f340,plain,(
  spl0_29 <=> ! [X8] : (~c0_1(X8) | c2_1(X8) | c3_1(X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_29])])).
fof(f315,plain,(
  spl0_23 <=> ! [X2] : (~c1_1(X2) | c2_1(X2) | ~c0_1(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_23])])).
fof(f3286,plain,(
  ( ! [X0] : (c2_1(X0) | ~c0_1(X0) | c3_1(X0)) ) | (~spl0_23 | ~spl0_42)),
  inference(duplicate_literal_removal,[],[f3275])).
fof(f3275,plain,(
  ( ! [X0] : (c2_1(X0) | ~c0_1(X0) | c3_1(X0) | c2_1(X0)) ) | (~spl0_23 | ~spl0_42)),
  inference(resolution,[],[f316,f395])).
fof(f316,plain,(
  ( ! [X2] : (~c1_1(X2) | c2_1(X2) | ~c0_1(X2)) ) | ~spl0_23),
  inference(avatar_component_clause,[],[f315])).
fof(f3228,plain,(
  ~spl0_30 | ~spl0_34 | ~spl0_45 | ~spl0_131 | ~spl0_132),
  inference(avatar_contradiction_clause,[],[f3227])).
fof(f3227,plain,(
  $false | (~spl0_30 | ~spl0_34 | ~spl0_45 | ~spl0_131 | ~spl0_132)),
  inference(subsumption_resolution,[],[f3218,f869])).
fof(f3218,plain,(
  ~c0_1(a468) | (~spl0_30 | ~spl0_34 | ~spl0_45 | ~spl0_131)),
  inference(resolution,[],[f3180,f864])).
fof(f3180,plain,(
  ( ! [X11] : (~c3_1(X11) | ~c0_1(X11)) ) | (~spl0_30 | ~spl0_34 | ~spl0_45)),
  inference(subsumption_resolution,[],[f360,f3113])).
fof(f3113,plain,(
  ( ! [X7] : (~c1_1(X7) | ~c3_1(X7)) ) | (~spl0_30 | ~spl0_45)),
  inference(subsumption_resolution,[],[f344,f410])).
fof(f344,plain,(
  ( ! [X7] : (~c1_1(X7) | ~c0_1(X7) | ~c3_1(X7)) ) | ~spl0_30),
  inference(avatar_component_clause,[],[f343])).
fof(f343,plain,(
  spl0_30 <=> ! [X7] : (~c3_1(X7) | ~c0_1(X7) | ~c1_1(X7))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_30])])).
fof(f360,plain,(
  ( ! [X11] : (~c3_1(X11) | c1_1(X11) | ~c0_1(X11)) ) | ~spl0_34),
  inference(avatar_component_clause,[],[f359])).
fof(f359,plain,(
  spl0_34 <=> ! [X11] : (~c3_1(X11) | c1_1(X11) | ~c0_1(X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_34])])).
fof(f3205,plain,(
  ~spl0_26 | ~spl0_42 | spl0_139 | spl0_140),
  inference(avatar_contradiction_clause,[],[f3204])).
fof(f3204,plain,(
  $false | (~spl0_26 | ~spl0_42 | spl0_139 | spl0_140)),
  inference(subsumption_resolution,[],[f3188,f912])).
fof(f912,plain,(
  ~c2_1(a465) | spl0_140),
  inference(avatar_component_clause,[],[f910])).
fof(f910,plain,(
  spl0_140 <=> c2_1(a465)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_140])])).
fof(f3188,plain,(
  c2_1(a465) | (~spl0_26 | ~spl0_42 | spl0_139)),
  inference(resolution,[],[f3179,f907])).
fof(f907,plain,(
  ~c3_1(a465) | spl0_139),
  inference(avatar_component_clause,[],[f905])).
fof(f905,plain,(
  spl0_139 <=> c3_1(a465)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_139])])).
fof(f3179,plain,(
  ( ! [X4] : (c3_1(X4) | c2_1(X4)) ) | (~spl0_26 | ~spl0_42)),
  inference(subsumption_resolution,[],[f329,f395])).
fof(f329,plain,(
  ( ! [X4] : (~c1_1(X4) | c2_1(X4) | c3_1(X4)) ) | ~spl0_26),
  inference(avatar_component_clause,[],[f328])).
fof(f328,plain,(
  spl0_26 <=> ! [X4] : (~c1_1(X4) | c2_1(X4) | c3_1(X4))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_26])])).
fof(f3203,plain,(
  ~spl0_26 | ~spl0_42 | spl0_145 | spl0_146),
  inference(avatar_contradiction_clause,[],[f3202])).
fof(f3202,plain,(
  $false | (~spl0_26 | ~spl0_42 | spl0_145 | spl0_146)),
  inference(subsumption_resolution,[],[f3186,f944])).
fof(f944,plain,(
  ~c2_1(a460) | spl0_146),
  inference(avatar_component_clause,[],[f942])).
fof(f942,plain,(
  spl0_146 <=> c2_1(a460)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_146])])).
fof(f3186,plain,(
  c2_1(a460) | (~spl0_26 | ~spl0_42 | spl0_145)),
  inference(resolution,[],[f3179,f939])).
fof(f939,plain,(
  ~c3_1(a460) | spl0_145),
  inference(avatar_component_clause,[],[f937])).
fof(f937,plain,(
  spl0_145 <=> c3_1(a460)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_145])])).
fof(f3131,plain,(
  spl0_159 | ~spl0_55 | spl0_85 | ~spl0_87),
  inference(avatar_split_clause,[],[f3130,f627,f617,f456,f2834])).
fof(f456,plain,(
  spl0_55 <=> ! [X64] : (~c2_1(X64) | c0_1(X64) | c1_1(X64))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_55])])).
fof(f617,plain,(
  spl0_85 <=> c1_1(a503)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_85])])).
fof(f3130,plain,(
  c0_1(a503) | (~spl0_55 | spl0_85 | ~spl0_87)),
  inference(subsumption_resolution,[],[f3090,f619])).
fof(f619,plain,(
  ~c1_1(a503) | spl0_85),
  inference(avatar_component_clause,[],[f617])).
fof(f3090,plain,(
  c0_1(a503) | c1_1(a503) | (~spl0_55 | ~spl0_87)),
  inference(resolution,[],[f457,f629])).
fof(f457,plain,(
  ( ! [X64] : (~c2_1(X64) | c0_1(X64) | c1_1(X64)) ) | ~spl0_55),
  inference(avatar_component_clause,[],[f456])).
fof(f3121,plain,(
  spl0_103 | ~spl0_42 | spl0_104 | spl0_105),
  inference(avatar_split_clause,[],[f3120,f723,f718,f394,f713])).
fof(f713,plain,(
  spl0_103 <=> c3_1(a488)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_103])])).
fof(f718,plain,(
  spl0_104 <=> c2_1(a488)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_104])])).
fof(f723,plain,(
  spl0_105 <=> c1_1(a488)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_105])])).
fof(f3120,plain,(
  c3_1(a488) | (~spl0_42 | spl0_104 | spl0_105)),
  inference(subsumption_resolution,[],[f3037,f720])).
fof(f720,plain,(
  ~c2_1(a488) | spl0_104),
  inference(avatar_component_clause,[],[f718])).
fof(f3037,plain,(
  c3_1(a488) | c2_1(a488) | (~spl0_42 | spl0_105)),
  inference(resolution,[],[f395,f725])).
fof(f725,plain,(
  ~c1_1(a488) | spl0_105),
  inference(avatar_component_clause,[],[f723])).
fof(f3119,plain,(
  spl0_94 | ~spl0_42 | spl0_95 | spl0_158),
  inference(avatar_split_clause,[],[f3118,f2631,f670,f394,f665])).
fof(f2631,plain,(
  spl0_158 <=> c2_1(a494)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_158])])).
fof(f3118,plain,(
  c3_1(a494) | (~spl0_42 | spl0_95 | spl0_158)),
  inference(subsumption_resolution,[],[f3039,f2632])).
fof(f2632,plain,(
  ~c2_1(a494) | spl0_158),
  inference(avatar_component_clause,[],[f2631])).
fof(f3039,plain,(
  c3_1(a494) | c2_1(a494) | (~spl0_42 | spl0_95)),
  inference(resolution,[],[f395,f672])).
fof(f3109,plain,(
  spl0_150 | ~spl0_55 | spl0_88 | ~spl0_90),
  inference(avatar_split_clause,[],[f3108,f643,f633,f456,f2222])).
fof(f2222,plain,(
  spl0_150 <=> c1_1(a502)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_150])])).
fof(f633,plain,(
  spl0_88 <=> c0_1(a502)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_88])])).
fof(f643,plain,(
  spl0_90 <=> c2_1(a502)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_90])])).
fof(f3108,plain,(
  c1_1(a502) | (~spl0_55 | spl0_88 | ~spl0_90)),
  inference(subsumption_resolution,[],[f3089,f635])).
fof(f635,plain,(
  ~c0_1(a502) | spl0_88),
  inference(avatar_component_clause,[],[f633])).
fof(f3089,plain,(
  c0_1(a502) | c1_1(a502) | (~spl0_55 | ~spl0_90)),
  inference(resolution,[],[f457,f645])).
fof(f645,plain,(
  c2_1(a502) | ~spl0_90),
  inference(avatar_component_clause,[],[f643])).
fof(f3101,plain,(
  ~spl0_55 | spl0_95 | spl0_96 | ~spl0_158),
  inference(avatar_contradiction_clause,[],[f3100])).
fof(f3100,plain,(
  $false | (~spl0_55 | spl0_95 | spl0_96 | ~spl0_158)),
  inference(subsumption_resolution,[],[f3099,f672])).
fof(f3099,plain,(
  c1_1(a494) | (~spl0_55 | spl0_96 | ~spl0_158)),
  inference(subsumption_resolution,[],[f3087,f677])).
fof(f3087,plain,(
  c0_1(a494) | c1_1(a494) | (~spl0_55 | ~spl0_158)),
  inference(resolution,[],[f457,f2633])).
fof(f2633,plain,(
  c2_1(a494) | ~spl0_158),
  inference(avatar_component_clause,[],[f2631])).
fof(f3014,plain,(
  ~spl0_23 | ~spl0_52 | spl0_97 | ~spl0_99),
  inference(avatar_contradiction_clause,[],[f3013])).
fof(f3013,plain,(
  $false | (~spl0_23 | ~spl0_52 | spl0_97 | ~spl0_99)),
  inference(subsumption_resolution,[],[f3003,f683])).
fof(f683,plain,(
  ~c2_1(a493) | spl0_97),
  inference(avatar_component_clause,[],[f681])).
fof(f681,plain,(
  spl0_97 <=> c2_1(a493)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_97])])).
fof(f3003,plain,(
  c2_1(a493) | (~spl0_23 | ~spl0_52 | ~spl0_99)),
  inference(resolution,[],[f2978,f693])).
fof(f693,plain,(
  c1_1(a493) | ~spl0_99),
  inference(avatar_component_clause,[],[f691])).
fof(f691,plain,(
  spl0_99 <=> c1_1(a493)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_99])])).
fof(f2978,plain,(
  ( ! [X2] : (~c1_1(X2) | c2_1(X2)) ) | (~spl0_23 | ~spl0_52)),
  inference(subsumption_resolution,[],[f316,f444])).
fof(f444,plain,(
  ( ! [X56] : (~c1_1(X56) | c0_1(X56) | c2_1(X56)) ) | ~spl0_52),
  inference(avatar_component_clause,[],[f443])).
fof(f443,plain,(
  spl0_52 <=> ! [X56] : (~c1_1(X56) | c0_1(X56) | c2_1(X56))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_52])])).
fof(f2994,plain,(
  ~spl0_29 | ~spl0_54 | spl0_139 | spl0_140),
  inference(avatar_contradiction_clause,[],[f2993])).
fof(f2993,plain,(
  $false | (~spl0_29 | ~spl0_54 | spl0_139 | spl0_140)),
  inference(subsumption_resolution,[],[f2980,f912])).
fof(f2980,plain,(
  c2_1(a465) | (~spl0_29 | ~spl0_54 | spl0_139)),
  inference(resolution,[],[f2977,f907])).
fof(f2977,plain,(
  ( ! [X8] : (c3_1(X8) | c2_1(X8)) ) | (~spl0_29 | ~spl0_54)),
  inference(subsumption_resolution,[],[f341,f452])).
fof(f452,plain,(
  ( ! [X61] : (c3_1(X61) | c0_1(X61) | c2_1(X61)) ) | ~spl0_54),
  inference(avatar_component_clause,[],[f451])).
fof(f451,plain,(
  spl0_54 <=> ! [X61] : (c3_1(X61) | c0_1(X61) | c2_1(X61))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_54])])).
fof(f341,plain,(
  ( ! [X8] : (~c0_1(X8) | c2_1(X8) | c3_1(X8)) ) | ~spl0_29),
  inference(avatar_component_clause,[],[f340])).
fof(f2971,plain,(
  ~spl0_23 | ~spl0_41 | spl0_130 | ~spl0_132),
  inference(avatar_contradiction_clause,[],[f2970])).
fof(f2970,plain,(
  $false | (~spl0_23 | ~spl0_41 | spl0_130 | ~spl0_132)),
  inference(subsumption_resolution,[],[f2956,f859])).
fof(f2956,plain,(
  c2_1(a468) | (~spl0_23 | ~spl0_41 | ~spl0_132)),
  inference(resolution,[],[f2928,f869])).
fof(f2928,plain,(
  ( ! [X2] : (~c0_1(X2) | c2_1(X2)) ) | (~spl0_23 | ~spl0_41)),
  inference(subsumption_resolution,[],[f316,f391])).
fof(f391,plain,(
  ( ! [X23] : (~c0_1(X23) | c1_1(X23) | c2_1(X23)) ) | ~spl0_41),
  inference(avatar_component_clause,[],[f390])).
fof(f390,plain,(
  spl0_41 <=> ! [X23] : (~c0_1(X23) | c1_1(X23) | c2_1(X23))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_41])])).
fof(f2925,plain,(
  ~spl0_20 | ~spl0_36 | ~spl0_61 | ~spl0_63),
  inference(avatar_contradiction_clause,[],[f2924])).
fof(f2924,plain,(
  $false | (~spl0_20 | ~spl0_36 | ~spl0_61 | ~spl0_63)),
  inference(subsumption_resolution,[],[f2908,f501])).
fof(f2908,plain,(
  ~c0_1(a490) | (~spl0_20 | ~spl0_36 | ~spl0_61)),
  inference(resolution,[],[f2861,f491])).
fof(f491,plain,(
  c2_1(a490) | ~spl0_61),
  inference(avatar_component_clause,[],[f489])).
fof(f2861,plain,(
  ( ! [X0] : (~c2_1(X0) | ~c0_1(X0)) ) | (~spl0_20 | ~spl0_36)),
  inference(subsumption_resolution,[],[f304,f367])).
fof(f304,plain,(
  ( ! [X0] : (~c2_1(X0) | c3_1(X0) | ~c0_1(X0)) ) | ~spl0_20),
  inference(avatar_component_clause,[],[f303])).
fof(f303,plain,(
  spl0_20 <=> ! [X0] : (~c2_1(X0) | c3_1(X0) | ~c0_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f2914,plain,(
  ~spl0_20 | ~spl0_36 | ~spl0_119 | ~spl0_120),
  inference(avatar_contradiction_clause,[],[f2913])).
fof(f2913,plain,(
  $false | (~spl0_20 | ~spl0_36 | ~spl0_119 | ~spl0_120)),
  inference(subsumption_resolution,[],[f2897,f805])).
fof(f2897,plain,(
  ~c0_1(a476) | (~spl0_20 | ~spl0_36 | ~spl0_119)),
  inference(resolution,[],[f2861,f800])).
fof(f2891,plain,(
  spl0_161 | ~spl0_45 | ~spl0_50 | ~spl0_69),
  inference(avatar_split_clause,[],[f2878,f531,f434,f409,f2852])).
fof(f434,plain,(
  spl0_50 <=> ! [X50] : (~c1_1(X50) | c0_1(X50) | c3_1(X50))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_50])])).
fof(f2878,plain,(
  c0_1(a470) | (~spl0_45 | ~spl0_50 | ~spl0_69)),
  inference(resolution,[],[f2859,f533])).
fof(f2859,plain,(
  ( ! [X36] : (~c1_1(X36) | c0_1(X36)) ) | (~spl0_45 | ~spl0_50)),
  inference(subsumption_resolution,[],[f410,f435])).
fof(f435,plain,(
  ( ! [X50] : (~c1_1(X50) | c0_1(X50) | c3_1(X50)) ) | ~spl0_50),
  inference(avatar_component_clause,[],[f434])).
fof(f2888,plain,(
  ~spl0_45 | ~spl0_50 | spl0_91 | ~spl0_93),
  inference(avatar_contradiction_clause,[],[f2887])).
fof(f2887,plain,(
  $false | (~spl0_45 | ~spl0_50 | spl0_91 | ~spl0_93)),
  inference(subsumption_resolution,[],[f2874,f651])).
fof(f651,plain,(
  ~c0_1(a500) | spl0_91),
  inference(avatar_component_clause,[],[f649])).
fof(f649,plain,(
  spl0_91 <=> c0_1(a500)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_91])])).
fof(f2874,plain,(
  c0_1(a500) | (~spl0_45 | ~spl0_50 | ~spl0_93)),
  inference(resolution,[],[f2859,f661])).
fof(f661,plain,(
  c1_1(a500) | ~spl0_93),
  inference(avatar_component_clause,[],[f659])).
fof(f659,plain,(
  spl0_93 <=> c1_1(a500)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_93])])).
fof(f2886,plain,(
  ~spl0_45 | ~spl0_50 | spl0_98 | ~spl0_99),
  inference(avatar_contradiction_clause,[],[f2885])).
fof(f2885,plain,(
  $false | (~spl0_45 | ~spl0_50 | spl0_98 | ~spl0_99)),
  inference(subsumption_resolution,[],[f2873,f688])).
fof(f688,plain,(
  ~c0_1(a493) | spl0_98),
  inference(avatar_component_clause,[],[f686])).
fof(f686,plain,(
  spl0_98 <=> c0_1(a493)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_98])])).
fof(f2873,plain,(
  c0_1(a493) | (~spl0_45 | ~spl0_50 | ~spl0_99)),
  inference(resolution,[],[f2859,f693])).
fof(f2882,plain,(
  ~spl0_45 | ~spl0_50 | spl0_136 | ~spl0_138),
  inference(avatar_contradiction_clause,[],[f2881])).
fof(f2881,plain,(
  $false | (~spl0_45 | ~spl0_50 | spl0_136 | ~spl0_138)),
  inference(subsumption_resolution,[],[f2867,f891])).
fof(f2867,plain,(
  c0_1(a466) | (~spl0_45 | ~spl0_50 | ~spl0_138)),
  inference(resolution,[],[f2859,f901])).
fof(f901,plain,(
  c1_1(a466) | ~spl0_138),
  inference(avatar_component_clause,[],[f899])).
fof(f899,plain,(
  spl0_138 <=> c1_1(a466)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_138])])).
fof(f2855,plain,(
  ~spl0_67 | ~spl0_161 | ~spl0_30 | ~spl0_69),
  inference(avatar_split_clause,[],[f2819,f531,f343,f2852,f521])).
fof(f521,plain,(
  spl0_67 <=> c3_1(a470)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_67])])).
fof(f2819,plain,(
  ~c0_1(a470) | ~c3_1(a470) | (~spl0_30 | ~spl0_69)),
  inference(resolution,[],[f344,f533])).
fof(f2845,plain,(
  ~spl0_64 | ~spl0_30 | ~spl0_65 | ~spl0_66),
  inference(avatar_split_clause,[],[f2839,f515,f510,f343,f505])).
fof(f2839,plain,(
  ~c3_1(a473) | (~spl0_30 | ~spl0_65 | ~spl0_66)),
  inference(subsumption_resolution,[],[f2820,f517])).
fof(f2820,plain,(
  ~c0_1(a473) | ~c3_1(a473) | (~spl0_30 | ~spl0_65)),
  inference(resolution,[],[f344,f512])).
fof(f2767,plain,(
  ~spl0_29 | ~spl0_54 | spl0_103 | spl0_104),
  inference(avatar_contradiction_clause,[],[f2766])).
fof(f2766,plain,(
  $false | (~spl0_29 | ~spl0_54 | spl0_103 | spl0_104)),
  inference(subsumption_resolution,[],[f2760,f720])).
fof(f2760,plain,(
  c2_1(a488) | (~spl0_29 | ~spl0_54 | spl0_103)),
  inference(resolution,[],[f2699,f715])).
fof(f715,plain,(
  ~c3_1(a488) | spl0_103),
  inference(avatar_component_clause,[],[f713])).
fof(f2699,plain,(
  ( ! [X61] : (c3_1(X61) | c2_1(X61)) ) | (~spl0_29 | ~spl0_54)),
  inference(subsumption_resolution,[],[f452,f341])).
fof(f2713,plain,(
  ~spl0_35 | ~spl0_44 | ~spl0_48 | spl0_128 | ~spl0_129),
  inference(avatar_contradiction_clause,[],[f2712])).
fof(f2712,plain,(
  $false | (~spl0_35 | ~spl0_44 | ~spl0_48 | spl0_128 | ~spl0_129)),
  inference(subsumption_resolution,[],[f2702,f848])).
fof(f2702,plain,(
  c1_1(a471) | (~spl0_35 | ~spl0_44 | ~spl0_48 | ~spl0_129)),
  inference(resolution,[],[f2675,f853])).
fof(f2675,plain,(
  ( ! [X13] : (~c2_1(X13) | c1_1(X13)) ) | (~spl0_35 | ~spl0_44 | ~spl0_48)),
  inference(subsumption_resolution,[],[f364,f2514])).
fof(f2514,plain,(
  ( ! [X28] : (~c2_1(X28) | c0_1(X28)) ) | (~spl0_44 | ~spl0_48)),
  inference(subsumption_resolution,[],[f404,f427])).
fof(f364,plain,(
  ( ! [X13] : (~c2_1(X13) | c1_1(X13) | ~c0_1(X13)) ) | ~spl0_35),
  inference(avatar_component_clause,[],[f363])).
fof(f363,plain,(
  spl0_35 <=> ! [X13] : (~c2_1(X13) | c1_1(X13) | ~c0_1(X13))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_35])])).
fof(f2697,plain,(
  ~spl0_50 | ~spl0_56 | spl0_139 | spl0_141),
  inference(avatar_contradiction_clause,[],[f2696])).
fof(f2696,plain,(
  $false | (~spl0_50 | ~spl0_56 | spl0_139 | spl0_141)),
  inference(subsumption_resolution,[],[f2684,f917])).
fof(f917,plain,(
  ~c0_1(a465) | spl0_141),
  inference(avatar_component_clause,[],[f915])).
fof(f915,plain,(
  spl0_141 <=> c0_1(a465)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_141])])).
fof(f2684,plain,(
  c0_1(a465) | (~spl0_50 | ~spl0_56 | spl0_139)),
  inference(resolution,[],[f2674,f907])).
fof(f2674,plain,(
  ( ! [X50] : (c3_1(X50) | c0_1(X50)) ) | (~spl0_50 | ~spl0_56)),
  inference(subsumption_resolution,[],[f435,f462])).
fof(f2673,plain,(
  ~spl0_23 | ~spl0_36 | ~spl0_64 | ~spl0_65 | ~spl0_66),
  inference(avatar_contradiction_clause,[],[f2672])).
fof(f2672,plain,(
  $false | (~spl0_23 | ~spl0_36 | ~spl0_64 | ~spl0_65 | ~spl0_66)),
  inference(subsumption_resolution,[],[f2671,f2590])).
fof(f2590,plain,(
  c2_1(a473) | (~spl0_23 | ~spl0_65 | ~spl0_66)),
  inference(subsumption_resolution,[],[f2586,f517])).
fof(f2586,plain,(
  c2_1(a473) | ~c0_1(a473) | (~spl0_23 | ~spl0_65)),
  inference(resolution,[],[f316,f512])).
fof(f2671,plain,(
  ~c2_1(a473) | (~spl0_36 | ~spl0_64 | ~spl0_66)),
  inference(subsumption_resolution,[],[f2667,f517])).
fof(f2667,plain,(
  ~c0_1(a473) | ~c2_1(a473) | (~spl0_36 | ~spl0_64)),
  inference(resolution,[],[f367,f507])).
fof(f2624,plain,(
  ~spl0_34 | spl0_82 | ~spl0_83 | ~spl0_84),
  inference(avatar_contradiction_clause,[],[f2623])).
fof(f2623,plain,(
  $false | (~spl0_34 | spl0_82 | ~spl0_83 | ~spl0_84)),
  inference(subsumption_resolution,[],[f2622,f613])).
fof(f613,plain,(
  c0_1(a512) | ~spl0_84),
  inference(avatar_component_clause,[],[f611])).
fof(f611,plain,(
  spl0_84 <=> c0_1(a512)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_84])])).
fof(f2622,plain,(
  ~c0_1(a512) | (~spl0_34 | spl0_82 | ~spl0_83)),
  inference(subsumption_resolution,[],[f2602,f603])).
fof(f603,plain,(
  ~c1_1(a512) | spl0_82),
  inference(avatar_component_clause,[],[f601])).
fof(f601,plain,(
  spl0_82 <=> c1_1(a512)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_82])])).
fof(f2602,plain,(
  c1_1(a512) | ~c0_1(a512) | (~spl0_34 | ~spl0_83)),
  inference(resolution,[],[f360,f608])).
fof(f608,plain,(
  c3_1(a512) | ~spl0_83),
  inference(avatar_component_clause,[],[f606])).
fof(f606,plain,(
  spl0_83 <=> c3_1(a512)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_83])])).
fof(f2548,plain,(
  ~spl0_37 | ~spl0_53 | spl0_118 | ~spl0_119),
  inference(avatar_contradiction_clause,[],[f2547])).
fof(f2547,plain,(
  $false | (~spl0_37 | ~spl0_53 | spl0_118 | ~spl0_119)),
  inference(subsumption_resolution,[],[f2535,f795])).
fof(f2535,plain,(
  c1_1(a476) | (~spl0_37 | ~spl0_53 | ~spl0_119)),
  inference(resolution,[],[f2503,f800])).
fof(f2503,plain,(
  ( ! [X16] : (~c2_1(X16) | c1_1(X16)) ) | (~spl0_37 | ~spl0_53)),
  inference(subsumption_resolution,[],[f372,f447])).
fof(f447,plain,(
  ( ! [X55] : (~c3_1(X55) | c1_1(X55) | ~c2_1(X55)) ) | ~spl0_53),
  inference(avatar_component_clause,[],[f446])).
fof(f446,plain,(
  spl0_53 <=> ! [X55] : (~c3_1(X55) | c1_1(X55) | ~c2_1(X55))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_53])])).
fof(f2384,plain,(
  ~spl0_36 | ~spl0_44 | ~spl0_86 | ~spl0_87),
  inference(avatar_contradiction_clause,[],[f2383])).
fof(f2383,plain,(
  $false | (~spl0_36 | ~spl0_44 | ~spl0_86 | ~spl0_87)),
  inference(subsumption_resolution,[],[f2375,f629])).
fof(f2375,plain,(
  ~c2_1(a503) | (~spl0_36 | ~spl0_44 | ~spl0_86)),
  inference(resolution,[],[f2357,f624])).
fof(f2357,plain,(
  ( ! [X12] : (~c3_1(X12) | ~c2_1(X12)) ) | (~spl0_36 | ~spl0_44)),
  inference(subsumption_resolution,[],[f367,f404])).
fof(f2362,plain,(
  ~spl0_45 | spl0_88 | ~spl0_89 | ~spl0_150),
  inference(avatar_contradiction_clause,[],[f2361])).
fof(f2361,plain,(
  $false | (~spl0_45 | spl0_88 | ~spl0_89 | ~spl0_150)),
  inference(subsumption_resolution,[],[f2360,f640])).
fof(f640,plain,(
  c3_1(a502) | ~spl0_89),
  inference(avatar_component_clause,[],[f638])).
fof(f638,plain,(
  spl0_89 <=> c3_1(a502)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_89])])).
fof(f2360,plain,(
  ~c3_1(a502) | (~spl0_45 | spl0_88 | ~spl0_150)),
  inference(subsumption_resolution,[],[f2359,f635])).
fof(f2359,plain,(
  c0_1(a502) | ~c3_1(a502) | (~spl0_45 | ~spl0_150)),
  inference(resolution,[],[f2224,f410])).
fof(f2224,plain,(
  c1_1(a502) | ~spl0_150),
  inference(avatar_component_clause,[],[f2222])).
fof(f2356,plain,(
  ~spl0_35 | ~spl0_55 | spl0_85 | ~spl0_87),
  inference(avatar_contradiction_clause,[],[f2355])).
fof(f2355,plain,(
  $false | (~spl0_35 | ~spl0_55 | spl0_85 | ~spl0_87)),
  inference(subsumption_resolution,[],[f2347,f619])).
fof(f2347,plain,(
  c1_1(a503) | (~spl0_35 | ~spl0_55 | ~spl0_87)),
  inference(resolution,[],[f2314,f629])).
fof(f2314,plain,(
  ( ! [X13] : (~c2_1(X13) | c1_1(X13)) ) | (~spl0_35 | ~spl0_55)),
  inference(subsumption_resolution,[],[f364,f457])).
fof(f2354,plain,(
  ~spl0_35 | ~spl0_55 | spl0_118 | ~spl0_119),
  inference(avatar_contradiction_clause,[],[f2353])).
fof(f2353,plain,(
  $false | (~spl0_35 | ~spl0_55 | spl0_118 | ~spl0_119)),
  inference(subsumption_resolution,[],[f2343,f795])).
fof(f2343,plain,(
  c1_1(a476) | (~spl0_35 | ~spl0_55 | ~spl0_119)),
  inference(resolution,[],[f2314,f800])).
fof(f2350,plain,(
  ~spl0_35 | ~spl0_55 | spl0_142 | ~spl0_144),
  inference(avatar_contradiction_clause,[],[f2349])).
fof(f2349,plain,(
  $false | (~spl0_35 | ~spl0_55 | spl0_142 | ~spl0_144)),
  inference(subsumption_resolution,[],[f2339,f923])).
fof(f923,plain,(
  ~c1_1(a463) | spl0_142),
  inference(avatar_component_clause,[],[f921])).
fof(f921,plain,(
  spl0_142 <=> c1_1(a463)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_142])])).
fof(f2339,plain,(
  c1_1(a463) | (~spl0_35 | ~spl0_55 | ~spl0_144)),
  inference(resolution,[],[f2314,f933])).
fof(f933,plain,(
  c2_1(a463) | ~spl0_144),
  inference(avatar_component_clause,[],[f931])).
fof(f931,plain,(
  spl0_144 <=> c2_1(a463)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_144])])).
fof(f2310,plain,(
  ~spl0_46 | ~spl0_55 | spl0_113 | ~spl0_114),
  inference(avatar_contradiction_clause,[],[f2309])).
fof(f2309,plain,(
  $false | (~spl0_46 | ~spl0_55 | spl0_113 | ~spl0_114)),
  inference(subsumption_resolution,[],[f2298,f768])).
fof(f768,plain,(
  ~c0_1(a478) | spl0_113),
  inference(avatar_component_clause,[],[f766])).
fof(f2298,plain,(
  c0_1(a478) | (~spl0_46 | ~spl0_55 | ~spl0_114)),
  inference(resolution,[],[f2281,f773])).
fof(f2281,plain,(
  ( ! [X64] : (~c2_1(X64) | c0_1(X64)) ) | (~spl0_46 | ~spl0_55)),
  inference(subsumption_resolution,[],[f457,f416])).
fof(f2226,plain,(
  ~spl0_90 | spl0_88 | ~spl0_44 | ~spl0_89),
  inference(avatar_split_clause,[],[f2112,f638,f403,f633,f643])).
fof(f2112,plain,(
  c0_1(a502) | ~c2_1(a502) | (~spl0_44 | ~spl0_89)),
  inference(resolution,[],[f404,f640])).
fof(f2220,plain,(
  ~spl0_53 | spl0_85 | ~spl0_86 | ~spl0_87),
  inference(avatar_contradiction_clause,[],[f2219])).
fof(f2219,plain,(
  $false | (~spl0_53 | spl0_85 | ~spl0_86 | ~spl0_87)),
  inference(subsumption_resolution,[],[f2218,f629])).
fof(f2218,plain,(
  ~c2_1(a503) | (~spl0_53 | spl0_85 | ~spl0_86)),
  inference(subsumption_resolution,[],[f2211,f619])).
fof(f2211,plain,(
  c1_1(a503) | ~c2_1(a503) | (~spl0_53 | ~spl0_86)),
  inference(resolution,[],[f447,f624])).
fof(f2191,plain,(
  spl0_91 | ~spl0_46 | ~spl0_92 | ~spl0_93),
  inference(avatar_split_clause,[],[f2190,f659,f654,f415,f649])).
fof(f654,plain,(
  spl0_92 <=> c2_1(a500)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_92])])).
fof(f2190,plain,(
  c0_1(a500) | (~spl0_46 | ~spl0_92 | ~spl0_93)),
  inference(subsumption_resolution,[],[f2177,f656])).
fof(f656,plain,(
  c2_1(a500) | ~spl0_92),
  inference(avatar_component_clause,[],[f654])).
fof(f2177,plain,(
  c0_1(a500) | ~c2_1(a500) | (~spl0_46 | ~spl0_93)),
  inference(resolution,[],[f416,f661])).
fof(f2140,plain,(
  ~spl0_45 | spl0_136 | ~spl0_137 | ~spl0_138),
  inference(avatar_contradiction_clause,[],[f2139])).
fof(f2139,plain,(
  $false | (~spl0_45 | spl0_136 | ~spl0_137 | ~spl0_138)),
  inference(subsumption_resolution,[],[f2138,f896])).
fof(f2138,plain,(
  ~c3_1(a466) | (~spl0_45 | spl0_136 | ~spl0_138)),
  inference(subsumption_resolution,[],[f2126,f891])).
fof(f2126,plain,(
  c0_1(a466) | ~c3_1(a466) | (~spl0_45 | ~spl0_138)),
  inference(resolution,[],[f410,f901])).
fof(f2026,plain,(
  ~spl0_51 | ~spl0_54 | spl0_109 | spl0_111),
  inference(avatar_contradiction_clause,[],[f2025])).
fof(f2025,plain,(
  $false | (~spl0_51 | ~spl0_54 | spl0_109 | spl0_111)),
  inference(subsumption_resolution,[],[f2012,f757])).
fof(f2012,plain,(
  c0_1(a480) | (~spl0_51 | ~spl0_54 | spl0_109)),
  inference(resolution,[],[f1977,f747])).
fof(f1977,plain,(
  ( ! [X61] : (c2_1(X61) | c0_1(X61)) ) | (~spl0_51 | ~spl0_54)),
  inference(subsumption_resolution,[],[f452,f439])).
fof(f2022,plain,(
  ~spl0_51 | ~spl0_54 | spl0_140 | spl0_141),
  inference(avatar_contradiction_clause,[],[f2021])).
fof(f2021,plain,(
  $false | (~spl0_51 | ~spl0_54 | spl0_140 | spl0_141)),
  inference(subsumption_resolution,[],[f2005,f917])).
fof(f2005,plain,(
  c0_1(a465) | (~spl0_51 | ~spl0_54 | spl0_140)),
  inference(resolution,[],[f1977,f912])).
fof(f1832,plain,(
  ~spl0_46 | ~spl0_52 | spl0_136 | ~spl0_138),
  inference(avatar_contradiction_clause,[],[f1831])).
fof(f1831,plain,(
  $false | (~spl0_46 | ~spl0_52 | spl0_136 | ~spl0_138)),
  inference(subsumption_resolution,[],[f1823,f891])).
fof(f1823,plain,(
  c0_1(a466) | (~spl0_46 | ~spl0_52 | ~spl0_138)),
  inference(resolution,[],[f1737,f901])).
fof(f1737,plain,(
  ( ! [X41] : (~c1_1(X41) | c0_1(X41)) ) | (~spl0_46 | ~spl0_52)),
  inference(subsumption_resolution,[],[f416,f444])).
fof(f1816,plain,(
  ~spl0_35 | ~spl0_41 | spl0_82 | ~spl0_84),
  inference(avatar_contradiction_clause,[],[f1815])).
fof(f1815,plain,(
  $false | (~spl0_35 | ~spl0_41 | spl0_82 | ~spl0_84)),
  inference(subsumption_resolution,[],[f1810,f603])).
fof(f1810,plain,(
  c1_1(a512) | (~spl0_35 | ~spl0_41 | ~spl0_84)),
  inference(resolution,[],[f1709,f613])).
fof(f1709,plain,(
  ( ! [X23] : (~c0_1(X23) | c1_1(X23)) ) | (~spl0_35 | ~spl0_41)),
  inference(subsumption_resolution,[],[f391,f364])).
fof(f1668,plain,(
  ~spl0_46 | ~spl0_51 | ~spl0_54 | ~spl0_55 | spl0_96),
  inference(avatar_contradiction_clause,[],[f1665])).
fof(f1665,plain,(
  $false | (~spl0_46 | ~spl0_51 | ~spl0_54 | ~spl0_55 | spl0_96)),
  inference(resolution,[],[f1654,f677])).
fof(f1654,plain,(
  ( ! [X64] : (c0_1(X64)) ) | (~spl0_46 | ~spl0_51 | ~spl0_54 | ~spl0_55)),
  inference(subsumption_resolution,[],[f1653,f1648])).
fof(f1648,plain,(
  ( ! [X61] : (c0_1(X61) | c2_1(X61)) ) | (~spl0_51 | ~spl0_54)),
  inference(subsumption_resolution,[],[f452,f439])).
fof(f1653,plain,(
  ( ! [X64] : (~c2_1(X64) | c0_1(X64)) ) | (~spl0_46 | ~spl0_55)),
  inference(subsumption_resolution,[],[f457,f416])).
fof(f1651,plain,(
  spl0_145 | ~spl0_29 | spl0_146 | ~spl0_147),
  inference(avatar_split_clause,[],[f1650,f947,f942,f340,f937])).
fof(f947,plain,(
  spl0_147 <=> c0_1(a460)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_147])])).
fof(f1650,plain,(
  c3_1(a460) | (~spl0_29 | spl0_146 | ~spl0_147)),
  inference(subsumption_resolution,[],[f1228,f944])).
fof(f1228,plain,(
  c2_1(a460) | c3_1(a460) | (~spl0_29 | ~spl0_147)),
  inference(resolution,[],[f949,f341])).
fof(f949,plain,(
  c0_1(a460) | ~spl0_147),
  inference(avatar_component_clause,[],[f947])).
fof(f1560,plain,(
  ~spl0_36 | ~spl0_55 | ~spl0_60 | spl0_85 | ~spl0_86 | ~spl0_87),
  inference(avatar_contradiction_clause,[],[f1559])).
fof(f1559,plain,(
  $false | (~spl0_36 | ~spl0_55 | ~spl0_60 | spl0_85 | ~spl0_86 | ~spl0_87)),
  inference(subsumption_resolution,[],[f1558,f629])).
fof(f1558,plain,(
  ~c2_1(a503) | (~spl0_36 | ~spl0_55 | ~spl0_60 | spl0_85 | ~spl0_86)),
  inference(subsumption_resolution,[],[f1421,f1234])).
fof(f1234,plain,(
  c0_1(a503) | (~spl0_55 | ~spl0_60 | spl0_85)),
  inference(resolution,[],[f1170,f619])).
fof(f1170,plain,(
  ( ! [X64] : (c1_1(X64) | c0_1(X64)) ) | (~spl0_55 | ~spl0_60)),
  inference(subsumption_resolution,[],[f457,f482])).
fof(f1421,plain,(
  ~c0_1(a503) | ~c2_1(a503) | (~spl0_36 | ~spl0_86)),
  inference(resolution,[],[f367,f624])).
fof(f1554,plain,(
  ~spl0_36 | ~spl0_70 | ~spl0_71 | ~spl0_72),
  inference(avatar_contradiction_clause,[],[f1553])).
fof(f1553,plain,(
  $false | (~spl0_36 | ~spl0_70 | ~spl0_71 | ~spl0_72)),
  inference(subsumption_resolution,[],[f1552,f544])).
fof(f544,plain,(
  c2_1(a461) | ~spl0_71),
  inference(avatar_component_clause,[],[f542])).
fof(f542,plain,(
  spl0_71 <=> c2_1(a461)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_71])])).
fof(f1552,plain,(
  ~c2_1(a461) | (~spl0_36 | ~spl0_70 | ~spl0_72)),
  inference(subsumption_resolution,[],[f1423,f549])).
fof(f549,plain,(
  c0_1(a461) | ~spl0_72),
  inference(avatar_component_clause,[],[f547])).
fof(f547,plain,(
  spl0_72 <=> c0_1(a461)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_72])])).
fof(f1423,plain,(
  ~c0_1(a461) | ~c2_1(a461) | (~spl0_36 | ~spl0_70)),
  inference(resolution,[],[f367,f539])).
fof(f539,plain,(
  c3_1(a461) | ~spl0_70),
  inference(avatar_component_clause,[],[f537])).
fof(f537,plain,(
  spl0_70 <=> c3_1(a461)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_70])])).
fof(f1240,plain,(
  spl0_96 | ~spl0_55 | ~spl0_60 | spl0_95),
  inference(avatar_split_clause,[],[f1233,f670,f481,f456,f675])).
fof(f1233,plain,(
  c0_1(a494) | (~spl0_55 | ~spl0_60 | spl0_95)),
  inference(resolution,[],[f1170,f672])).
fof(f1216,plain,(
  ~spl0_23 | ~spl0_32 | ~spl0_62 | ~spl0_63),
  inference(avatar_contradiction_clause,[],[f1215])).
fof(f1215,plain,(
  $false | (~spl0_23 | ~spl0_32 | ~spl0_62 | ~spl0_63)),
  inference(subsumption_resolution,[],[f1212,f501])).
fof(f1212,plain,(
  ~c0_1(a490) | (~spl0_23 | ~spl0_32 | ~spl0_62)),
  inference(resolution,[],[f1143,f496])).
fof(f1143,plain,(
  ( ! [X2] : (~c1_1(X2) | ~c0_1(X2)) ) | (~spl0_23 | ~spl0_32)),
  inference(subsumption_resolution,[],[f316,f352])).
fof(f1156,plain,(
  ~spl0_39 | ~spl0_56 | spl0_103 | spl0_105),
  inference(avatar_contradiction_clause,[],[f1155])).
fof(f1155,plain,(
  $false | (~spl0_39 | ~spl0_56 | spl0_103 | spl0_105)),
  inference(subsumption_resolution,[],[f1149,f715])).
fof(f1149,plain,(
  c3_1(a488) | (~spl0_39 | ~spl0_56 | spl0_105)),
  inference(resolution,[],[f1083,f725])).
fof(f1083,plain,(
  ( ! [X18] : (c1_1(X18) | c3_1(X18)) ) | (~spl0_39 | ~spl0_56)),
  inference(subsumption_resolution,[],[f381,f462])).
fof(f1154,plain,(
  ~spl0_39 | ~spl0_56 | spl0_127 | spl0_128),
  inference(avatar_contradiction_clause,[],[f1153])).
fof(f1153,plain,(
  $false | (~spl0_39 | ~spl0_56 | spl0_127 | spl0_128)),
  inference(subsumption_resolution,[],[f1146,f843])).
fof(f843,plain,(
  ~c3_1(a471) | spl0_127),
  inference(avatar_component_clause,[],[f841])).
fof(f1146,plain,(
  c3_1(a471) | (~spl0_39 | ~spl0_56 | spl0_128)),
  inference(resolution,[],[f1083,f848])).
fof(f1109,plain,(
  ~spl0_30 | ~spl0_45 | ~spl0_64 | ~spl0_65),
  inference(avatar_contradiction_clause,[],[f1108])).
fof(f1108,plain,(
  $false | (~spl0_30 | ~spl0_45 | ~spl0_64 | ~spl0_65)),
  inference(subsumption_resolution,[],[f1103,f507])).
fof(f1103,plain,(
  ~c3_1(a473) | (~spl0_30 | ~spl0_45 | ~spl0_65)),
  inference(resolution,[],[f1082,f512])).
fof(f1082,plain,(
  ( ! [X7] : (~c1_1(X7) | ~c3_1(X7)) ) | (~spl0_30 | ~spl0_45)),
  inference(subsumption_resolution,[],[f344,f410])).
fof(f1107,plain,(
  ~spl0_30 | ~spl0_45 | ~spl0_67 | ~spl0_69),
  inference(avatar_contradiction_clause,[],[f1106])).
fof(f1106,plain,(
  $false | (~spl0_30 | ~spl0_45 | ~spl0_67 | ~spl0_69)),
  inference(subsumption_resolution,[],[f1102,f523])).
fof(f523,plain,(
  c3_1(a470) | ~spl0_67),
  inference(avatar_component_clause,[],[f521])).
fof(f1102,plain,(
  ~c3_1(a470) | (~spl0_30 | ~spl0_45 | ~spl0_69)),
  inference(resolution,[],[f1082,f533])).
fof(f1105,plain,(
  ~spl0_30 | ~spl0_45 | ~spl0_101 | ~spl0_102),
  inference(avatar_contradiction_clause,[],[f1104])).
fof(f1104,plain,(
  $false | (~spl0_30 | ~spl0_45 | ~spl0_101 | ~spl0_102)),
  inference(subsumption_resolution,[],[f1099,f704])).
fof(f704,plain,(
  c3_1(a492) | ~spl0_101),
  inference(avatar_component_clause,[],[f702])).
fof(f702,plain,(
  spl0_101 <=> c3_1(a492)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_101])])).
fof(f1099,plain,(
  ~c3_1(a492) | (~spl0_30 | ~spl0_45 | ~spl0_102)),
  inference(resolution,[],[f1082,f709])).
fof(f709,plain,(
  c1_1(a492) | ~spl0_102),
  inference(avatar_component_clause,[],[f707])).
fof(f707,plain,(
  spl0_102 <=> c1_1(a492)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_102])])).
fof(f1048,plain,(
  ~spl0_32 | ~spl0_46 | ~spl0_68 | ~spl0_69),
  inference(avatar_contradiction_clause,[],[f1047])).
fof(f1047,plain,(
  $false | (~spl0_32 | ~spl0_46 | ~spl0_68 | ~spl0_69)),
  inference(subsumption_resolution,[],[f1044,f528])).
fof(f1044,plain,(
  ~c2_1(a470) | (~spl0_32 | ~spl0_46 | ~spl0_69)),
  inference(resolution,[],[f1037,f533])).
fof(f1037,plain,(
  ( ! [X9] : (~c1_1(X9) | ~c2_1(X9)) ) | (~spl0_32 | ~spl0_46)),
  inference(subsumption_resolution,[],[f352,f416])).
fof(f985,plain,(
  ~spl0_52 | ~spl0_60 | spl0_140 | spl0_141),
  inference(avatar_contradiction_clause,[],[f984])).
fof(f984,plain,(
  $false | (~spl0_52 | ~spl0_60 | spl0_140 | spl0_141)),
  inference(subsumption_resolution,[],[f980,f917])).
fof(f980,plain,(
  c0_1(a465) | (~spl0_52 | ~spl0_60 | spl0_140)),
  inference(resolution,[],[f952,f912])).
fof(f952,plain,(
  ( ! [X56] : (c2_1(X56) | c0_1(X56)) ) | (~spl0_52 | ~spl0_60)),
  inference(subsumption_resolution,[],[f444,f482])).
fof(f950,plain,(
  ~spl0_49 | spl0_147),
  inference(avatar_split_clause,[],[f8,f947,f429])).
fof(f429,plain,(
  spl0_49 <=> hskp0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_49])])).
fof(f8,plain,(
  c0_1(a460) | ~hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f6,plain,(
  (hskp12 | hskp6 | hskp20) & (hskp16 | hskp15 | hskp24) & (hskp11 | hskp5) & (hskp6 | hskp10 | hskp21) & (hskp20 | hskp23 | hskp9) & (hskp14 | hskp5 | hskp25) & (hskp16 | hskp15 | hskp8) & (hskp17 | hskp22 | hskp27) & (hskp10 | hskp26 | ! [X0] : (~c2_1(X0) | ~c0_1(X0) | c3_1(X0) | ~ndr1_0)) & (hskp12 | hskp15 | ! [X1] : (~c3_1(X1) | ~c1_1(X1) | c2_1(X1) | ~ndr1_0)) & (hskp2 | hskp15 | ! [X2] : (~c1_1(X2) | ~c0_1(X2) | c2_1(X2) | ~ndr1_0)) & (hskp20 | hskp18 | ! [X3] : (~c1_1(X3) | ~c0_1(X3) | c2_1(X3) | ~ndr1_0)) & (hskp7 | hskp11 | ! [X4] : (~c1_1(X4) | c3_1(X4) | c2_1(X4) | ~ndr1_0)) & (hskp21 | ! [X5] : (~c3_1(X5) | ~c0_1(X5) | c2_1(X5) | ~ndr1_0) | ! [X6] : (~c1_1(X6) | c3_1(X6) | c2_1(X6) | ~ndr1_0)) & (hskp3 | ! [X7] : (~c3_1(X7) | ~c1_1(X7) | ~c0_1(X7) | ~ndr1_0) | ! [X8] : (~c0_1(X8) | c3_1(X8) | c2_1(X8) | ~ndr1_0)) & (hskp1 | ! [X9] : (~c2_1(X9) | ~c1_1(X9) | ~c0_1(X9) | ~ndr1_0) | ! [X10] : (~c0_1(X10) | c3_1(X10) | c2_1(X10) | ~ndr1_0)) & (hskp12 | hskp3 | ! [X11] : (~c3_1(X11) | ~c0_1(X11) | c1_1(X11) | ~ndr1_0)) & (hskp17 | ! [X12] : (~c3_1(X12) | ~c2_1(X12) | ~c0_1(X12) | ~ndr1_0) | ! [X13] : (~c2_1(X13) | ~c0_1(X13) | c1_1(X13) | ~ndr1_0)) & (hskp9 | ! [X14] : (~c3_1(X14) | ~c0_1(X14) | c1_1(X14) | ~ndr1_0) | ! [X15] : (~c2_1(X15) | ~c0_1(X15) | c1_1(X15) | ~ndr1_0)) & (hskp14 | hskp12 | ! [X16] : (~c2_1(X16) | c3_1(X16) | c1_1(X16) | ~ndr1_0)) & (hskp20 | hskp19 | ! [X17] : (~c2_1(X17) | c3_1(X17) | c1_1(X17) | ~ndr1_0)) & (hskp2 | hskp18 | ! [X18] : (~c0_1(X18) | c3_1(X18) | c1_1(X18) | ~ndr1_0)) & (hskp16 | hskp27 | ! [X19] : (~c0_1(X19) | c3_1(X19) | c1_1(X19) | ~ndr1_0)) & (hskp12 | ! [X20] : (~c0_1(X20) | c3_1(X20) | c2_1(X20) | ~ndr1_0) | ! [X21] : (~c0_1(X21) | c3_1(X21) | c1_1(X21) | ~ndr1_0)) & (hskp12 | hskp7 | ! [X22] : (~c3_1(X22) | c2_1(X22) | c1_1(X22) | ~ndr1_0)) & (hskp17 | hskp16 | ! [X23] : (~c0_1(X23) | c2_1(X23) | c1_1(X23) | ~ndr1_0)) & (hskp15 | hskp3 | ! [X24] : (c3_1(X24) | c2_1(X24) | c1_1(X24) | ~ndr1_0)) & (hskp28 | ! [X25] : (~c3_1(X25) | ~c0_1(X25) | c2_1(X25) | ~ndr1_0) | ! [X26] : (c3_1(X26) | c2_1(X26) | c1_1(X26) | ~ndr1_0)) & (hskp10 | ! [X27] : (~c2_1(X27) | ~c1_1(X27) | ~c0_1(X27) | ~ndr1_0) | ! [X28] : (~c3_1(X28) | ~c2_1(X28) | c0_1(X28) | ~ndr1_0)) & (! [X29] : (~c1_1(X29) | ~c0_1(X29) | c2_1(X29) | ~ndr1_0) | ! [X30] : (~c2_1(X30) | c3_1(X30) | c1_1(X30) | ~ndr1_0) | ! [X31] : (~c3_1(X31) | ~c2_1(X31) | c0_1(X31) | ~ndr1_0)) & (! [X32] : (~c3_1(X32) | ~c1_1(X32) | ~c0_1(X32) | ~ndr1_0) | ! [X33] : (c3_1(X33) | c2_1(X33) | c1_1(X33) | ~ndr1_0) | ! [X34] : (~c3_1(X34) | ~c2_1(X34) | c0_1(X34) | ~ndr1_0)) & (hskp14 | ! [X35] : (~c3_1(X35) | ~c2_1(X35) | ~c0_1(X35) | ~ndr1_0) | ! [X36] : (~c3_1(X36) | ~c1_1(X36) | c0_1(X36) | ~ndr1_0)) & (hskp7 | ! [X37] : (~c2_1(X37) | ~c1_1(X37) | ~c0_1(X37) | ~ndr1_0) | ! [X38] : (~c3_1(X38) | ~c1_1(X38) | c0_1(X38) | ~ndr1_0)) & (hskp8 | ! [X39] : (~c2_1(X39) | ~c0_1(X39) | c3_1(X39) | ~ndr1_0) | ! [X40] : (~c3_1(X40) | ~c1_1(X40) | c0_1(X40) | ~ndr1_0)) & (hskp10 | hskp27 | ! [X41] : (~c2_1(X41) | ~c1_1(X41) | c0_1(X41) | ~ndr1_0)) & (hskp13 | ! [X42] : (~c0_1(X42) | c3_1(X42) | c1_1(X42) | ~ndr1_0) | ! [X43] : (~c2_1(X43) | ~c1_1(X43) | c0_1(X43) | ~ndr1_0)) & (hskp9 | ! [X44] : (~c0_1(X44) | c3_1(X44) | c1_1(X44) | ~ndr1_0) | ! [X45] : (~c2_1(X45) | ~c1_1(X45) | c0_1(X45) | ~ndr1_0)) & (hskp9 | ! [X46] : (~c0_1(X46) | c2_1(X46) | c1_1(X46) | ~ndr1_0) | ! [X47] : (~c2_1(X47) | ~c1_1(X47) | c0_1(X47) | ~ndr1_0)) & (hskp12 | hskp0 | ! [X48] : (~c2_1(X48) | c3_1(X48) | c0_1(X48) | ~ndr1_0)) & (hskp11 | ! [X49] : (~c3_1(X49) | ~c2_1(X49) | ~c0_1(X49) | ~ndr1_0) | ! [X50] : (~c1_1(X50) | c3_1(X50) | c0_1(X50) | ~ndr1_0)) & (hskp10 | ! [X51] : (~c2_1(X51) | ~c1_1(X51) | ~c0_1(X51) | ~ndr1_0) | ! [X52] : (~c3_1(X52) | c2_1(X52) | c0_1(X52) | ~ndr1_0)) & (hskp9 | ! [X53] : (~c3_1(X53) | ~c2_1(X53) | c0_1(X53) | ~ndr1_0) | ! [X54] : (~c3_1(X54) | c2_1(X54) | c0_1(X54) | ~ndr1_0)) & (hskp8 | ! [X55] : (~c3_1(X55) | ~c2_1(X55) | c1_1(X55) | ~ndr1_0) | ! [X56] : (~c1_1(X56) | c2_1(X56) | c0_1(X56) | ~ndr1_0)) & (! [X57] : (~c3_1(X57) | ~c2_1(X57) | c0_1(X57) | ~ndr1_0) | ! [X58] : (~c3_1(X58) | ~c1_1(X58) | c0_1(X58) | ~ndr1_0) | ! [X59] : (~c1_1(X59) | c2_1(X59) | c0_1(X59) | ~ndr1_0)) & (hskp7 | ! [X60] : (~c0_1(X60) | c3_1(X60) | c2_1(X60) | ~ndr1_0) | ! [X61] : (c3_1(X61) | c2_1(X61) | c0_1(X61) | ~ndr1_0)) & (hskp27 | ! [X62] : (c3_1(X62) | c2_1(X62) | c1_1(X62) | ~ndr1_0) | ! [X63] : (c3_1(X63) | c2_1(X63) | c0_1(X63) | ~ndr1_0)) & (hskp2 | hskp6 | ! [X64] : (~c2_1(X64) | c1_1(X64) | c0_1(X64) | ~ndr1_0)) & (hskp26 | ! [X65] : (~c3_1(X65) | ~c2_1(X65) | ~c0_1(X65) | ~ndr1_0) | ! [X66] : (~c2_1(X66) | c1_1(X66) | c0_1(X66) | ~ndr1_0)) & (hskp4 | ! [X67] : (~c3_1(X67) | ~c2_1(X67) | ~c1_1(X67) | ~ndr1_0) | ! [X68] : (c3_1(X68) | c1_1(X68) | c0_1(X68) | ~ndr1_0)) & (hskp5 | ! [X69] : (~c2_1(X69) | ~c1_1(X69) | ~c0_1(X69) | ~ndr1_0) | ! [X70] : (c3_1(X70) | c1_1(X70) | c0_1(X70) | ~ndr1_0)) & (hskp4 | ! [X71] : (~c3_1(X71) | ~c1_1(X71) | c2_1(X71) | ~ndr1_0) | ! [X72] : (c3_1(X72) | c1_1(X72) | c0_1(X72) | ~ndr1_0)) & (! [X73] : (~c2_1(X73) | ~c1_1(X73) | c3_1(X73) | ~ndr1_0) | ! [X74] : (c3_1(X74) | c2_1(X74) | c1_1(X74) | ~ndr1_0) | ! [X75] : (c3_1(X75) | c1_1(X75) | c0_1(X75) | ~ndr1_0)) & (hskp3 | ! [X76] : (~c3_1(X76) | ~c1_1(X76) | c0_1(X76) | ~ndr1_0) | ! [X77] : (c3_1(X77) | c1_1(X77) | c0_1(X77) | ~ndr1_0)) & (hskp2 | ! [X78] : (~c2_1(X78) | ~c1_1(X78) | c0_1(X78) | ~ndr1_0) | ! [X79] : (c3_1(X79) | c1_1(X79) | c0_1(X79) | ~ndr1_0)) & (hskp0 | ! [X80] : (c3_1(X80) | c2_1(X80) | c0_1(X80) | ~ndr1_0) | ! [X81] : (c3_1(X81) | c1_1(X81) | c0_1(X81) | ~ndr1_0)) & (hskp1 | hskp0 | ! [X82] : (c2_1(X82) | c1_1(X82) | c0_1(X82) | ~ndr1_0)) & (! [X83] : (~c2_1(X83) | ~c0_1(X83) | c3_1(X83) | ~ndr1_0) | ! [X84] : (~c3_1(X84) | ~c2_1(X84) | c1_1(X84) | ~ndr1_0) | ! [X85] : (c2_1(X85) | c1_1(X85) | c0_1(X85) | ~ndr1_0)) & (hskp25 | ! [X86] : (~c0_1(X86) | c3_1(X86) | c1_1(X86) | ~ndr1_0) | ! [X87] : (c2_1(X87) | c1_1(X87) | c0_1(X87) | ~ndr1_0)) & (! [X88] : (~c0_1(X88) | c3_1(X88) | c2_1(X88) | ~ndr1_0) | ! [X89] : (~c3_1(X89) | ~c1_1(X89) | c0_1(X89) | ~ndr1_0) | ! [X90] : (c2_1(X90) | c1_1(X90) | c0_1(X90) | ~ndr1_0)) & (hskp0 | ! [X91] : (~c2_1(X91) | c1_1(X91) | c0_1(X91) | ~ndr1_0) | ! [X92] : (c2_1(X92) | c1_1(X92) | c0_1(X92) | ~ndr1_0)) & ((c2_1(a490) & c1_1(a490) & c0_1(a490) & ndr1_0) | ~hskp28) & ((c3_1(a473) & c1_1(a473) & c0_1(a473) & ndr1_0) | ~hskp27) & ((c3_1(a470) & c2_1(a470) & c1_1(a470) & ndr1_0) | ~hskp26) & ((c3_1(a461) & c2_1(a461) & c0_1(a461) & ndr1_0) | ~hskp25) & ((~c2_1(a540) & ~c1_1(a540) & c0_1(a540) & ndr1_0) | ~hskp24) & ((~c3_1(a533) & ~c1_1(a533) & c0_1(a533) & ndr1_0) | ~hskp23) & ((~c2_1(a524) & c1_1(a524) & c0_1(a524) & ndr1_0) | ~hskp22) & ((~c1_1(a512) & c3_1(a512) & c0_1(a512) & ndr1_0) | ~hskp21) & ((~c1_1(a503) & c3_1(a503) & c2_1(a503) & ndr1_0) | ~hskp20) & ((~c0_1(a502) & c3_1(a502) & c2_1(a502) & ndr1_0) | ~hskp19) & ((~c0_1(a500) & c2_1(a500) & c1_1(a500) & ndr1_0) | ~hskp18) & ((~c3_1(a494) & ~c1_1(a494) & ~c0_1(a494) & ndr1_0) | ~hskp17) & ((~c2_1(a493) & ~c0_1(a493) & c1_1(a493) & ndr1_0) | ~hskp16) & ((~c2_1(a492) & c3_1(a492) & c1_1(a492) & ndr1_0) | ~hskp15) & ((~c3_1(a488) & ~c2_1(a488) & ~c1_1(a488) & ndr1_0) | ~hskp14) & ((~c3_1(a483) & c2_1(a483) & c0_1(a483) & ndr1_0) | ~hskp13) & ((~c2_1(a480) & ~c1_1(a480) & ~c0_1(a480) & ndr1_0) | ~hskp12) & ((~c3_1(a478) & ~c0_1(a478) & c2_1(a478) & ndr1_0) | ~hskp11) & ((~c3_1(a477) & c2_1(a477) & c1_1(a477) & ndr1_0) | ~hskp10) & ((~c1_1(a476) & c2_1(a476) & c0_1(a476) & ndr1_0) | ~hskp9) & ((~c3_1(a475) & c1_1(a475) & c0_1(a475) & ndr1_0) | ~hskp8) & ((~c2_1(a474) & ~c1_1(a474) & c3_1(a474) & ndr1_0) | ~hskp7) & ((~c3_1(a471) & ~c1_1(a471) & c2_1(a471) & ndr1_0) | ~hskp6) & ((~c2_1(a468) & c3_1(a468) & c0_1(a468) & ndr1_0) | ~hskp5) & ((~c1_1(a467) & ~c0_1(a467) & c3_1(a467) & ndr1_0) | ~hskp4) & ((~c0_1(a466) & c3_1(a466) & c1_1(a466) & ndr1_0) | ~hskp3) & ((~c3_1(a465) & ~c2_1(a465) & ~c0_1(a465) & ndr1_0) | ~hskp2) & ((~c1_1(a463) & ~c0_1(a463) & c2_1(a463) & ndr1_0) | ~hskp1) & ((~c3_1(a460) & ~c2_1(a460) & c0_1(a460) & ndr1_0) | ~hskp0)),
  inference(flattening,[],[f5])).
fof(f5,plain,(
  (hskp12 | hskp6 | hskp20) & (hskp16 | hskp15 | hskp24) & (hskp11 | hskp5) & (hskp6 | hskp10 | hskp21) & (hskp20 | hskp23 | hskp9) & (hskp14 | hskp5 | hskp25) & (hskp16 | hskp15 | hskp8) & (hskp17 | hskp22 | hskp27) & (hskp10 | hskp26 | ! [X0] : ((~c2_1(X0) | ~c0_1(X0) | c3_1(X0)) | ~ndr1_0)) & (hskp12 | hskp15 | ! [X1] : ((~c3_1(X1) | ~c1_1(X1) | c2_1(X1)) | ~ndr1_0)) & (hskp2 | hskp15 | ! [X2] : ((~c1_1(X2) | ~c0_1(X2) | c2_1(X2)) | ~ndr1_0)) & (hskp20 | hskp18 | ! [X3] : ((~c1_1(X3) | ~c0_1(X3) | c2_1(X3)) | ~ndr1_0)) & (hskp7 | hskp11 | ! [X4] : ((~c1_1(X4) | c3_1(X4) | c2_1(X4)) | ~ndr1_0)) & (hskp21 | ! [X5] : ((~c3_1(X5) | ~c0_1(X5) | c2_1(X5)) | ~ndr1_0) | ! [X6] : ((~c1_1(X6) | c3_1(X6) | c2_1(X6)) | ~ndr1_0)) & (hskp3 | ! [X7] : ((~c3_1(X7) | ~c1_1(X7) | ~c0_1(X7)) | ~ndr1_0) | ! [X8] : ((~c0_1(X8) | c3_1(X8) | c2_1(X8)) | ~ndr1_0)) & (hskp1 | ! [X9] : ((~c2_1(X9) | ~c1_1(X9) | ~c0_1(X9)) | ~ndr1_0) | ! [X10] : ((~c0_1(X10) | c3_1(X10) | c2_1(X10)) | ~ndr1_0)) & (hskp12 | hskp3 | ! [X11] : ((~c3_1(X11) | ~c0_1(X11) | c1_1(X11)) | ~ndr1_0)) & (hskp17 | ! [X12] : ((~c3_1(X12) | ~c2_1(X12) | ~c0_1(X12)) | ~ndr1_0) | ! [X13] : ((~c2_1(X13) | ~c0_1(X13) | c1_1(X13)) | ~ndr1_0)) & (hskp9 | ! [X14] : ((~c3_1(X14) | ~c0_1(X14) | c1_1(X14)) | ~ndr1_0) | ! [X15] : ((~c2_1(X15) | ~c0_1(X15) | c1_1(X15)) | ~ndr1_0)) & (hskp14 | hskp12 | ! [X16] : ((~c2_1(X16) | c3_1(X16) | c1_1(X16)) | ~ndr1_0)) & (hskp20 | hskp19 | ! [X17] : ((~c2_1(X17) | c3_1(X17) | c1_1(X17)) | ~ndr1_0)) & (hskp2 | hskp18 | ! [X18] : ((~c0_1(X18) | c3_1(X18) | c1_1(X18)) | ~ndr1_0)) & (hskp16 | hskp27 | ! [X19] : ((~c0_1(X19) | c3_1(X19) | c1_1(X19)) | ~ndr1_0)) & (hskp12 | ! [X20] : ((~c0_1(X20) | c3_1(X20) | c2_1(X20)) | ~ndr1_0) | ! [X21] : ((~c0_1(X21) | c3_1(X21) | c1_1(X21)) | ~ndr1_0)) & (hskp12 | hskp7 | ! [X22] : ((~c3_1(X22) | c2_1(X22) | c1_1(X22)) | ~ndr1_0)) & (hskp17 | hskp16 | ! [X23] : ((~c0_1(X23) | c2_1(X23) | c1_1(X23)) | ~ndr1_0)) & (hskp15 | hskp3 | ! [X24] : ((c3_1(X24) | c2_1(X24) | c1_1(X24)) | ~ndr1_0)) & (hskp28 | ! [X25] : ((~c3_1(X25) | ~c0_1(X25) | c2_1(X25)) | ~ndr1_0) | ! [X26] : ((c3_1(X26) | c2_1(X26) | c1_1(X26)) | ~ndr1_0)) & (hskp10 | ! [X27] : ((~c2_1(X27) | ~c1_1(X27) | ~c0_1(X27)) | ~ndr1_0) | ! [X28] : ((~c3_1(X28) | ~c2_1(X28) | c0_1(X28)) | ~ndr1_0)) & (! [X29] : ((~c1_1(X29) | ~c0_1(X29) | c2_1(X29)) | ~ndr1_0) | ! [X30] : ((~c2_1(X30) | c3_1(X30) | c1_1(X30)) | ~ndr1_0) | ! [X31] : ((~c3_1(X31) | ~c2_1(X31) | c0_1(X31)) | ~ndr1_0)) & (! [X32] : ((~c3_1(X32) | ~c1_1(X32) | ~c0_1(X32)) | ~ndr1_0) | ! [X33] : ((c3_1(X33) | c2_1(X33) | c1_1(X33)) | ~ndr1_0) | ! [X34] : ((~c3_1(X34) | ~c2_1(X34) | c0_1(X34)) | ~ndr1_0)) & (hskp14 | ! [X35] : ((~c3_1(X35) | ~c2_1(X35) | ~c0_1(X35)) | ~ndr1_0) | ! [X36] : ((~c3_1(X36) | ~c1_1(X36) | c0_1(X36)) | ~ndr1_0)) & (hskp7 | ! [X37] : ((~c2_1(X37) | ~c1_1(X37) | ~c0_1(X37)) | ~ndr1_0) | ! [X38] : ((~c3_1(X38) | ~c1_1(X38) | c0_1(X38)) | ~ndr1_0)) & (hskp8 | ! [X39] : ((~c2_1(X39) | ~c0_1(X39) | c3_1(X39)) | ~ndr1_0) | ! [X40] : ((~c3_1(X40) | ~c1_1(X40) | c0_1(X40)) | ~ndr1_0)) & (hskp10 | hskp27 | ! [X41] : ((~c2_1(X41) | ~c1_1(X41) | c0_1(X41)) | ~ndr1_0)) & (hskp13 | ! [X42] : ((~c0_1(X42) | c3_1(X42) | c1_1(X42)) | ~ndr1_0) | ! [X43] : ((~c2_1(X43) | ~c1_1(X43) | c0_1(X43)) | ~ndr1_0)) & (hskp9 | ! [X44] : ((~c0_1(X44) | c3_1(X44) | c1_1(X44)) | ~ndr1_0) | ! [X45] : ((~c2_1(X45) | ~c1_1(X45) | c0_1(X45)) | ~ndr1_0)) & (hskp9 | ! [X46] : ((~c0_1(X46) | c2_1(X46) | c1_1(X46)) | ~ndr1_0) | ! [X47] : ((~c2_1(X47) | ~c1_1(X47) | c0_1(X47)) | ~ndr1_0)) & (hskp12 | hskp0 | ! [X48] : ((~c2_1(X48) | c3_1(X48) | c0_1(X48)) | ~ndr1_0)) & (hskp11 | ! [X49] : ((~c3_1(X49) | ~c2_1(X49) | ~c0_1(X49)) | ~ndr1_0) | ! [X50] : ((~c1_1(X50) | c3_1(X50) | c0_1(X50)) | ~ndr1_0)) & (hskp10 | ! [X51] : ((~c2_1(X51) | ~c1_1(X51) | ~c0_1(X51)) | ~ndr1_0) | ! [X52] : ((~c3_1(X52) | c2_1(X52) | c0_1(X52)) | ~ndr1_0)) & (hskp9 | ! [X53] : ((~c3_1(X53) | ~c2_1(X53) | c0_1(X53)) | ~ndr1_0) | ! [X54] : ((~c3_1(X54) | c2_1(X54) | c0_1(X54)) | ~ndr1_0)) & (hskp8 | ! [X55] : ((~c3_1(X55) | ~c2_1(X55) | c1_1(X55)) | ~ndr1_0) | ! [X56] : ((~c1_1(X56) | c2_1(X56) | c0_1(X56)) | ~ndr1_0)) & (! [X57] : ((~c3_1(X57) | ~c2_1(X57) | c0_1(X57)) | ~ndr1_0) | ! [X58] : ((~c3_1(X58) | ~c1_1(X58) | c0_1(X58)) | ~ndr1_0) | ! [X59] : ((~c1_1(X59) | c2_1(X59) | c0_1(X59)) | ~ndr1_0)) & (hskp7 | ! [X60] : ((~c0_1(X60) | c3_1(X60) | c2_1(X60)) | ~ndr1_0) | ! [X61] : ((c3_1(X61) | c2_1(X61) | c0_1(X61)) | ~ndr1_0)) & (hskp27 | ! [X62] : ((c3_1(X62) | c2_1(X62) | c1_1(X62)) | ~ndr1_0) | ! [X63] : ((c3_1(X63) | c2_1(X63) | c0_1(X63)) | ~ndr1_0)) & (hskp2 | hskp6 | ! [X64] : ((~c2_1(X64) | c1_1(X64) | c0_1(X64)) | ~ndr1_0)) & (hskp26 | ! [X65] : ((~c3_1(X65) | ~c2_1(X65) | ~c0_1(X65)) | ~ndr1_0) | ! [X66] : ((~c2_1(X66) | c1_1(X66) | c0_1(X66)) | ~ndr1_0)) & (hskp4 | ! [X67] : ((~c3_1(X67) | ~c2_1(X67) | ~c1_1(X67)) | ~ndr1_0) | ! [X68] : ((c3_1(X68) | c1_1(X68) | c0_1(X68)) | ~ndr1_0)) & (hskp5 | ! [X69] : ((~c2_1(X69) | ~c1_1(X69) | ~c0_1(X69)) | ~ndr1_0) | ! [X70] : ((c3_1(X70) | c1_1(X70) | c0_1(X70)) | ~ndr1_0)) & (hskp4 | ! [X71] : ((~c3_1(X71) | ~c1_1(X71) | c2_1(X71)) | ~ndr1_0) | ! [X72] : ((c3_1(X72) | c1_1(X72) | c0_1(X72)) | ~ndr1_0)) & (! [X73] : ((~c2_1(X73) | ~c1_1(X73) | c3_1(X73)) | ~ndr1_0) | ! [X74] : ((c3_1(X74) | c2_1(X74) | c1_1(X74)) | ~ndr1_0) | ! [X75] : ((c3_1(X75) | c1_1(X75) | c0_1(X75)) | ~ndr1_0)) & (hskp3 | ! [X76] : ((~c3_1(X76) | ~c1_1(X76) | c0_1(X76)) | ~ndr1_0) | ! [X77] : ((c3_1(X77) | c1_1(X77) | c0_1(X77)) | ~ndr1_0)) & (hskp2 | ! [X78] : ((~c2_1(X78) | ~c1_1(X78) | c0_1(X78)) | ~ndr1_0) | ! [X79] : ((c3_1(X79) | c1_1(X79) | c0_1(X79)) | ~ndr1_0)) & (hskp0 | ! [X80] : ((c3_1(X80) | c2_1(X80) | c0_1(X80)) | ~ndr1_0) | ! [X81] : ((c3_1(X81) | c1_1(X81) | c0_1(X81)) | ~ndr1_0)) & (hskp1 | hskp0 | ! [X82] : ((c2_1(X82) | c1_1(X82) | c0_1(X82)) | ~ndr1_0)) & (! [X83] : ((~c2_1(X83) | ~c0_1(X83) | c3_1(X83)) | ~ndr1_0) | ! [X84] : ((~c3_1(X84) | ~c2_1(X84) | c1_1(X84)) | ~ndr1_0) | ! [X85] : ((c2_1(X85) | c1_1(X85) | c0_1(X85)) | ~ndr1_0)) & (hskp25 | ! [X86] : ((~c0_1(X86) | c3_1(X86) | c1_1(X86)) | ~ndr1_0) | ! [X87] : ((c2_1(X87) | c1_1(X87) | c0_1(X87)) | ~ndr1_0)) & (! [X88] : ((~c0_1(X88) | c3_1(X88) | c2_1(X88)) | ~ndr1_0) | ! [X89] : ((~c3_1(X89) | ~c1_1(X89) | c0_1(X89)) | ~ndr1_0) | ! [X90] : ((c2_1(X90) | c1_1(X90) | c0_1(X90)) | ~ndr1_0)) & (hskp0 | ! [X91] : ((~c2_1(X91) | c1_1(X91) | c0_1(X91)) | ~ndr1_0) | ! [X92] : ((c2_1(X92) | c1_1(X92) | c0_1(X92)) | ~ndr1_0)) & ((c2_1(a490) & c1_1(a490) & c0_1(a490) & ndr1_0) | ~hskp28) & ((c3_1(a473) & c1_1(a473) & c0_1(a473) & ndr1_0) | ~hskp27) & ((c3_1(a470) & c2_1(a470) & c1_1(a470) & ndr1_0) | ~hskp26) & ((c3_1(a461) & c2_1(a461) & c0_1(a461) & ndr1_0) | ~hskp25) & ((~c2_1(a540) & ~c1_1(a540) & c0_1(a540) & ndr1_0) | ~hskp24) & ((~c3_1(a533) & ~c1_1(a533) & c0_1(a533) & ndr1_0) | ~hskp23) & ((~c2_1(a524) & c1_1(a524) & c0_1(a524) & ndr1_0) | ~hskp22) & ((~c1_1(a512) & c3_1(a512) & c0_1(a512) & ndr1_0) | ~hskp21) & ((~c1_1(a503) & c3_1(a503) & c2_1(a503) & ndr1_0) | ~hskp20) & ((~c0_1(a502) & c3_1(a502) & c2_1(a502) & ndr1_0) | ~hskp19) & ((~c0_1(a500) & c2_1(a500) & c1_1(a500) & ndr1_0) | ~hskp18) & ((~c3_1(a494) & ~c1_1(a494) & ~c0_1(a494) & ndr1_0) | ~hskp17) & ((~c2_1(a493) & ~c0_1(a493) & c1_1(a493) & ndr1_0) | ~hskp16) & ((~c2_1(a492) & c3_1(a492) & c1_1(a492) & ndr1_0) | ~hskp15) & ((~c3_1(a488) & ~c2_1(a488) & ~c1_1(a488) & ndr1_0) | ~hskp14) & ((~c3_1(a483) & c2_1(a483) & c0_1(a483) & ndr1_0) | ~hskp13) & ((~c2_1(a480) & ~c1_1(a480) & ~c0_1(a480) & ndr1_0) | ~hskp12) & ((~c3_1(a478) & ~c0_1(a478) & c2_1(a478) & ndr1_0) | ~hskp11) & ((~c3_1(a477) & c2_1(a477) & c1_1(a477) & ndr1_0) | ~hskp10) & ((~c1_1(a476) & c2_1(a476) & c0_1(a476) & ndr1_0) | ~hskp9) & ((~c3_1(a475) & c1_1(a475) & c0_1(a475) & ndr1_0) | ~hskp8) & ((~c2_1(a474) & ~c1_1(a474) & c3_1(a474) & ndr1_0) | ~hskp7) & ((~c3_1(a471) & ~c1_1(a471) & c2_1(a471) & ndr1_0) | ~hskp6) & ((~c2_1(a468) & c3_1(a468) & c0_1(a468) & ndr1_0) | ~hskp5) & ((~c1_1(a467) & ~c0_1(a467) & c3_1(a467) & ndr1_0) | ~hskp4) & ((~c0_1(a466) & c3_1(a466) & c1_1(a466) & ndr1_0) | ~hskp3) & ((~c3_1(a465) & ~c2_1(a465) & ~c0_1(a465) & ndr1_0) | ~hskp2) & ((~c1_1(a463) & ~c0_1(a463) & c2_1(a463) & ndr1_0) | ~hskp1) & ((~c3_1(a460) & ~c2_1(a460) & c0_1(a460) & ndr1_0) | ~hskp0)),
  inference(ennf_transformation,[],[f4])).
fof(f4,plain,(
  (hskp12 | hskp6 | hskp20) & (hskp16 | hskp15 | hskp24) & (hskp11 | hskp5) & (hskp6 | hskp10 | hskp21) & (hskp20 | hskp23 | hskp9) & (hskp14 | hskp5 | hskp25) & (hskp16 | hskp15 | hskp8) & (hskp17 | hskp22 | hskp27) & (hskp10 | hskp26 | ! [X0] : (ndr1_0 => (~c2_1(X0) | ~c0_1(X0) | c3_1(X0)))) & (hskp12 | hskp15 | ! [X1] : (ndr1_0 => (~c3_1(X1) | ~c1_1(X1) | c2_1(X1)))) & (hskp2 | hskp15 | ! [X2] : (ndr1_0 => (~c1_1(X2) | ~c0_1(X2) | c2_1(X2)))) & (hskp20 | hskp18 | ! [X3] : (ndr1_0 => (~c1_1(X3) | ~c0_1(X3) | c2_1(X3)))) & (hskp7 | hskp11 | ! [X4] : (ndr1_0 => (~c1_1(X4) | c3_1(X4) | c2_1(X4)))) & (hskp21 | ! [X5] : (ndr1_0 => (~c3_1(X5) | ~c0_1(X5) | c2_1(X5))) | ! [X6] : (ndr1_0 => (~c1_1(X6) | c3_1(X6) | c2_1(X6)))) & (hskp3 | ! [X7] : (ndr1_0 => (~c3_1(X7) | ~c1_1(X7) | ~c0_1(X7))) | ! [X8] : (ndr1_0 => (~c0_1(X8) | c3_1(X8) | c2_1(X8)))) & (hskp1 | ! [X9] : (ndr1_0 => (~c2_1(X9) | ~c1_1(X9) | ~c0_1(X9))) | ! [X10] : (ndr1_0 => (~c0_1(X10) | c3_1(X10) | c2_1(X10)))) & (hskp12 | hskp3 | ! [X11] : (ndr1_0 => (~c3_1(X11) | ~c0_1(X11) | c1_1(X11)))) & (hskp17 | ! [X12] : (ndr1_0 => (~c3_1(X12) | ~c2_1(X12) | ~c0_1(X12))) | ! [X13] : (ndr1_0 => (~c2_1(X13) | ~c0_1(X13) | c1_1(X13)))) & (hskp9 | ! [X14] : (ndr1_0 => (~c3_1(X14) | ~c0_1(X14) | c1_1(X14))) | ! [X15] : (ndr1_0 => (~c2_1(X15) | ~c0_1(X15) | c1_1(X15)))) & (hskp14 | hskp12 | ! [X16] : (ndr1_0 => (~c2_1(X16) | c3_1(X16) | c1_1(X16)))) & (hskp20 | hskp19 | ! [X17] : (ndr1_0 => (~c2_1(X17) | c3_1(X17) | c1_1(X17)))) & (hskp2 | hskp18 | ! [X18] : (ndr1_0 => (~c0_1(X18) | c3_1(X18) | c1_1(X18)))) & (hskp16 | hskp27 | ! [X19] : (ndr1_0 => (~c0_1(X19) | c3_1(X19) | c1_1(X19)))) & (hskp12 | ! [X20] : (ndr1_0 => (~c0_1(X20) | c3_1(X20) | c2_1(X20))) | ! [X21] : (ndr1_0 => (~c0_1(X21) | c3_1(X21) | c1_1(X21)))) & (hskp12 | hskp7 | ! [X22] : (ndr1_0 => (~c3_1(X22) | c2_1(X22) | c1_1(X22)))) & (hskp17 | hskp16 | ! [X23] : (ndr1_0 => (~c0_1(X23) | c2_1(X23) | c1_1(X23)))) & (hskp15 | hskp3 | ! [X24] : (ndr1_0 => (c3_1(X24) | c2_1(X24) | c1_1(X24)))) & (hskp28 | ! [X25] : (ndr1_0 => (~c3_1(X25) | ~c0_1(X25) | c2_1(X25))) | ! [X26] : (ndr1_0 => (c3_1(X26) | c2_1(X26) | c1_1(X26)))) & (hskp10 | ! [X27] : (ndr1_0 => (~c2_1(X27) | ~c1_1(X27) | ~c0_1(X27))) | ! [X28] : (ndr1_0 => (~c3_1(X28) | ~c2_1(X28) | c0_1(X28)))) & (! [X29] : (ndr1_0 => (~c1_1(X29) | ~c0_1(X29) | c2_1(X29))) | ! [X30] : (ndr1_0 => (~c2_1(X30) | c3_1(X30) | c1_1(X30))) | ! [X31] : (ndr1_0 => (~c3_1(X31) | ~c2_1(X31) | c0_1(X31)))) & (! [X32] : (ndr1_0 => (~c3_1(X32) | ~c1_1(X32) | ~c0_1(X32))) | ! [X33] : (ndr1_0 => (c3_1(X33) | c2_1(X33) | c1_1(X33))) | ! [X34] : (ndr1_0 => (~c3_1(X34) | ~c2_1(X34) | c0_1(X34)))) & (hskp14 | ! [X35] : (ndr1_0 => (~c3_1(X35) | ~c2_1(X35) | ~c0_1(X35))) | ! [X36] : (ndr1_0 => (~c3_1(X36) | ~c1_1(X36) | c0_1(X36)))) & (hskp7 | ! [X37] : (ndr1_0 => (~c2_1(X37) | ~c1_1(X37) | ~c0_1(X37))) | ! [X38] : (ndr1_0 => (~c3_1(X38) | ~c1_1(X38) | c0_1(X38)))) & (hskp8 | ! [X39] : (ndr1_0 => (~c2_1(X39) | ~c0_1(X39) | c3_1(X39))) | ! [X40] : (ndr1_0 => (~c3_1(X40) | ~c1_1(X40) | c0_1(X40)))) & (hskp10 | hskp27 | ! [X41] : (ndr1_0 => (~c2_1(X41) | ~c1_1(X41) | c0_1(X41)))) & (hskp13 | ! [X42] : (ndr1_0 => (~c0_1(X42) | c3_1(X42) | c1_1(X42))) | ! [X43] : (ndr1_0 => (~c2_1(X43) | ~c1_1(X43) | c0_1(X43)))) & (hskp9 | ! [X44] : (ndr1_0 => (~c0_1(X44) | c3_1(X44) | c1_1(X44))) | ! [X45] : (ndr1_0 => (~c2_1(X45) | ~c1_1(X45) | c0_1(X45)))) & (hskp9 | ! [X46] : (ndr1_0 => (~c0_1(X46) | c2_1(X46) | c1_1(X46))) | ! [X47] : (ndr1_0 => (~c2_1(X47) | ~c1_1(X47) | c0_1(X47)))) & (hskp12 | hskp0 | ! [X48] : (ndr1_0 => (~c2_1(X48) | c3_1(X48) | c0_1(X48)))) & (hskp11 | ! [X49] : (ndr1_0 => (~c3_1(X49) | ~c2_1(X49) | ~c0_1(X49))) | ! [X50] : (ndr1_0 => (~c1_1(X50) | c3_1(X50) | c0_1(X50)))) & (hskp10 | ! [X51] : (ndr1_0 => (~c2_1(X51) | ~c1_1(X51) | ~c0_1(X51))) | ! [X52] : (ndr1_0 => (~c3_1(X52) | c2_1(X52) | c0_1(X52)))) & (hskp9 | ! [X53] : (ndr1_0 => (~c3_1(X53) | ~c2_1(X53) | c0_1(X53))) | ! [X54] : (ndr1_0 => (~c3_1(X54) | c2_1(X54) | c0_1(X54)))) & (hskp8 | ! [X55] : (ndr1_0 => (~c3_1(X55) | ~c2_1(X55) | c1_1(X55))) | ! [X56] : (ndr1_0 => (~c1_1(X56) | c2_1(X56) | c0_1(X56)))) & (! [X57] : (ndr1_0 => (~c3_1(X57) | ~c2_1(X57) | c0_1(X57))) | ! [X58] : (ndr1_0 => (~c3_1(X58) | ~c1_1(X58) | c0_1(X58))) | ! [X59] : (ndr1_0 => (~c1_1(X59) | c2_1(X59) | c0_1(X59)))) & (hskp7 | ! [X60] : (ndr1_0 => (~c0_1(X60) | c3_1(X60) | c2_1(X60))) | ! [X61] : (ndr1_0 => (c3_1(X61) | c2_1(X61) | c0_1(X61)))) & (hskp27 | ! [X62] : (ndr1_0 => (c3_1(X62) | c2_1(X62) | c1_1(X62))) | ! [X63] : (ndr1_0 => (c3_1(X63) | c2_1(X63) | c0_1(X63)))) & (hskp2 | hskp6 | ! [X64] : (ndr1_0 => (~c2_1(X64) | c1_1(X64) | c0_1(X64)))) & (hskp26 | ! [X65] : (ndr1_0 => (~c3_1(X65) | ~c2_1(X65) | ~c0_1(X65))) | ! [X66] : (ndr1_0 => (~c2_1(X66) | c1_1(X66) | c0_1(X66)))) & (hskp4 | ! [X67] : (ndr1_0 => (~c3_1(X67) | ~c2_1(X67) | ~c1_1(X67))) | ! [X68] : (ndr1_0 => (c3_1(X68) | c1_1(X68) | c0_1(X68)))) & (hskp5 | ! [X69] : (ndr1_0 => (~c2_1(X69) | ~c1_1(X69) | ~c0_1(X69))) | ! [X70] : (ndr1_0 => (c3_1(X70) | c1_1(X70) | c0_1(X70)))) & (hskp4 | ! [X71] : (ndr1_0 => (~c3_1(X71) | ~c1_1(X71) | c2_1(X71))) | ! [X72] : (ndr1_0 => (c3_1(X72) | c1_1(X72) | c0_1(X72)))) & (! [X73] : (ndr1_0 => (~c2_1(X73) | ~c1_1(X73) | c3_1(X73))) | ! [X74] : (ndr1_0 => (c3_1(X74) | c2_1(X74) | c1_1(X74))) | ! [X75] : (ndr1_0 => (c3_1(X75) | c1_1(X75) | c0_1(X75)))) & (hskp3 | ! [X76] : (ndr1_0 => (~c3_1(X76) | ~c1_1(X76) | c0_1(X76))) | ! [X77] : (ndr1_0 => (c3_1(X77) | c1_1(X77) | c0_1(X77)))) & (hskp2 | ! [X78] : (ndr1_0 => (~c2_1(X78) | ~c1_1(X78) | c0_1(X78))) | ! [X79] : (ndr1_0 => (c3_1(X79) | c1_1(X79) | c0_1(X79)))) & (hskp0 | ! [X80] : (ndr1_0 => (c3_1(X80) | c2_1(X80) | c0_1(X80))) | ! [X81] : (ndr1_0 => (c3_1(X81) | c1_1(X81) | c0_1(X81)))) & (hskp1 | hskp0 | ! [X82] : (ndr1_0 => (c2_1(X82) | c1_1(X82) | c0_1(X82)))) & (! [X83] : (ndr1_0 => (~c2_1(X83) | ~c0_1(X83) | c3_1(X83))) | ! [X84] : (ndr1_0 => (~c3_1(X84) | ~c2_1(X84) | c1_1(X84))) | ! [X85] : (ndr1_0 => (c2_1(X85) | c1_1(X85) | c0_1(X85)))) & (hskp25 | ! [X86] : (ndr1_0 => (~c0_1(X86) | c3_1(X86) | c1_1(X86))) | ! [X87] : (ndr1_0 => (c2_1(X87) | c1_1(X87) | c0_1(X87)))) & (! [X88] : (ndr1_0 => (~c0_1(X88) | c3_1(X88) | c2_1(X88))) | ! [X89] : (ndr1_0 => (~c3_1(X89) | ~c1_1(X89) | c0_1(X89))) | ! [X90] : (ndr1_0 => (c2_1(X90) | c1_1(X90) | c0_1(X90)))) & (hskp0 | ! [X91] : (ndr1_0 => (~c2_1(X91) | c1_1(X91) | c0_1(X91))) | ! [X92] : (ndr1_0 => (c2_1(X92) | c1_1(X92) | c0_1(X92)))) & ((c2_1(a490) & c1_1(a490) & c0_1(a490) & ndr1_0) | ~hskp28) & ((c3_1(a473) & c1_1(a473) & c0_1(a473) & ndr1_0) | ~hskp27) & ((c3_1(a470) & c2_1(a470) & c1_1(a470) & ndr1_0) | ~hskp26) & ((c3_1(a461) & c2_1(a461) & c0_1(a461) & ndr1_0) | ~hskp25) & ((~c2_1(a540) & ~c1_1(a540) & c0_1(a540) & ndr1_0) | ~hskp24) & ((~c3_1(a533) & ~c1_1(a533) & c0_1(a533) & ndr1_0) | ~hskp23) & ((~c2_1(a524) & c1_1(a524) & c0_1(a524) & ndr1_0) | ~hskp22) & ((~c1_1(a512) & c3_1(a512) & c0_1(a512) & ndr1_0) | ~hskp21) & ((~c1_1(a503) & c3_1(a503) & c2_1(a503) & ndr1_0) | ~hskp20) & ((~c0_1(a502) & c3_1(a502) & c2_1(a502) & ndr1_0) | ~hskp19) & ((~c0_1(a500) & c2_1(a500) & c1_1(a500) & ndr1_0) | ~hskp18) & ((~c3_1(a494) & ~c1_1(a494) & ~c0_1(a494) & ndr1_0) | ~hskp17) & ((~c2_1(a493) & ~c0_1(a493) & c1_1(a493) & ndr1_0) | ~hskp16) & ((~c2_1(a492) & c3_1(a492) & c1_1(a492) & ndr1_0) | ~hskp15) & ((~c3_1(a488) & ~c2_1(a488) & ~c1_1(a488) & ndr1_0) | ~hskp14) & ((~c3_1(a483) & c2_1(a483) & c0_1(a483) & ndr1_0) | ~hskp13) & ((~c2_1(a480) & ~c1_1(a480) & ~c0_1(a480) & ndr1_0) | ~hskp12) & ((~c3_1(a478) & ~c0_1(a478) & c2_1(a478) & ndr1_0) | ~hskp11) & ((~c3_1(a477) & c2_1(a477) & c1_1(a477) & ndr1_0) | ~hskp10) & ((~c1_1(a476) & c2_1(a476) & c0_1(a476) & ndr1_0) | ~hskp9) & ((~c3_1(a475) & c1_1(a475) & c0_1(a475) & ndr1_0) | ~hskp8) & ((~c2_1(a474) & ~c1_1(a474) & c3_1(a474) & ndr1_0) | ~hskp7) & ((~c3_1(a471) & ~c1_1(a471) & c2_1(a471) & ndr1_0) | ~hskp6) & ((~c2_1(a468) & c3_1(a468) & c0_1(a468) & ndr1_0) | ~hskp5) & ((~c1_1(a467) & ~c0_1(a467) & c3_1(a467) & ndr1_0) | ~hskp4) & ((~c0_1(a466) & c3_1(a466) & c1_1(a466) & ndr1_0) | ~hskp3) & ((~c3_1(a465) & ~c2_1(a465) & ~c0_1(a465) & ndr1_0) | ~hskp2) & ((~c1_1(a463) & ~c0_1(a463) & c2_1(a463) & ndr1_0) | ~hskp1) & ((~c3_1(a460) & ~c2_1(a460) & c0_1(a460) & ndr1_0) | ~hskp0)),
  inference(flattening,[],[f3])).
fof(f3,plain,(
  ~~((hskp12 | hskp6 | hskp20) & (hskp16 | hskp15 | hskp24) & (hskp11 | hskp5) & (hskp6 | hskp10 | hskp21) & (hskp20 | hskp23 | hskp9) & (hskp14 | hskp5 | hskp25) & (hskp16 | hskp15 | hskp8) & (hskp17 | hskp22 | hskp27) & (hskp10 | hskp26 | ! [X0] : (ndr1_0 => (~c2_1(X0) | ~c0_1(X0) | c3_1(X0)))) & (hskp12 | hskp15 | ! [X1] : (ndr1_0 => (~c3_1(X1) | ~c1_1(X1) | c2_1(X1)))) & (hskp2 | hskp15 | ! [X2] : (ndr1_0 => (~c1_1(X2) | ~c0_1(X2) | c2_1(X2)))) & (hskp20 | hskp18 | ! [X3] : (ndr1_0 => (~c1_1(X3) | ~c0_1(X3) | c2_1(X3)))) & (hskp7 | hskp11 | ! [X4] : (ndr1_0 => (~c1_1(X4) | c3_1(X4) | c2_1(X4)))) & (hskp21 | ! [X5] : (ndr1_0 => (~c3_1(X5) | ~c0_1(X5) | c2_1(X5))) | ! [X6] : (ndr1_0 => (~c1_1(X6) | c3_1(X6) | c2_1(X6)))) & (hskp3 | ! [X7] : (ndr1_0 => (~c3_1(X7) | ~c1_1(X7) | ~c0_1(X7))) | ! [X8] : (ndr1_0 => (~c0_1(X8) | c3_1(X8) | c2_1(X8)))) & (hskp1 | ! [X9] : (ndr1_0 => (~c2_1(X9) | ~c1_1(X9) | ~c0_1(X9))) | ! [X10] : (ndr1_0 => (~c0_1(X10) | c3_1(X10) | c2_1(X10)))) & (hskp12 | hskp3 | ! [X11] : (ndr1_0 => (~c3_1(X11) | ~c0_1(X11) | c1_1(X11)))) & (hskp17 | ! [X12] : (ndr1_0 => (~c3_1(X12) | ~c2_1(X12) | ~c0_1(X12))) | ! [X13] : (ndr1_0 => (~c2_1(X13) | ~c0_1(X13) | c1_1(X13)))) & (hskp9 | ! [X14] : (ndr1_0 => (~c3_1(X14) | ~c0_1(X14) | c1_1(X14))) | ! [X15] : (ndr1_0 => (~c2_1(X15) | ~c0_1(X15) | c1_1(X15)))) & (hskp14 | hskp12 | ! [X16] : (ndr1_0 => (~c2_1(X16) | c3_1(X16) | c1_1(X16)))) & (hskp20 | hskp19 | ! [X17] : (ndr1_0 => (~c2_1(X17) | c3_1(X17) | c1_1(X17)))) & (hskp2 | hskp18 | ! [X18] : (ndr1_0 => (~c0_1(X18) | c3_1(X18) | c1_1(X18)))) & (hskp16 | hskp27 | ! [X19] : (ndr1_0 => (~c0_1(X19) | c3_1(X19) | c1_1(X19)))) & (hskp12 | ! [X20] : (ndr1_0 => (~c0_1(X20) | c3_1(X20) | c2_1(X20))) | ! [X21] : (ndr1_0 => (~c0_1(X21) | c3_1(X21) | c1_1(X21)))) & (hskp12 | hskp7 | ! [X22] : (ndr1_0 => (~c3_1(X22) | c2_1(X22) | c1_1(X22)))) & (hskp17 | hskp16 | ! [X23] : (ndr1_0 => (~c0_1(X23) | c2_1(X23) | c1_1(X23)))) & (hskp15 | hskp3 | ! [X24] : (ndr1_0 => (c3_1(X24) | c2_1(X24) | c1_1(X24)))) & (hskp28 | ! [X25] : (ndr1_0 => (~c3_1(X25) | ~c0_1(X25) | c2_1(X25))) | ! [X26] : (ndr1_0 => (c3_1(X26) | c2_1(X26) | c1_1(X26)))) & (hskp10 | ! [X27] : (ndr1_0 => (~c2_1(X27) | ~c1_1(X27) | ~c0_1(X27))) | ! [X28] : (ndr1_0 => (~c3_1(X28) | ~c2_1(X28) | c0_1(X28)))) & (! [X29] : (ndr1_0 => (~c1_1(X29) | ~c0_1(X29) | c2_1(X29))) | ! [X30] : (ndr1_0 => (~c2_1(X30) | c3_1(X30) | c1_1(X30))) | ! [X31] : (ndr1_0 => (~c3_1(X31) | ~c2_1(X31) | c0_1(X31)))) & (! [X32] : (ndr1_0 => (~c3_1(X32) | ~c1_1(X32) | ~c0_1(X32))) | ! [X33] : (ndr1_0 => (c3_1(X33) | c2_1(X33) | c1_1(X33))) | ! [X34] : (ndr1_0 => (~c3_1(X34) | ~c2_1(X34) | c0_1(X34)))) & (hskp14 | ! [X35] : (ndr1_0 => (~c3_1(X35) | ~c2_1(X35) | ~c0_1(X35))) | ! [X36] : (ndr1_0 => (~c3_1(X36) | ~c1_1(X36) | c0_1(X36)))) & (hskp7 | ! [X37] : (ndr1_0 => (~c2_1(X37) | ~c1_1(X37) | ~c0_1(X37))) | ! [X38] : (ndr1_0 => (~c3_1(X38) | ~c1_1(X38) | c0_1(X38)))) & (hskp8 | ! [X39] : (ndr1_0 => (~c2_1(X39) | ~c0_1(X39) | c3_1(X39))) | ! [X40] : (ndr1_0 => (~c3_1(X40) | ~c1_1(X40) | c0_1(X40)))) & (hskp10 | hskp27 | ! [X41] : (ndr1_0 => (~c2_1(X41) | ~c1_1(X41) | c0_1(X41)))) & (hskp13 | ! [X42] : (ndr1_0 => (~c0_1(X42) | c3_1(X42) | c1_1(X42))) | ! [X43] : (ndr1_0 => (~c2_1(X43) | ~c1_1(X43) | c0_1(X43)))) & (hskp9 | ! [X44] : (ndr1_0 => (~c0_1(X44) | c3_1(X44) | c1_1(X44))) | ! [X45] : (ndr1_0 => (~c2_1(X45) | ~c1_1(X45) | c0_1(X45)))) & (hskp9 | ! [X46] : (ndr1_0 => (~c0_1(X46) | c2_1(X46) | c1_1(X46))) | ! [X47] : (ndr1_0 => (~c2_1(X47) | ~c1_1(X47) | c0_1(X47)))) & (hskp12 | hskp0 | ! [X48] : (ndr1_0 => (~c2_1(X48) | c3_1(X48) | c0_1(X48)))) & (hskp11 | ! [X49] : (ndr1_0 => (~c3_1(X49) | ~c2_1(X49) | ~c0_1(X49))) | ! [X50] : (ndr1_0 => (~c1_1(X50) | c3_1(X50) | c0_1(X50)))) & (hskp10 | ! [X51] : (ndr1_0 => (~c2_1(X51) | ~c1_1(X51) | ~c0_1(X51))) | ! [X52] : (ndr1_0 => (~c3_1(X52) | c2_1(X52) | c0_1(X52)))) & (hskp9 | ! [X53] : (ndr1_0 => (~c3_1(X53) | ~c2_1(X53) | c0_1(X53))) | ! [X54] : (ndr1_0 => (~c3_1(X54) | c2_1(X54) | c0_1(X54)))) & (hskp8 | ! [X55] : (ndr1_0 => (~c3_1(X55) | ~c2_1(X55) | c1_1(X55))) | ! [X56] : (ndr1_0 => (~c1_1(X56) | c2_1(X56) | c0_1(X56)))) & (! [X57] : (ndr1_0 => (~c3_1(X57) | ~c2_1(X57) | c0_1(X57))) | ! [X58] : (ndr1_0 => (~c3_1(X58) | ~c1_1(X58) | c0_1(X58))) | ! [X59] : (ndr1_0 => (~c1_1(X59) | c2_1(X59) | c0_1(X59)))) & (hskp7 | ! [X60] : (ndr1_0 => (~c0_1(X60) | c3_1(X60) | c2_1(X60))) | ! [X61] : (ndr1_0 => (c3_1(X61) | c2_1(X61) | c0_1(X61)))) & (hskp27 | ! [X62] : (ndr1_0 => (c3_1(X62) | c2_1(X62) | c1_1(X62))) | ! [X63] : (ndr1_0 => (c3_1(X63) | c2_1(X63) | c0_1(X63)))) & (hskp2 | hskp6 | ! [X64] : (ndr1_0 => (~c2_1(X64) | c1_1(X64) | c0_1(X64)))) & (hskp26 | ! [X65] : (ndr1_0 => (~c3_1(X65) | ~c2_1(X65) | ~c0_1(X65))) | ! [X66] : (ndr1_0 => (~c2_1(X66) | c1_1(X66) | c0_1(X66)))) & (hskp4 | ! [X67] : (ndr1_0 => (~c3_1(X67) | ~c2_1(X67) | ~c1_1(X67))) | ! [X68] : (ndr1_0 => (c3_1(X68) | c1_1(X68) | c0_1(X68)))) & (hskp5 | ! [X69] : (ndr1_0 => (~c2_1(X69) | ~c1_1(X69) | ~c0_1(X69))) | ! [X70] : (ndr1_0 => (c3_1(X70) | c1_1(X70) | c0_1(X70)))) & (hskp4 | ! [X71] : (ndr1_0 => (~c3_1(X71) | ~c1_1(X71) | c2_1(X71))) | ! [X72] : (ndr1_0 => (c3_1(X72) | c1_1(X72) | c0_1(X72)))) & (! [X73] : (ndr1_0 => (~c2_1(X73) | ~c1_1(X73) | c3_1(X73))) | ! [X74] : (ndr1_0 => (c3_1(X74) | c2_1(X74) | c1_1(X74))) | ! [X75] : (ndr1_0 => (c3_1(X75) | c1_1(X75) | c0_1(X75)))) & (hskp3 | ! [X76] : (ndr1_0 => (~c3_1(X76) | ~c1_1(X76) | c0_1(X76))) | ! [X77] : (ndr1_0 => (c3_1(X77) | c1_1(X77) | c0_1(X77)))) & (hskp2 | ! [X78] : (ndr1_0 => (~c2_1(X78) | ~c1_1(X78) | c0_1(X78))) | ! [X79] : (ndr1_0 => (c3_1(X79) | c1_1(X79) | c0_1(X79)))) & (hskp0 | ! [X80] : (ndr1_0 => (c3_1(X80) | c2_1(X80) | c0_1(X80))) | ! [X81] : (ndr1_0 => (c3_1(X81) | c1_1(X81) | c0_1(X81)))) & (hskp1 | hskp0 | ! [X82] : (ndr1_0 => (c2_1(X82) | c1_1(X82) | c0_1(X82)))) & (! [X83] : (ndr1_0 => (~c2_1(X83) | ~c0_1(X83) | c3_1(X83))) | ! [X84] : (ndr1_0 => (~c3_1(X84) | ~c2_1(X84) | c1_1(X84))) | ! [X85] : (ndr1_0 => (c2_1(X85) | c1_1(X85) | c0_1(X85)))) & (hskp25 | ! [X86] : (ndr1_0 => (~c0_1(X86) | c3_1(X86) | c1_1(X86))) | ! [X87] : (ndr1_0 => (c2_1(X87) | c1_1(X87) | c0_1(X87)))) & (! [X88] : (ndr1_0 => (~c0_1(X88) | c3_1(X88) | c2_1(X88))) | ! [X89] : (ndr1_0 => (~c3_1(X89) | ~c1_1(X89) | c0_1(X89))) | ! [X90] : (ndr1_0 => (c2_1(X90) | c1_1(X90) | c0_1(X90)))) & (hskp0 | ! [X91] : (ndr1_0 => (~c2_1(X91) | c1_1(X91) | c0_1(X91))) | ! [X92] : (ndr1_0 => (c2_1(X92) | c1_1(X92) | c0_1(X92)))) & ((c2_1(a490) & c1_1(a490) & c0_1(a490) & ndr1_0) | ~hskp28) & ((c3_1(a473) & c1_1(a473) & c0_1(a473) & ndr1_0) | ~hskp27) & ((c3_1(a470) & c2_1(a470) & c1_1(a470) & ndr1_0) | ~hskp26) & ((c3_1(a461) & c2_1(a461) & c0_1(a461) & ndr1_0) | ~hskp25) & ((~c2_1(a540) & ~c1_1(a540) & c0_1(a540) & ndr1_0) | ~hskp24) & ((~c3_1(a533) & ~c1_1(a533) & c0_1(a533) & ndr1_0) | ~hskp23) & ((~c2_1(a524) & c1_1(a524) & c0_1(a524) & ndr1_0) | ~hskp22) & ((~c1_1(a512) & c3_1(a512) & c0_1(a512) & ndr1_0) | ~hskp21) & ((~c1_1(a503) & c3_1(a503) & c2_1(a503) & ndr1_0) | ~hskp20) & ((~c0_1(a502) & c3_1(a502) & c2_1(a502) & ndr1_0) | ~hskp19) & ((~c0_1(a500) & c2_1(a500) & c1_1(a500) & ndr1_0) | ~hskp18) & ((~c3_1(a494) & ~c1_1(a494) & ~c0_1(a494) & ndr1_0) | ~hskp17) & ((~c2_1(a493) & ~c0_1(a493) & c1_1(a493) & ndr1_0) | ~hskp16) & ((~c2_1(a492) & c3_1(a492) & c1_1(a492) & ndr1_0) | ~hskp15) & ((~c3_1(a488) & ~c2_1(a488) & ~c1_1(a488) & ndr1_0) | ~hskp14) & ((~c3_1(a483) & c2_1(a483) & c0_1(a483) & ndr1_0) | ~hskp13) & ((~c2_1(a480) & ~c1_1(a480) & ~c0_1(a480) & ndr1_0) | ~hskp12) & ((~c3_1(a478) & ~c0_1(a478) & c2_1(a478) & ndr1_0) | ~hskp11) & ((~c3_1(a477) & c2_1(a477) & c1_1(a477) & ndr1_0) | ~hskp10) & ((~c1_1(a476) & c2_1(a476) & c0_1(a476) & ndr1_0) | ~hskp9) & ((~c3_1(a475) & c1_1(a475) & c0_1(a475) & ndr1_0) | ~hskp8) & ((~c2_1(a474) & ~c1_1(a474) & c3_1(a474) & ndr1_0) | ~hskp7) & ((~c3_1(a471) & ~c1_1(a471) & c2_1(a471) & ndr1_0) | ~hskp6) & ((~c2_1(a468) & c3_1(a468) & c0_1(a468) & ndr1_0) | ~hskp5) & ((~c1_1(a467) & ~c0_1(a467) & c3_1(a467) & ndr1_0) | ~hskp4) & ((~c0_1(a466) & c3_1(a466) & c1_1(a466) & ndr1_0) | ~hskp3) & ((~c3_1(a465) & ~c2_1(a465) & ~c0_1(a465) & ndr1_0) | ~hskp2) & ((~c1_1(a463) & ~c0_1(a463) & c2_1(a463) & ndr1_0) | ~hskp1) & ((~c3_1(a460) & ~c2_1(a460) & c0_1(a460) & ndr1_0) | ~hskp0))),
  inference(rectify,[],[f2])).
fof(f2,negated_conjecture,(
  ~~((hskp12 | hskp6 | hskp20) & (hskp16 | hskp15 | hskp24) & (hskp11 | hskp5) & (hskp6 | hskp10 | hskp21) & (hskp20 | hskp23 | hskp9) & (hskp14 | hskp5 | hskp25) & (hskp16 | hskp15 | hskp8) & (hskp17 | hskp22 | hskp27) & (hskp10 | hskp26 | ! [X92] : (ndr1_0 => (~c2_1(X92) | ~c0_1(X92) | c3_1(X92)))) & (hskp12 | hskp15 | ! [X91] : (ndr1_0 => (~c3_1(X91) | ~c1_1(X91) | c2_1(X91)))) & (hskp2 | hskp15 | ! [X90] : (ndr1_0 => (~c1_1(X90) | ~c0_1(X90) | c2_1(X90)))) & (hskp20 | hskp18 | ! [X89] : (ndr1_0 => (~c1_1(X89) | ~c0_1(X89) | c2_1(X89)))) & (hskp7 | hskp11 | ! [X88] : (ndr1_0 => (~c1_1(X88) | c3_1(X88) | c2_1(X88)))) & (hskp21 | ! [X87] : (ndr1_0 => (~c3_1(X87) | ~c0_1(X87) | c2_1(X87))) | ! [X86] : (ndr1_0 => (~c1_1(X86) | c3_1(X86) | c2_1(X86)))) & (hskp3 | ! [X85] : (ndr1_0 => (~c3_1(X85) | ~c1_1(X85) | ~c0_1(X85))) | ! [X84] : (ndr1_0 => (~c0_1(X84) | c3_1(X84) | c2_1(X84)))) & (hskp1 | ! [X83] : (ndr1_0 => (~c2_1(X83) | ~c1_1(X83) | ~c0_1(X83))) | ! [X82] : (ndr1_0 => (~c0_1(X82) | c3_1(X82) | c2_1(X82)))) & (hskp12 | hskp3 | ! [X81] : (ndr1_0 => (~c3_1(X81) | ~c0_1(X81) | c1_1(X81)))) & (hskp17 | ! [X80] : (ndr1_0 => (~c3_1(X80) | ~c2_1(X80) | ~c0_1(X80))) | ! [X79] : (ndr1_0 => (~c2_1(X79) | ~c0_1(X79) | c1_1(X79)))) & (hskp9 | ! [X78] : (ndr1_0 => (~c3_1(X78) | ~c0_1(X78) | c1_1(X78))) | ! [X77] : (ndr1_0 => (~c2_1(X77) | ~c0_1(X77) | c1_1(X77)))) & (hskp14 | hskp12 | ! [X76] : (ndr1_0 => (~c2_1(X76) | c3_1(X76) | c1_1(X76)))) & (hskp20 | hskp19 | ! [X75] : (ndr1_0 => (~c2_1(X75) | c3_1(X75) | c1_1(X75)))) & (hskp2 | hskp18 | ! [X74] : (ndr1_0 => (~c0_1(X74) | c3_1(X74) | c1_1(X74)))) & (hskp16 | hskp27 | ! [X73] : (ndr1_0 => (~c0_1(X73) | c3_1(X73) | c1_1(X73)))) & (hskp12 | ! [X72] : (ndr1_0 => (~c0_1(X72) | c3_1(X72) | c2_1(X72))) | ! [X71] : (ndr1_0 => (~c0_1(X71) | c3_1(X71) | c1_1(X71)))) & (hskp12 | hskp7 | ! [X70] : (ndr1_0 => (~c3_1(X70) | c2_1(X70) | c1_1(X70)))) & (hskp17 | hskp16 | ! [X69] : (ndr1_0 => (~c0_1(X69) | c2_1(X69) | c1_1(X69)))) & (hskp15 | hskp3 | ! [X68] : (ndr1_0 => (c3_1(X68) | c2_1(X68) | c1_1(X68)))) & (hskp28 | ! [X67] : (ndr1_0 => (~c3_1(X67) | ~c0_1(X67) | c2_1(X67))) | ! [X66] : (ndr1_0 => (c3_1(X66) | c2_1(X66) | c1_1(X66)))) & (hskp10 | ! [X65] : (ndr1_0 => (~c2_1(X65) | ~c1_1(X65) | ~c0_1(X65))) | ! [X64] : (ndr1_0 => (~c3_1(X64) | ~c2_1(X64) | c0_1(X64)))) & (! [X63] : (ndr1_0 => (~c1_1(X63) | ~c0_1(X63) | c2_1(X63))) | ! [X62] : (ndr1_0 => (~c2_1(X62) | c3_1(X62) | c1_1(X62))) | ! [X61] : (ndr1_0 => (~c3_1(X61) | ~c2_1(X61) | c0_1(X61)))) & (! [X60] : (ndr1_0 => (~c3_1(X60) | ~c1_1(X60) | ~c0_1(X60))) | ! [X59] : (ndr1_0 => (c3_1(X59) | c2_1(X59) | c1_1(X59))) | ! [X58] : (ndr1_0 => (~c3_1(X58) | ~c2_1(X58) | c0_1(X58)))) & (hskp14 | ! [X57] : (ndr1_0 => (~c3_1(X57) | ~c2_1(X57) | ~c0_1(X57))) | ! [X56] : (ndr1_0 => (~c3_1(X56) | ~c1_1(X56) | c0_1(X56)))) & (hskp7 | ! [X55] : (ndr1_0 => (~c2_1(X55) | ~c1_1(X55) | ~c0_1(X55))) | ! [X54] : (ndr1_0 => (~c3_1(X54) | ~c1_1(X54) | c0_1(X54)))) & (hskp8 | ! [X53] : (ndr1_0 => (~c2_1(X53) | ~c0_1(X53) | c3_1(X53))) | ! [X52] : (ndr1_0 => (~c3_1(X52) | ~c1_1(X52) | c0_1(X52)))) & (hskp10 | hskp27 | ! [X51] : (ndr1_0 => (~c2_1(X51) | ~c1_1(X51) | c0_1(X51)))) & (hskp13 | ! [X50] : (ndr1_0 => (~c0_1(X50) | c3_1(X50) | c1_1(X50))) | ! [X49] : (ndr1_0 => (~c2_1(X49) | ~c1_1(X49) | c0_1(X49)))) & (hskp9 | ! [X48] : (ndr1_0 => (~c0_1(X48) | c3_1(X48) | c1_1(X48))) | ! [X47] : (ndr1_0 => (~c2_1(X47) | ~c1_1(X47) | c0_1(X47)))) & (hskp9 | ! [X46] : (ndr1_0 => (~c0_1(X46) | c2_1(X46) | c1_1(X46))) | ! [X45] : (ndr1_0 => (~c2_1(X45) | ~c1_1(X45) | c0_1(X45)))) & (hskp12 | hskp0 | ! [X44] : (ndr1_0 => (~c2_1(X44) | c3_1(X44) | c0_1(X44)))) & (hskp11 | ! [X43] : (ndr1_0 => (~c3_1(X43) | ~c2_1(X43) | ~c0_1(X43))) | ! [X42] : (ndr1_0 => (~c1_1(X42) | c3_1(X42) | c0_1(X42)))) & (hskp10 | ! [X41] : (ndr1_0 => (~c2_1(X41) | ~c1_1(X41) | ~c0_1(X41))) | ! [X40] : (ndr1_0 => (~c3_1(X40) | c2_1(X40) | c0_1(X40)))) & (hskp9 | ! [X39] : (ndr1_0 => (~c3_1(X39) | ~c2_1(X39) | c0_1(X39))) | ! [X38] : (ndr1_0 => (~c3_1(X38) | c2_1(X38) | c0_1(X38)))) & (hskp8 | ! [X37] : (ndr1_0 => (~c3_1(X37) | ~c2_1(X37) | c1_1(X37))) | ! [X36] : (ndr1_0 => (~c1_1(X36) | c2_1(X36) | c0_1(X36)))) & (! [X35] : (ndr1_0 => (~c3_1(X35) | ~c2_1(X35) | c0_1(X35))) | ! [X34] : (ndr1_0 => (~c3_1(X34) | ~c1_1(X34) | c0_1(X34))) | ! [X33] : (ndr1_0 => (~c1_1(X33) | c2_1(X33) | c0_1(X33)))) & (hskp7 | ! [X32] : (ndr1_0 => (~c0_1(X32) | c3_1(X32) | c2_1(X32))) | ! [X31] : (ndr1_0 => (c3_1(X31) | c2_1(X31) | c0_1(X31)))) & (hskp27 | ! [X30] : (ndr1_0 => (c3_1(X30) | c2_1(X30) | c1_1(X30))) | ! [X29] : (ndr1_0 => (c3_1(X29) | c2_1(X29) | c0_1(X29)))) & (hskp2 | hskp6 | ! [X28] : (ndr1_0 => (~c2_1(X28) | c1_1(X28) | c0_1(X28)))) & (hskp26 | ! [X27] : (ndr1_0 => (~c3_1(X27) | ~c2_1(X27) | ~c0_1(X27))) | ! [X26] : (ndr1_0 => (~c2_1(X26) | c1_1(X26) | c0_1(X26)))) & (hskp4 | ! [X25] : (ndr1_0 => (~c3_1(X25) | ~c2_1(X25) | ~c1_1(X25))) | ! [X24] : (ndr1_0 => (c3_1(X24) | c1_1(X24) | c0_1(X24)))) & (hskp5 | ! [X23] : (ndr1_0 => (~c2_1(X23) | ~c1_1(X23) | ~c0_1(X23))) | ! [X22] : (ndr1_0 => (c3_1(X22) | c1_1(X22) | c0_1(X22)))) & (hskp4 | ! [X21] : (ndr1_0 => (~c3_1(X21) | ~c1_1(X21) | c2_1(X21))) | ! [X20] : (ndr1_0 => (c3_1(X20) | c1_1(X20) | c0_1(X20)))) & (! [X19] : (ndr1_0 => (~c2_1(X19) | ~c1_1(X19) | c3_1(X19))) | ! [X18] : (ndr1_0 => (c3_1(X18) | c2_1(X18) | c1_1(X18))) | ! [X17] : (ndr1_0 => (c3_1(X17) | c1_1(X17) | c0_1(X17)))) & (hskp3 | ! [X16] : (ndr1_0 => (~c3_1(X16) | ~c1_1(X16) | c0_1(X16))) | ! [X15] : (ndr1_0 => (c3_1(X15) | c1_1(X15) | c0_1(X15)))) & (hskp2 | ! [X14] : (ndr1_0 => (~c2_1(X14) | ~c1_1(X14) | c0_1(X14))) | ! [X13] : (ndr1_0 => (c3_1(X13) | c1_1(X13) | c0_1(X13)))) & (hskp0 | ! [X12] : (ndr1_0 => (c3_1(X12) | c2_1(X12) | c0_1(X12))) | ! [X11] : (ndr1_0 => (c3_1(X11) | c1_1(X11) | c0_1(X11)))) & (hskp1 | hskp0 | ! [X10] : (ndr1_0 => (c2_1(X10) | c1_1(X10) | c0_1(X10)))) & (! [X9] : (ndr1_0 => (~c2_1(X9) | ~c0_1(X9) | c3_1(X9))) | ! [X8] : (ndr1_0 => (~c3_1(X8) | ~c2_1(X8) | c1_1(X8))) | ! [X7] : (ndr1_0 => (c2_1(X7) | c1_1(X7) | c0_1(X7)))) & (hskp25 | ! [X6] : (ndr1_0 => (~c0_1(X6) | c3_1(X6) | c1_1(X6))) | ! [X5] : (ndr1_0 => (c2_1(X5) | c1_1(X5) | c0_1(X5)))) & (! [X4] : (ndr1_0 => (~c0_1(X4) | c3_1(X4) | c2_1(X4))) | ! [X3] : (ndr1_0 => (~c3_1(X3) | ~c1_1(X3) | c0_1(X3))) | ! [X2] : (ndr1_0 => (c2_1(X2) | c1_1(X2) | c0_1(X2)))) & (hskp0 | ! [X1] : (ndr1_0 => (~c2_1(X1) | c1_1(X1) | c0_1(X1))) | ! [X0] : (ndr1_0 => (c2_1(X0) | c1_1(X0) | c0_1(X0)))) & ((c2_1(a490) & c1_1(a490) & c0_1(a490) & ndr1_0) | ~hskp28) & ((c3_1(a473) & c1_1(a473) & c0_1(a473) & ndr1_0) | ~hskp27) & ((c3_1(a470) & c2_1(a470) & c1_1(a470) & ndr1_0) | ~hskp26) & ((c3_1(a461) & c2_1(a461) & c0_1(a461) & ndr1_0) | ~hskp25) & ((~c2_1(a540) & ~c1_1(a540) & c0_1(a540) & ndr1_0) | ~hskp24) & ((~c3_1(a533) & ~c1_1(a533) & c0_1(a533) & ndr1_0) | ~hskp23) & ((~c2_1(a524) & c1_1(a524) & c0_1(a524) & ndr1_0) | ~hskp22) & ((~c1_1(a512) & c3_1(a512) & c0_1(a512) & ndr1_0) | ~hskp21) & ((~c1_1(a503) & c3_1(a503) & c2_1(a503) & ndr1_0) | ~hskp20) & ((~c0_1(a502) & c3_1(a502) & c2_1(a502) & ndr1_0) | ~hskp19) & ((~c0_1(a500) & c2_1(a500) & c1_1(a500) & ndr1_0) | ~hskp18) & ((~c3_1(a494) & ~c1_1(a494) & ~c0_1(a494) & ndr1_0) | ~hskp17) & ((~c2_1(a493) & ~c0_1(a493) & c1_1(a493) & ndr1_0) | ~hskp16) & ((~c2_1(a492) & c3_1(a492) & c1_1(a492) & ndr1_0) | ~hskp15) & ((~c3_1(a488) & ~c2_1(a488) & ~c1_1(a488) & ndr1_0) | ~hskp14) & ((~c3_1(a483) & c2_1(a483) & c0_1(a483) & ndr1_0) | ~hskp13) & ((~c2_1(a480) & ~c1_1(a480) & ~c0_1(a480) & ndr1_0) | ~hskp12) & ((~c3_1(a478) & ~c0_1(a478) & c2_1(a478) & ndr1_0) | ~hskp11) & ((~c3_1(a477) & c2_1(a477) & c1_1(a477) & ndr1_0) | ~hskp10) & ((~c1_1(a476) & c2_1(a476) & c0_1(a476) & ndr1_0) | ~hskp9) & ((~c3_1(a475) & c1_1(a475) & c0_1(a475) & ndr1_0) | ~hskp8) & ((~c2_1(a474) & ~c1_1(a474) & c3_1(a474) & ndr1_0) | ~hskp7) & ((~c3_1(a471) & ~c1_1(a471) & c2_1(a471) & ndr1_0) | ~hskp6) & ((~c2_1(a468) & c3_1(a468) & c0_1(a468) & ndr1_0) | ~hskp5) & ((~c1_1(a467) & ~c0_1(a467) & c3_1(a467) & ndr1_0) | ~hskp4) & ((~c0_1(a466) & c3_1(a466) & c1_1(a466) & ndr1_0) | ~hskp3) & ((~c3_1(a465) & ~c2_1(a465) & ~c0_1(a465) & ndr1_0) | ~hskp2) & ((~c1_1(a463) & ~c0_1(a463) & c2_1(a463) & ndr1_0) | ~hskp1) & ((~c3_1(a460) & ~c2_1(a460) & c0_1(a460) & ndr1_0) | ~hskp0))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  ~((hskp12 | hskp6 | hskp20) & (hskp16 | hskp15 | hskp24) & (hskp11 | hskp5) & (hskp6 | hskp10 | hskp21) & (hskp20 | hskp23 | hskp9) & (hskp14 | hskp5 | hskp25) & (hskp16 | hskp15 | hskp8) & (hskp17 | hskp22 | hskp27) & (hskp10 | hskp26 | ! [X92] : (ndr1_0 => (~c2_1(X92) | ~c0_1(X92) | c3_1(X92)))) & (hskp12 | hskp15 | ! [X91] : (ndr1_0 => (~c3_1(X91) | ~c1_1(X91) | c2_1(X91)))) & (hskp2 | hskp15 | ! [X90] : (ndr1_0 => (~c1_1(X90) | ~c0_1(X90) | c2_1(X90)))) & (hskp20 | hskp18 | ! [X89] : (ndr1_0 => (~c1_1(X89) | ~c0_1(X89) | c2_1(X89)))) & (hskp7 | hskp11 | ! [X88] : (ndr1_0 => (~c1_1(X88) | c3_1(X88) | c2_1(X88)))) & (hskp21 | ! [X87] : (ndr1_0 => (~c3_1(X87) | ~c0_1(X87) | c2_1(X87))) | ! [X86] : (ndr1_0 => (~c1_1(X86) | c3_1(X86) | c2_1(X86)))) & (hskp3 | ! [X85] : (ndr1_0 => (~c3_1(X85) | ~c1_1(X85) | ~c0_1(X85))) | ! [X84] : (ndr1_0 => (~c0_1(X84) | c3_1(X84) | c2_1(X84)))) & (hskp1 | ! [X83] : (ndr1_0 => (~c2_1(X83) | ~c1_1(X83) | ~c0_1(X83))) | ! [X82] : (ndr1_0 => (~c0_1(X82) | c3_1(X82) | c2_1(X82)))) & (hskp12 | hskp3 | ! [X81] : (ndr1_0 => (~c3_1(X81) | ~c0_1(X81) | c1_1(X81)))) & (hskp17 | ! [X80] : (ndr1_0 => (~c3_1(X80) | ~c2_1(X80) | ~c0_1(X80))) | ! [X79] : (ndr1_0 => (~c2_1(X79) | ~c0_1(X79) | c1_1(X79)))) & (hskp9 | ! [X78] : (ndr1_0 => (~c3_1(X78) | ~c0_1(X78) | c1_1(X78))) | ! [X77] : (ndr1_0 => (~c2_1(X77) | ~c0_1(X77) | c1_1(X77)))) & (hskp14 | hskp12 | ! [X76] : (ndr1_0 => (~c2_1(X76) | c3_1(X76) | c1_1(X76)))) & (hskp20 | hskp19 | ! [X75] : (ndr1_0 => (~c2_1(X75) | c3_1(X75) | c1_1(X75)))) & (hskp2 | hskp18 | ! [X74] : (ndr1_0 => (~c0_1(X74) | c3_1(X74) | c1_1(X74)))) & (hskp16 | hskp27 | ! [X73] : (ndr1_0 => (~c0_1(X73) | c3_1(X73) | c1_1(X73)))) & (hskp12 | ! [X72] : (ndr1_0 => (~c0_1(X72) | c3_1(X72) | c2_1(X72))) | ! [X71] : (ndr1_0 => (~c0_1(X71) | c3_1(X71) | c1_1(X71)))) & (hskp12 | hskp7 | ! [X70] : (ndr1_0 => (~c3_1(X70) | c2_1(X70) | c1_1(X70)))) & (hskp17 | hskp16 | ! [X69] : (ndr1_0 => (~c0_1(X69) | c2_1(X69) | c1_1(X69)))) & (hskp15 | hskp3 | ! [X68] : (ndr1_0 => (c3_1(X68) | c2_1(X68) | c1_1(X68)))) & (hskp28 | ! [X67] : (ndr1_0 => (~c3_1(X67) | ~c0_1(X67) | c2_1(X67))) | ! [X66] : (ndr1_0 => (c3_1(X66) | c2_1(X66) | c1_1(X66)))) & (hskp10 | ! [X65] : (ndr1_0 => (~c2_1(X65) | ~c1_1(X65) | ~c0_1(X65))) | ! [X64] : (ndr1_0 => (~c3_1(X64) | ~c2_1(X64) | c0_1(X64)))) & (! [X63] : (ndr1_0 => (~c1_1(X63) | ~c0_1(X63) | c2_1(X63))) | ! [X62] : (ndr1_0 => (~c2_1(X62) | c3_1(X62) | c1_1(X62))) | ! [X61] : (ndr1_0 => (~c3_1(X61) | ~c2_1(X61) | c0_1(X61)))) & (! [X60] : (ndr1_0 => (~c3_1(X60) | ~c1_1(X60) | ~c0_1(X60))) | ! [X59] : (ndr1_0 => (c3_1(X59) | c2_1(X59) | c1_1(X59))) | ! [X58] : (ndr1_0 => (~c3_1(X58) | ~c2_1(X58) | c0_1(X58)))) & (hskp14 | ! [X57] : (ndr1_0 => (~c3_1(X57) | ~c2_1(X57) | ~c0_1(X57))) | ! [X56] : (ndr1_0 => (~c3_1(X56) | ~c1_1(X56) | c0_1(X56)))) & (hskp7 | ! [X55] : (ndr1_0 => (~c2_1(X55) | ~c1_1(X55) | ~c0_1(X55))) | ! [X54] : (ndr1_0 => (~c3_1(X54) | ~c1_1(X54) | c0_1(X54)))) & (hskp8 | ! [X53] : (ndr1_0 => (~c2_1(X53) | ~c0_1(X53) | c3_1(X53))) | ! [X52] : (ndr1_0 => (~c3_1(X52) | ~c1_1(X52) | c0_1(X52)))) & (hskp10 | hskp27 | ! [X51] : (ndr1_0 => (~c2_1(X51) | ~c1_1(X51) | c0_1(X51)))) & (hskp13 | ! [X50] : (ndr1_0 => (~c0_1(X50) | c3_1(X50) | c1_1(X50))) | ! [X49] : (ndr1_0 => (~c2_1(X49) | ~c1_1(X49) | c0_1(X49)))) & (hskp9 | ! [X48] : (ndr1_0 => (~c0_1(X48) | c3_1(X48) | c1_1(X48))) | ! [X47] : (ndr1_0 => (~c2_1(X47) | ~c1_1(X47) | c0_1(X47)))) & (hskp9 | ! [X46] : (ndr1_0 => (~c0_1(X46) | c2_1(X46) | c1_1(X46))) | ! [X45] : (ndr1_0 => (~c2_1(X45) | ~c1_1(X45) | c0_1(X45)))) & (hskp12 | hskp0 | ! [X44] : (ndr1_0 => (~c2_1(X44) | c3_1(X44) | c0_1(X44)))) & (hskp11 | ! [X43] : (ndr1_0 => (~c3_1(X43) | ~c2_1(X43) | ~c0_1(X43))) | ! [X42] : (ndr1_0 => (~c1_1(X42) | c3_1(X42) | c0_1(X42)))) & (hskp10 | ! [X41] : (ndr1_0 => (~c2_1(X41) | ~c1_1(X41) | ~c0_1(X41))) | ! [X40] : (ndr1_0 => (~c3_1(X40) | c2_1(X40) | c0_1(X40)))) & (hskp9 | ! [X39] : (ndr1_0 => (~c3_1(X39) | ~c2_1(X39) | c0_1(X39))) | ! [X38] : (ndr1_0 => (~c3_1(X38) | c2_1(X38) | c0_1(X38)))) & (hskp8 | ! [X37] : (ndr1_0 => (~c3_1(X37) | ~c2_1(X37) | c1_1(X37))) | ! [X36] : (ndr1_0 => (~c1_1(X36) | c2_1(X36) | c0_1(X36)))) & (! [X35] : (ndr1_0 => (~c3_1(X35) | ~c2_1(X35) | c0_1(X35))) | ! [X34] : (ndr1_0 => (~c3_1(X34) | ~c1_1(X34) | c0_1(X34))) | ! [X33] : (ndr1_0 => (~c1_1(X33) | c2_1(X33) | c0_1(X33)))) & (hskp7 | ! [X32] : (ndr1_0 => (~c0_1(X32) | c3_1(X32) | c2_1(X32))) | ! [X31] : (ndr1_0 => (c3_1(X31) | c2_1(X31) | c0_1(X31)))) & (hskp27 | ! [X30] : (ndr1_0 => (c3_1(X30) | c2_1(X30) | c1_1(X30))) | ! [X29] : (ndr1_0 => (c3_1(X29) | c2_1(X29) | c0_1(X29)))) & (hskp2 | hskp6 | ! [X28] : (ndr1_0 => (~c2_1(X28) | c1_1(X28) | c0_1(X28)))) & (hskp26 | ! [X27] : (ndr1_0 => (~c3_1(X27) | ~c2_1(X27) | ~c0_1(X27))) | ! [X26] : (ndr1_0 => (~c2_1(X26) | c1_1(X26) | c0_1(X26)))) & (hskp4 | ! [X25] : (ndr1_0 => (~c3_1(X25) | ~c2_1(X25) | ~c1_1(X25))) | ! [X24] : (ndr1_0 => (c3_1(X24) | c1_1(X24) | c0_1(X24)))) & (hskp5 | ! [X23] : (ndr1_0 => (~c2_1(X23) | ~c1_1(X23) | ~c0_1(X23))) | ! [X22] : (ndr1_0 => (c3_1(X22) | c1_1(X22) | c0_1(X22)))) & (hskp4 | ! [X21] : (ndr1_0 => (~c3_1(X21) | ~c1_1(X21) | c2_1(X21))) | ! [X20] : (ndr1_0 => (c3_1(X20) | c1_1(X20) | c0_1(X20)))) & (! [X19] : (ndr1_0 => (~c2_1(X19) | ~c1_1(X19) | c3_1(X19))) | ! [X18] : (ndr1_0 => (c3_1(X18) | c2_1(X18) | c1_1(X18))) | ! [X17] : (ndr1_0 => (c3_1(X17) | c1_1(X17) | c0_1(X17)))) & (hskp3 | ! [X16] : (ndr1_0 => (~c3_1(X16) | ~c1_1(X16) | c0_1(X16))) | ! [X15] : (ndr1_0 => (c3_1(X15) | c1_1(X15) | c0_1(X15)))) & (hskp2 | ! [X14] : (ndr1_0 => (~c2_1(X14) | ~c1_1(X14) | c0_1(X14))) | ! [X13] : (ndr1_0 => (c3_1(X13) | c1_1(X13) | c0_1(X13)))) & (hskp0 | ! [X12] : (ndr1_0 => (c3_1(X12) | c2_1(X12) | c0_1(X12))) | ! [X11] : (ndr1_0 => (c3_1(X11) | c1_1(X11) | c0_1(X11)))) & (hskp1 | hskp0 | ! [X10] : (ndr1_0 => (c2_1(X10) | c1_1(X10) | c0_1(X10)))) & (! [X9] : (ndr1_0 => (~c2_1(X9) | ~c0_1(X9) | c3_1(X9))) | ! [X8] : (ndr1_0 => (~c3_1(X8) | ~c2_1(X8) | c1_1(X8))) | ! [X7] : (ndr1_0 => (c2_1(X7) | c1_1(X7) | c0_1(X7)))) & (hskp25 | ! [X6] : (ndr1_0 => (~c0_1(X6) | c3_1(X6) | c1_1(X6))) | ! [X5] : (ndr1_0 => (c2_1(X5) | c1_1(X5) | c0_1(X5)))) & (! [X4] : (ndr1_0 => (~c0_1(X4) | c3_1(X4) | c2_1(X4))) | ! [X3] : (ndr1_0 => (~c3_1(X3) | ~c1_1(X3) | c0_1(X3))) | ! [X2] : (ndr1_0 => (c2_1(X2) | c1_1(X2) | c0_1(X2)))) & (hskp0 | ! [X1] : (ndr1_0 => (~c2_1(X1) | c1_1(X1) | c0_1(X1))) | ! [X0] : (ndr1_0 => (c2_1(X0) | c1_1(X0) | c0_1(X0)))) & ((c2_1(a490) & c1_1(a490) & c0_1(a490) & ndr1_0) | ~hskp28) & ((c3_1(a473) & c1_1(a473) & c0_1(a473) & ndr1_0) | ~hskp27) & ((c3_1(a470) & c2_1(a470) & c1_1(a470) & ndr1_0) | ~hskp26) & ((c3_1(a461) & c2_1(a461) & c0_1(a461) & ndr1_0) | ~hskp25) & ((~c2_1(a540) & ~c1_1(a540) & c0_1(a540) & ndr1_0) | ~hskp24) & ((~c3_1(a533) & ~c1_1(a533) & c0_1(a533) & ndr1_0) | ~hskp23) & ((~c2_1(a524) & c1_1(a524) & c0_1(a524) & ndr1_0) | ~hskp22) & ((~c1_1(a512) & c3_1(a512) & c0_1(a512) & ndr1_0) | ~hskp21) & ((~c1_1(a503) & c3_1(a503) & c2_1(a503) & ndr1_0) | ~hskp20) & ((~c0_1(a502) & c3_1(a502) & c2_1(a502) & ndr1_0) | ~hskp19) & ((~c0_1(a500) & c2_1(a500) & c1_1(a500) & ndr1_0) | ~hskp18) & ((~c3_1(a494) & ~c1_1(a494) & ~c0_1(a494) & ndr1_0) | ~hskp17) & ((~c2_1(a493) & ~c0_1(a493) & c1_1(a493) & ndr1_0) | ~hskp16) & ((~c2_1(a492) & c3_1(a492) & c1_1(a492) & ndr1_0) | ~hskp15) & ((~c3_1(a488) & ~c2_1(a488) & ~c1_1(a488) & ndr1_0) | ~hskp14) & ((~c3_1(a483) & c2_1(a483) & c0_1(a483) & ndr1_0) | ~hskp13) & ((~c2_1(a480) & ~c1_1(a480) & ~c0_1(a480) & ndr1_0) | ~hskp12) & ((~c3_1(a478) & ~c0_1(a478) & c2_1(a478) & ndr1_0) | ~hskp11) & ((~c3_1(a477) & c2_1(a477) & c1_1(a477) & ndr1_0) | ~hskp10) & ((~c1_1(a476) & c2_1(a476) & c0_1(a476) & ndr1_0) | ~hskp9) & ((~c3_1(a475) & c1_1(a475) & c0_1(a475) & ndr1_0) | ~hskp8) & ((~c2_1(a474) & ~c1_1(a474) & c3_1(a474) & ndr1_0) | ~hskp7) & ((~c3_1(a471) & ~c1_1(a471) & c2_1(a471) & ndr1_0) | ~hskp6) & ((~c2_1(a468) & c3_1(a468) & c0_1(a468) & ndr1_0) | ~hskp5) & ((~c1_1(a467) & ~c0_1(a467) & c3_1(a467) & ndr1_0) | ~hskp4) & ((~c0_1(a466) & c3_1(a466) & c1_1(a466) & ndr1_0) | ~hskp3) & ((~c3_1(a465) & ~c2_1(a465) & ~c0_1(a465) & ndr1_0) | ~hskp2) & ((~c1_1(a463) & ~c0_1(a463) & c2_1(a463) & ndr1_0) | ~hskp1) & ((~c3_1(a460) & ~c2_1(a460) & c0_1(a460) & ndr1_0) | ~hskp0))),
  file('SYN448+1.tptp',unknown)).
fof(f945,plain,(
  ~spl0_49 | ~spl0_146),
  inference(avatar_split_clause,[],[f9,f942,f429])).
fof(f9,plain,(
  ~c2_1(a460) | ~hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f940,plain,(
  ~spl0_49 | ~spl0_145),
  inference(avatar_split_clause,[],[f10,f937,f429])).
fof(f10,plain,(
  ~c3_1(a460) | ~hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f934,plain,(
  ~spl0_33 | spl0_144),
  inference(avatar_split_clause,[],[f12,f931,f354])).
fof(f354,plain,(
  spl0_33 <=> hskp1),
  introduced(avatar_definition,[new_symbols(naming,[spl0_33])])).
fof(f12,plain,(
  c2_1(a463) | ~hskp1),
  inference(cnf_transformation,[],[f6])).
fof(f924,plain,(
  ~spl0_33 | ~spl0_142),
  inference(avatar_split_clause,[],[f14,f921,f354])).
fof(f14,plain,(
  ~c1_1(a463) | ~hskp1),
  inference(cnf_transformation,[],[f6])).
fof(f918,plain,(
  ~spl0_24 | ~spl0_141),
  inference(avatar_split_clause,[],[f16,f915,f318])).
fof(f318,plain,(
  spl0_24 <=> hskp2),
  introduced(avatar_definition,[new_symbols(naming,[spl0_24])])).
fof(f16,plain,(
  ~c0_1(a465) | ~hskp2),
  inference(cnf_transformation,[],[f6])).
fof(f913,plain,(
  ~spl0_24 | ~spl0_140),
  inference(avatar_split_clause,[],[f17,f910,f318])).
fof(f17,plain,(
  ~c2_1(a465) | ~hskp2),
  inference(cnf_transformation,[],[f6])).
fof(f908,plain,(
  ~spl0_24 | ~spl0_139),
  inference(avatar_split_clause,[],[f18,f905,f318])).
fof(f18,plain,(
  ~c3_1(a465) | ~hskp2),
  inference(cnf_transformation,[],[f6])).
fof(f902,plain,(
  ~spl0_31 | spl0_138),
  inference(avatar_split_clause,[],[f20,f899,f346])).
fof(f346,plain,(
  spl0_31 <=> hskp3),
  introduced(avatar_definition,[new_symbols(naming,[spl0_31])])).
fof(f20,plain,(
  c1_1(a466) | ~hskp3),
  inference(cnf_transformation,[],[f6])).
fof(f897,plain,(
  ~spl0_31 | spl0_137),
  inference(avatar_split_clause,[],[f21,f894,f346])).
fof(f21,plain,(
  c3_1(a466) | ~hskp3),
  inference(cnf_transformation,[],[f6])).
fof(f892,plain,(
  ~spl0_31 | ~spl0_136),
  inference(avatar_split_clause,[],[f22,f889,f346])).
fof(f22,plain,(
  ~c0_1(a466) | ~hskp3),
  inference(cnf_transformation,[],[f6])).
fof(f871,plain,(
  ~spl0_7 | spl0_19),
  inference(avatar_split_clause,[],[f27,f299,f245])).
fof(f245,plain,(
  spl0_7 <=> hskp5),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f299,plain,(
  spl0_19 <=> ndr1_0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f27,plain,(
  ndr1_0 | ~hskp5),
  inference(cnf_transformation,[],[f6])).
fof(f870,plain,(
  ~spl0_7 | spl0_132),
  inference(avatar_split_clause,[],[f28,f867,f245])).
fof(f28,plain,(
  c0_1(a468) | ~hskp5),
  inference(cnf_transformation,[],[f6])).
fof(f865,plain,(
  ~spl0_7 | spl0_131),
  inference(avatar_split_clause,[],[f29,f862,f245])).
fof(f29,plain,(
  c3_1(a468) | ~hskp5),
  inference(cnf_transformation,[],[f6])).
fof(f860,plain,(
  ~spl0_7 | ~spl0_130),
  inference(avatar_split_clause,[],[f30,f857,f245])).
fof(f30,plain,(
  ~c2_1(a468) | ~hskp5),
  inference(cnf_transformation,[],[f6])).
fof(f854,plain,(
  ~spl0_2 | spl0_129),
  inference(avatar_split_clause,[],[f32,f851,f223])).
fof(f223,plain,(
  spl0_2 <=> hskp6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f32,plain,(
  c2_1(a471) | ~hskp6),
  inference(cnf_transformation,[],[f6])).
fof(f849,plain,(
  ~spl0_2 | ~spl0_128),
  inference(avatar_split_clause,[],[f33,f846,f223])).
fof(f33,plain,(
  ~c1_1(a471) | ~hskp6),
  inference(cnf_transformation,[],[f6])).
fof(f844,plain,(
  ~spl0_2 | ~spl0_127),
  inference(avatar_split_clause,[],[f34,f841,f223])).
fof(f34,plain,(
  ~c3_1(a471) | ~hskp6),
  inference(cnf_transformation,[],[f6])).
fof(f806,plain,(
  ~spl0_11 | spl0_120),
  inference(avatar_split_clause,[],[f44,f803,f263])).
fof(f263,plain,(
  spl0_11 <=> hskp9),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f44,plain,(
  c0_1(a476) | ~hskp9),
  inference(cnf_transformation,[],[f6])).
fof(f801,plain,(
  ~spl0_11 | spl0_119),
  inference(avatar_split_clause,[],[f45,f798,f263])).
fof(f45,plain,(
  c2_1(a476) | ~hskp9),
  inference(cnf_transformation,[],[f6])).
fof(f796,plain,(
  ~spl0_11 | ~spl0_118),
  inference(avatar_split_clause,[],[f46,f793,f263])).
fof(f46,plain,(
  ~c1_1(a476) | ~hskp9),
  inference(cnf_transformation,[],[f6])).
fof(f790,plain,(
  ~spl0_10 | spl0_117),
  inference(avatar_split_clause,[],[f48,f787,f258])).
fof(f258,plain,(
  spl0_10 <=> hskp10),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f48,plain,(
  c1_1(a477) | ~hskp10),
  inference(cnf_transformation,[],[f6])).
fof(f785,plain,(
  ~spl0_10 | spl0_116),
  inference(avatar_split_clause,[],[f49,f782,f258])).
fof(f49,plain,(
  c2_1(a477) | ~hskp10),
  inference(cnf_transformation,[],[f6])).
fof(f780,plain,(
  ~spl0_10 | ~spl0_115),
  inference(avatar_split_clause,[],[f50,f777,f258])).
fof(f50,plain,(
  ~c3_1(a477) | ~hskp10),
  inference(cnf_transformation,[],[f6])).
fof(f775,plain,(
  ~spl0_8 | spl0_19),
  inference(avatar_split_clause,[],[f51,f299,f249])).
fof(f249,plain,(
  spl0_8 <=> hskp11),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f51,plain,(
  ndr1_0 | ~hskp11),
  inference(cnf_transformation,[],[f6])).
fof(f774,plain,(
  ~spl0_8 | spl0_114),
  inference(avatar_split_clause,[],[f52,f771,f249])).
fof(f52,plain,(
  c2_1(a478) | ~hskp11),
  inference(cnf_transformation,[],[f6])).
fof(f769,plain,(
  ~spl0_8 | ~spl0_113),
  inference(avatar_split_clause,[],[f53,f766,f249])).
fof(f53,plain,(
  ~c0_1(a478) | ~hskp11),
  inference(cnf_transformation,[],[f6])).
fof(f764,plain,(
  ~spl0_8 | ~spl0_112),
  inference(avatar_split_clause,[],[f54,f761,f249])).
fof(f54,plain,(
  ~c3_1(a478) | ~hskp11),
  inference(cnf_transformation,[],[f6])).
fof(f758,plain,(
  ~spl0_3 | ~spl0_111),
  inference(avatar_split_clause,[],[f56,f755,f227])).
fof(f227,plain,(
  spl0_3 <=> hskp12),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f56,plain,(
  ~c0_1(a480) | ~hskp12),
  inference(cnf_transformation,[],[f6])).
fof(f753,plain,(
  ~spl0_3 | ~spl0_110),
  inference(avatar_split_clause,[],[f57,f750,f227])).
fof(f57,plain,(
  ~c1_1(a480) | ~hskp12),
  inference(cnf_transformation,[],[f6])).
fof(f748,plain,(
  ~spl0_3 | ~spl0_109),
  inference(avatar_split_clause,[],[f58,f745,f227])).
fof(f58,plain,(
  ~c2_1(a480) | ~hskp12),
  inference(cnf_transformation,[],[f6])).
fof(f726,plain,(
  ~spl0_14 | ~spl0_105),
  inference(avatar_split_clause,[],[f64,f723,f276])).
fof(f276,plain,(
  spl0_14 <=> hskp14),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f64,plain,(
  ~c1_1(a488) | ~hskp14),
  inference(cnf_transformation,[],[f6])).
fof(f721,plain,(
  ~spl0_14 | ~spl0_104),
  inference(avatar_split_clause,[],[f65,f718,f276])).
fof(f65,plain,(
  ~c2_1(a488) | ~hskp14),
  inference(cnf_transformation,[],[f6])).
fof(f716,plain,(
  ~spl0_14 | ~spl0_103),
  inference(avatar_split_clause,[],[f66,f713,f276])).
fof(f66,plain,(
  ~c3_1(a488) | ~hskp14),
  inference(cnf_transformation,[],[f6])).
fof(f710,plain,(
  ~spl0_5 | spl0_102),
  inference(avatar_split_clause,[],[f68,f707,f236])).
fof(f236,plain,(
  spl0_5 <=> hskp15),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f68,plain,(
  c1_1(a492) | ~hskp15),
  inference(cnf_transformation,[],[f6])).
fof(f705,plain,(
  ~spl0_5 | spl0_101),
  inference(avatar_split_clause,[],[f69,f702,f236])).
fof(f69,plain,(
  c3_1(a492) | ~hskp15),
  inference(cnf_transformation,[],[f6])).
fof(f694,plain,(
  ~spl0_6 | spl0_99),
  inference(avatar_split_clause,[],[f72,f691,f240])).
fof(f240,plain,(
  spl0_6 <=> hskp16),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f72,plain,(
  c1_1(a493) | ~hskp16),
  inference(cnf_transformation,[],[f6])).
fof(f689,plain,(
  ~spl0_6 | ~spl0_98),
  inference(avatar_split_clause,[],[f73,f686,f240])).
fof(f73,plain,(
  ~c0_1(a493) | ~hskp16),
  inference(cnf_transformation,[],[f6])).
fof(f684,plain,(
  ~spl0_6 | ~spl0_97),
  inference(avatar_split_clause,[],[f74,f681,f240])).
fof(f74,plain,(
  ~c2_1(a493) | ~hskp16),
  inference(cnf_transformation,[],[f6])).
fof(f678,plain,(
  ~spl0_18 | ~spl0_96),
  inference(avatar_split_clause,[],[f76,f675,f294])).
fof(f294,plain,(
  spl0_18 <=> hskp17),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f76,plain,(
  ~c0_1(a494) | ~hskp17),
  inference(cnf_transformation,[],[f6])).
fof(f673,plain,(
  ~spl0_18 | ~spl0_95),
  inference(avatar_split_clause,[],[f77,f670,f294])).
fof(f77,plain,(
  ~c1_1(a494) | ~hskp17),
  inference(cnf_transformation,[],[f6])).
fof(f668,plain,(
  ~spl0_18 | ~spl0_94),
  inference(avatar_split_clause,[],[f78,f665,f294])).
fof(f78,plain,(
  ~c3_1(a494) | ~hskp17),
  inference(cnf_transformation,[],[f6])).
fof(f662,plain,(
  ~spl0_25 | spl0_93),
  inference(avatar_split_clause,[],[f80,f659,f323])).
fof(f323,plain,(
  spl0_25 <=> hskp18),
  introduced(avatar_definition,[new_symbols(naming,[spl0_25])])).
fof(f80,plain,(
  c1_1(a500) | ~hskp18),
  inference(cnf_transformation,[],[f6])).
fof(f657,plain,(
  ~spl0_25 | spl0_92),
  inference(avatar_split_clause,[],[f81,f654,f323])).
fof(f81,plain,(
  c2_1(a500) | ~hskp18),
  inference(cnf_transformation,[],[f6])).
fof(f652,plain,(
  ~spl0_25 | ~spl0_91),
  inference(avatar_split_clause,[],[f82,f649,f323])).
fof(f82,plain,(
  ~c0_1(a500) | ~hskp18),
  inference(cnf_transformation,[],[f6])).
fof(f646,plain,(
  ~spl0_38 | spl0_90),
  inference(avatar_split_clause,[],[f84,f643,f375])).
fof(f375,plain,(
  spl0_38 <=> hskp19),
  introduced(avatar_definition,[new_symbols(naming,[spl0_38])])).
fof(f84,plain,(
  c2_1(a502) | ~hskp19),
  inference(cnf_transformation,[],[f6])).
fof(f641,plain,(
  ~spl0_38 | spl0_89),
  inference(avatar_split_clause,[],[f85,f638,f375])).
fof(f85,plain,(
  c3_1(a502) | ~hskp19),
  inference(cnf_transformation,[],[f6])).
fof(f636,plain,(
  ~spl0_38 | ~spl0_88),
  inference(avatar_split_clause,[],[f86,f633,f375])).
fof(f86,plain,(
  ~c0_1(a502) | ~hskp19),
  inference(cnf_transformation,[],[f6])).
fof(f630,plain,(
  ~spl0_1 | spl0_87),
  inference(avatar_split_clause,[],[f88,f627,f219])).
fof(f219,plain,(
  spl0_1 <=> hskp20),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f88,plain,(
  c2_1(a503) | ~hskp20),
  inference(cnf_transformation,[],[f6])).
fof(f625,plain,(
  ~spl0_1 | spl0_86),
  inference(avatar_split_clause,[],[f89,f622,f219])).
fof(f89,plain,(
  c3_1(a503) | ~hskp20),
  inference(cnf_transformation,[],[f6])).
fof(f620,plain,(
  ~spl0_1 | ~spl0_85),
  inference(avatar_split_clause,[],[f90,f617,f219])).
fof(f90,plain,(
  ~c1_1(a503) | ~hskp20),
  inference(cnf_transformation,[],[f6])).
fof(f614,plain,(
  ~spl0_9 | spl0_84),
  inference(avatar_split_clause,[],[f92,f611,f254])).
fof(f254,plain,(
  spl0_9 <=> hskp21),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f92,plain,(
  c0_1(a512) | ~hskp21),
  inference(cnf_transformation,[],[f6])).
fof(f609,plain,(
  ~spl0_9 | spl0_83),
  inference(avatar_split_clause,[],[f93,f606,f254])).
fof(f93,plain,(
  c3_1(a512) | ~hskp21),
  inference(cnf_transformation,[],[f6])).
fof(f604,plain,(
  ~spl0_9 | ~spl0_82),
  inference(avatar_split_clause,[],[f94,f601,f254])).
fof(f94,plain,(
  ~c1_1(a512) | ~hskp21),
  inference(cnf_transformation,[],[f6])).
fof(f582,plain,(
  ~spl0_12 | spl0_78),
  inference(avatar_split_clause,[],[f100,f579,f267])).
fof(f267,plain,(
  spl0_12 <=> hskp23),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f100,plain,(
  c0_1(a533) | ~hskp23),
  inference(cnf_transformation,[],[f6])).
fof(f577,plain,(
  ~spl0_12 | ~spl0_77),
  inference(avatar_split_clause,[],[f101,f574,f267])).
fof(f101,plain,(
  ~c1_1(a533) | ~hskp23),
  inference(cnf_transformation,[],[f6])).
fof(f572,plain,(
  ~spl0_12 | ~spl0_76),
  inference(avatar_split_clause,[],[f102,f569,f267])).
fof(f102,plain,(
  ~c3_1(a533) | ~hskp23),
  inference(cnf_transformation,[],[f6])).
fof(f550,plain,(
  ~spl0_13 | spl0_72),
  inference(avatar_split_clause,[],[f108,f547,f272])).
fof(f272,plain,(
  spl0_13 <=> hskp25),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f108,plain,(
  c0_1(a461) | ~hskp25),
  inference(cnf_transformation,[],[f6])).
fof(f545,plain,(
  ~spl0_13 | spl0_71),
  inference(avatar_split_clause,[],[f109,f542,f272])).
fof(f109,plain,(
  c2_1(a461) | ~hskp25),
  inference(cnf_transformation,[],[f6])).
fof(f540,plain,(
  ~spl0_13 | spl0_70),
  inference(avatar_split_clause,[],[f110,f537,f272])).
fof(f110,plain,(
  c3_1(a461) | ~hskp25),
  inference(cnf_transformation,[],[f6])).
fof(f534,plain,(
  ~spl0_21 | spl0_69),
  inference(avatar_split_clause,[],[f112,f531,f306])).
fof(f306,plain,(
  spl0_21 <=> hskp26),
  introduced(avatar_definition,[new_symbols(naming,[spl0_21])])).
fof(f112,plain,(
  c1_1(a470) | ~hskp26),
  inference(cnf_transformation,[],[f6])).
fof(f529,plain,(
  ~spl0_21 | spl0_68),
  inference(avatar_split_clause,[],[f113,f526,f306])).
fof(f113,plain,(
  c2_1(a470) | ~hskp26),
  inference(cnf_transformation,[],[f6])).
fof(f524,plain,(
  ~spl0_21 | spl0_67),
  inference(avatar_split_clause,[],[f114,f521,f306])).
fof(f114,plain,(
  c3_1(a470) | ~hskp26),
  inference(cnf_transformation,[],[f6])).
fof(f518,plain,(
  ~spl0_16 | spl0_66),
  inference(avatar_split_clause,[],[f116,f515,f286])).
fof(f286,plain,(
  spl0_16 <=> hskp27),
  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).
fof(f116,plain,(
  c0_1(a473) | ~hskp27),
  inference(cnf_transformation,[],[f6])).
fof(f513,plain,(
  ~spl0_16 | spl0_65),
  inference(avatar_split_clause,[],[f117,f510,f286])).
fof(f117,plain,(
  c1_1(a473) | ~hskp27),
  inference(cnf_transformation,[],[f6])).
fof(f508,plain,(
  ~spl0_16 | spl0_64),
  inference(avatar_split_clause,[],[f118,f505,f286])).
fof(f118,plain,(
  c3_1(a473) | ~hskp27),
  inference(cnf_transformation,[],[f6])).
fof(f502,plain,(
  ~spl0_43 | spl0_63),
  inference(avatar_split_clause,[],[f120,f499,f398])).
fof(f398,plain,(
  spl0_43 <=> hskp28),
  introduced(avatar_definition,[new_symbols(naming,[spl0_43])])).
fof(f120,plain,(
  c0_1(a490) | ~hskp28),
  inference(cnf_transformation,[],[f6])).
fof(f497,plain,(
  ~spl0_43 | spl0_62),
  inference(avatar_split_clause,[],[f121,f494,f398])).
fof(f121,plain,(
  c1_1(a490) | ~hskp28),
  inference(cnf_transformation,[],[f6])).
fof(f492,plain,(
  ~spl0_43 | spl0_61),
  inference(avatar_split_clause,[],[f122,f489,f398])).
fof(f122,plain,(
  c2_1(a490) | ~hskp28),
  inference(cnf_transformation,[],[f6])).
fof(f487,plain,(
  spl0_60 | ~spl0_19 | spl0_55 | spl0_49),
  inference(avatar_split_clause,[],[f183,f429,f456,f299,f481])).
fof(f183,plain,(
  ( ! [X92,X91] : (hskp0 | ~c2_1(X91) | c1_1(X91) | c0_1(X91) | ~ndr1_0 | c2_1(X92) | c1_1(X92) | c0_1(X92)) )),
  inference(duplicate_literal_removal,[],[f123])).
fof(f123,plain,(
  ( ! [X92,X91] : (hskp0 | ~c2_1(X91) | c1_1(X91) | c0_1(X91) | ~ndr1_0 | c2_1(X92) | c1_1(X92) | c0_1(X92) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f486,plain,(
  spl0_60 | spl0_45 | ~spl0_19 | spl0_29),
  inference(avatar_split_clause,[],[f184,f340,f299,f409,f481])).
fof(f184,plain,(
  ( ! [X90,X88,X89] : (~c0_1(X88) | c3_1(X88) | c2_1(X88) | ~ndr1_0 | ~c3_1(X89) | ~c1_1(X89) | c0_1(X89) | c2_1(X90) | c1_1(X90) | c0_1(X90)) )),
  inference(duplicate_literal_removal,[],[f124])).
fof(f124,plain,(
  ( ! [X90,X88,X89] : (~c0_1(X88) | c3_1(X88) | c2_1(X88) | ~ndr1_0 | ~c3_1(X89) | ~c1_1(X89) | c0_1(X89) | ~ndr1_0 | c2_1(X90) | c1_1(X90) | c0_1(X90) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f485,plain,(
  spl0_60 | ~spl0_19 | spl0_39 | spl0_13),
  inference(avatar_split_clause,[],[f185,f272,f380,f299,f481])).
fof(f185,plain,(
  ( ! [X87,X86] : (hskp25 | ~c0_1(X86) | c3_1(X86) | c1_1(X86) | ~ndr1_0 | c2_1(X87) | c1_1(X87) | c0_1(X87)) )),
  inference(duplicate_literal_removal,[],[f125])).
fof(f125,plain,(
  ( ! [X87,X86] : (hskp25 | ~c0_1(X86) | c3_1(X86) | c1_1(X86) | ~ndr1_0 | c2_1(X87) | c1_1(X87) | c0_1(X87) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f484,plain,(
  spl0_60 | spl0_53 | ~spl0_19 | spl0_20),
  inference(avatar_split_clause,[],[f186,f303,f299,f446,f481])).
fof(f186,plain,(
  ( ! [X85,X83,X84] : (~c2_1(X83) | ~c0_1(X83) | c3_1(X83) | ~ndr1_0 | ~c3_1(X84) | ~c2_1(X84) | c1_1(X84) | c2_1(X85) | c1_1(X85) | c0_1(X85)) )),
  inference(duplicate_literal_removal,[],[f126])).
fof(f126,plain,(
  ( ! [X85,X83,X84] : (~c2_1(X83) | ~c0_1(X83) | c3_1(X83) | ~ndr1_0 | ~c3_1(X84) | ~c2_1(X84) | c1_1(X84) | ~ndr1_0 | c2_1(X85) | c1_1(X85) | c0_1(X85) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f483,plain,(
  ~spl0_19 | spl0_60 | spl0_49 | spl0_33),
  inference(avatar_split_clause,[],[f127,f354,f429,f481,f299])).
fof(f127,plain,(
  ( ! [X82] : (hskp1 | hskp0 | c2_1(X82) | c1_1(X82) | c0_1(X82) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f479,plain,(
  spl0_56 | ~spl0_19 | spl0_54 | spl0_49),
  inference(avatar_split_clause,[],[f187,f429,f451,f299,f461])).
fof(f187,plain,(
  ( ! [X80,X81] : (hskp0 | c3_1(X80) | c2_1(X80) | c0_1(X80) | ~ndr1_0 | c3_1(X81) | c1_1(X81) | c0_1(X81)) )),
  inference(duplicate_literal_removal,[],[f128])).
fof(f128,plain,(
  ( ! [X80,X81] : (hskp0 | c3_1(X80) | c2_1(X80) | c0_1(X80) | ~ndr1_0 | c3_1(X81) | c1_1(X81) | c0_1(X81) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f478,plain,(
  spl0_56 | ~spl0_19 | spl0_46 | spl0_24),
  inference(avatar_split_clause,[],[f188,f318,f415,f299,f461])).
fof(f188,plain,(
  ( ! [X78,X79] : (hskp2 | ~c2_1(X78) | ~c1_1(X78) | c0_1(X78) | ~ndr1_0 | c3_1(X79) | c1_1(X79) | c0_1(X79)) )),
  inference(duplicate_literal_removal,[],[f129])).
fof(f129,plain,(
  ( ! [X78,X79] : (hskp2 | ~c2_1(X78) | ~c1_1(X78) | c0_1(X78) | ~ndr1_0 | c3_1(X79) | c1_1(X79) | c0_1(X79) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f477,plain,(
  spl0_56 | ~spl0_19 | spl0_45 | spl0_31),
  inference(avatar_split_clause,[],[f189,f346,f409,f299,f461])).
fof(f189,plain,(
  ( ! [X76,X77] : (hskp3 | ~c3_1(X76) | ~c1_1(X76) | c0_1(X76) | ~ndr1_0 | c3_1(X77) | c1_1(X77) | c0_1(X77)) )),
  inference(duplicate_literal_removal,[],[f130])).
fof(f130,plain,(
  ( ! [X76,X77] : (hskp3 | ~c3_1(X76) | ~c1_1(X76) | c0_1(X76) | ~ndr1_0 | c3_1(X77) | c1_1(X77) | c0_1(X77) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f476,plain,(
  spl0_56 | spl0_42 | ~spl0_19 | spl0_59),
  inference(avatar_split_clause,[],[f190,f474,f299,f394,f461])).
fof(f190,plain,(
  ( ! [X74,X75,X73] : (~c2_1(X73) | ~c1_1(X73) | c3_1(X73) | ~ndr1_0 | c3_1(X74) | c2_1(X74) | c1_1(X74) | c3_1(X75) | c1_1(X75) | c0_1(X75)) )),
  inference(duplicate_literal_removal,[],[f131])).
fof(f131,plain,(
  ( ! [X74,X75,X73] : (~c2_1(X73) | ~c1_1(X73) | c3_1(X73) | ~ndr1_0 | c3_1(X74) | c2_1(X74) | c1_1(X74) | ~ndr1_0 | c3_1(X75) | c1_1(X75) | c0_1(X75) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f471,plain,(
  spl0_56 | ~spl0_19 | spl0_32 | spl0_7),
  inference(avatar_split_clause,[],[f192,f245,f351,f299,f461])).
fof(f192,plain,(
  ( ! [X70,X69] : (hskp5 | ~c2_1(X69) | ~c1_1(X69) | ~c0_1(X69) | ~ndr1_0 | c3_1(X70) | c1_1(X70) | c0_1(X70)) )),
  inference(duplicate_literal_removal,[],[f133])).
fof(f133,plain,(
  ( ! [X70,X69] : (hskp5 | ~c2_1(X69) | ~c1_1(X69) | ~c0_1(X69) | ~ndr1_0 | c3_1(X70) | c1_1(X70) | c0_1(X70) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f459,plain,(
  spl0_55 | ~spl0_19 | spl0_36 | spl0_21),
  inference(avatar_split_clause,[],[f194,f306,f366,f299,f456])).
fof(f194,plain,(
  ( ! [X66,X65] : (hskp26 | ~c3_1(X65) | ~c2_1(X65) | ~c0_1(X65) | ~ndr1_0 | ~c2_1(X66) | c1_1(X66) | c0_1(X66)) )),
  inference(duplicate_literal_removal,[],[f135])).
fof(f135,plain,(
  ( ! [X66,X65] : (hskp26 | ~c3_1(X65) | ~c2_1(X65) | ~c0_1(X65) | ~ndr1_0 | ~c2_1(X66) | c1_1(X66) | c0_1(X66) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f458,plain,(
  ~spl0_19 | spl0_55 | spl0_2 | spl0_24),
  inference(avatar_split_clause,[],[f136,f318,f223,f456,f299])).
fof(f136,plain,(
  ( ! [X64] : (hskp2 | hskp6 | ~c2_1(X64) | c1_1(X64) | c0_1(X64) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f454,plain,(
  spl0_54 | ~spl0_19 | spl0_42 | spl0_16),
  inference(avatar_split_clause,[],[f195,f286,f394,f299,f451])).
fof(f195,plain,(
  ( ! [X62,X63] : (hskp27 | c3_1(X62) | c2_1(X62) | c1_1(X62) | ~ndr1_0 | c3_1(X63) | c2_1(X63) | c0_1(X63)) )),
  inference(duplicate_literal_removal,[],[f137])).
fof(f137,plain,(
  ( ! [X62,X63] : (hskp27 | c3_1(X62) | c2_1(X62) | c1_1(X62) | ~ndr1_0 | c3_1(X63) | c2_1(X63) | c0_1(X63) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f449,plain,(
  spl0_52 | spl0_45 | ~spl0_19 | spl0_44),
  inference(avatar_split_clause,[],[f197,f403,f299,f409,f443])).
fof(f197,plain,(
  ( ! [X59,X57,X58] : (~c3_1(X57) | ~c2_1(X57) | c0_1(X57) | ~ndr1_0 | ~c3_1(X58) | ~c1_1(X58) | c0_1(X58) | ~c1_1(X59) | c2_1(X59) | c0_1(X59)) )),
  inference(duplicate_literal_removal,[],[f139])).
fof(f139,plain,(
  ( ! [X59,X57,X58] : (~c3_1(X57) | ~c2_1(X57) | c0_1(X57) | ~ndr1_0 | ~c3_1(X58) | ~c1_1(X58) | c0_1(X58) | ~ndr1_0 | ~c1_1(X59) | c2_1(X59) | c0_1(X59) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f441,plain,(
  spl0_51 | ~spl0_19 | spl0_44 | spl0_11),
  inference(avatar_split_clause,[],[f199,f263,f403,f299,f438])).
fof(f199,plain,(
  ( ! [X54,X53] : (hskp9 | ~c3_1(X53) | ~c2_1(X53) | c0_1(X53) | ~ndr1_0 | ~c3_1(X54) | c2_1(X54) | c0_1(X54)) )),
  inference(duplicate_literal_removal,[],[f141])).
fof(f141,plain,(
  ( ! [X54,X53] : (hskp9 | ~c3_1(X53) | ~c2_1(X53) | c0_1(X53) | ~ndr1_0 | ~c3_1(X54) | c2_1(X54) | c0_1(X54) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f440,plain,(
  spl0_51 | ~spl0_19 | spl0_32 | spl0_10),
  inference(avatar_split_clause,[],[f200,f258,f351,f299,f438])).
fof(f200,plain,(
  ( ! [X52,X51] : (hskp10 | ~c2_1(X51) | ~c1_1(X51) | ~c0_1(X51) | ~ndr1_0 | ~c3_1(X52) | c2_1(X52) | c0_1(X52)) )),
  inference(duplicate_literal_removal,[],[f142])).
fof(f142,plain,(
  ( ! [X52,X51] : (hskp10 | ~c2_1(X51) | ~c1_1(X51) | ~c0_1(X51) | ~ndr1_0 | ~c3_1(X52) | c2_1(X52) | c0_1(X52) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f436,plain,(
  spl0_50 | ~spl0_19 | spl0_36 | spl0_8),
  inference(avatar_split_clause,[],[f201,f249,f366,f299,f434])).
fof(f201,plain,(
  ( ! [X50,X49] : (hskp11 | ~c3_1(X49) | ~c2_1(X49) | ~c0_1(X49) | ~ndr1_0 | ~c1_1(X50) | c3_1(X50) | c0_1(X50)) )),
  inference(duplicate_literal_removal,[],[f143])).
fof(f143,plain,(
  ( ! [X50,X49] : (hskp11 | ~c3_1(X49) | ~c2_1(X49) | ~c0_1(X49) | ~ndr1_0 | ~c1_1(X50) | c3_1(X50) | c0_1(X50) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f432,plain,(
  ~spl0_19 | spl0_48 | spl0_49 | spl0_3),
  inference(avatar_split_clause,[],[f144,f227,f429,f426,f299])).
fof(f144,plain,(
  ( ! [X48] : (hskp12 | hskp0 | ~c2_1(X48) | c3_1(X48) | c0_1(X48) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f424,plain,(
  spl0_46 | ~spl0_19 | spl0_41 | spl0_11),
  inference(avatar_split_clause,[],[f202,f263,f390,f299,f415])).
fof(f202,plain,(
  ( ! [X47,X46] : (hskp9 | ~c0_1(X46) | c2_1(X46) | c1_1(X46) | ~ndr1_0 | ~c2_1(X47) | ~c1_1(X47) | c0_1(X47)) )),
  inference(duplicate_literal_removal,[],[f145])).
fof(f145,plain,(
  ( ! [X47,X46] : (hskp9 | ~c0_1(X46) | c2_1(X46) | c1_1(X46) | ~ndr1_0 | ~c2_1(X47) | ~c1_1(X47) | c0_1(X47) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f423,plain,(
  spl0_46 | ~spl0_19 | spl0_39 | spl0_11),
  inference(avatar_split_clause,[],[f203,f263,f380,f299,f415])).
fof(f203,plain,(
  ( ! [X45,X44] : (hskp9 | ~c0_1(X44) | c3_1(X44) | c1_1(X44) | ~ndr1_0 | ~c2_1(X45) | ~c1_1(X45) | c0_1(X45)) )),
  inference(duplicate_literal_removal,[],[f146])).
fof(f146,plain,(
  ( ! [X45,X44] : (hskp9 | ~c0_1(X44) | c3_1(X44) | c1_1(X44) | ~ndr1_0 | ~c2_1(X45) | ~c1_1(X45) | c0_1(X45) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f417,plain,(
  ~spl0_19 | spl0_46 | spl0_16 | spl0_10),
  inference(avatar_split_clause,[],[f148,f258,f286,f415,f299])).
fof(f148,plain,(
  ( ! [X41] : (hskp10 | hskp27 | ~c2_1(X41) | ~c1_1(X41) | c0_1(X41) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f411,plain,(
  spl0_45 | ~spl0_19 | spl0_36 | spl0_14),
  inference(avatar_split_clause,[],[f207,f276,f366,f299,f409])).
fof(f207,plain,(
  ( ! [X35,X36] : (hskp14 | ~c3_1(X35) | ~c2_1(X35) | ~c0_1(X35) | ~ndr1_0 | ~c3_1(X36) | ~c1_1(X36) | c0_1(X36)) )),
  inference(duplicate_literal_removal,[],[f151])).
fof(f151,plain,(
  ( ! [X35,X36] : (hskp14 | ~c3_1(X35) | ~c2_1(X35) | ~c0_1(X35) | ~ndr1_0 | ~c3_1(X36) | ~c1_1(X36) | c0_1(X36) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f407,plain,(
  spl0_44 | spl0_42 | ~spl0_19 | spl0_30),
  inference(avatar_split_clause,[],[f208,f343,f299,f394,f403])).
fof(f208,plain,(
  ( ! [X33,X34,X32] : (~c3_1(X32) | ~c1_1(X32) | ~c0_1(X32) | ~ndr1_0 | c3_1(X33) | c2_1(X33) | c1_1(X33) | ~c3_1(X34) | ~c2_1(X34) | c0_1(X34)) )),
  inference(duplicate_literal_removal,[],[f152])).
fof(f152,plain,(
  ( ! [X33,X34,X32] : (~c3_1(X32) | ~c1_1(X32) | ~c0_1(X32) | ~ndr1_0 | c3_1(X33) | c2_1(X33) | c1_1(X33) | ~ndr1_0 | ~c3_1(X34) | ~c2_1(X34) | c0_1(X34) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f406,plain,(
  spl0_44 | spl0_37 | ~spl0_19 | spl0_23),
  inference(avatar_split_clause,[],[f209,f315,f299,f371,f403])).
fof(f209,plain,(
  ( ! [X30,X31,X29] : (~c1_1(X29) | ~c0_1(X29) | c2_1(X29) | ~ndr1_0 | ~c2_1(X30) | c3_1(X30) | c1_1(X30) | ~c3_1(X31) | ~c2_1(X31) | c0_1(X31)) )),
  inference(duplicate_literal_removal,[],[f153])).
fof(f153,plain,(
  ( ! [X30,X31,X29] : (~c1_1(X29) | ~c0_1(X29) | c2_1(X29) | ~ndr1_0 | ~c2_1(X30) | c3_1(X30) | c1_1(X30) | ~ndr1_0 | ~c3_1(X31) | ~c2_1(X31) | c0_1(X31) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f405,plain,(
  spl0_44 | ~spl0_19 | spl0_32 | spl0_10),
  inference(avatar_split_clause,[],[f210,f258,f351,f299,f403])).
fof(f210,plain,(
  ( ! [X28,X27] : (hskp10 | ~c2_1(X27) | ~c1_1(X27) | ~c0_1(X27) | ~ndr1_0 | ~c3_1(X28) | ~c2_1(X28) | c0_1(X28)) )),
  inference(duplicate_literal_removal,[],[f154])).
fof(f154,plain,(
  ( ! [X28,X27] : (hskp10 | ~c2_1(X27) | ~c1_1(X27) | ~c0_1(X27) | ~ndr1_0 | ~c3_1(X28) | ~c2_1(X28) | c0_1(X28) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f401,plain,(
  spl0_42 | ~spl0_19 | spl0_28 | spl0_43),
  inference(avatar_split_clause,[],[f211,f398,f336,f299,f394])).
fof(f211,plain,(
  ( ! [X26,X25] : (hskp28 | ~c3_1(X25) | ~c0_1(X25) | c2_1(X25) | ~ndr1_0 | c3_1(X26) | c2_1(X26) | c1_1(X26)) )),
  inference(duplicate_literal_removal,[],[f155])).
fof(f155,plain,(
  ( ! [X26,X25] : (hskp28 | ~c3_1(X25) | ~c0_1(X25) | c2_1(X25) | ~ndr1_0 | c3_1(X26) | c2_1(X26) | c1_1(X26) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f396,plain,(
  ~spl0_19 | spl0_42 | spl0_31 | spl0_5),
  inference(avatar_split_clause,[],[f156,f236,f346,f394,f299])).
fof(f156,plain,(
  ( ! [X24] : (hskp15 | hskp3 | c3_1(X24) | c2_1(X24) | c1_1(X24) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f392,plain,(
  ~spl0_19 | spl0_41 | spl0_6 | spl0_18),
  inference(avatar_split_clause,[],[f157,f294,f240,f390,f299])).
fof(f157,plain,(
  ( ! [X23] : (hskp17 | hskp16 | ~c0_1(X23) | c2_1(X23) | c1_1(X23) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f384,plain,(
  spl0_39 | ~spl0_19 | spl0_29 | spl0_3),
  inference(avatar_split_clause,[],[f212,f227,f340,f299,f380])).
fof(f212,plain,(
  ( ! [X21,X20] : (hskp12 | ~c0_1(X20) | c3_1(X20) | c2_1(X20) | ~ndr1_0 | ~c0_1(X21) | c3_1(X21) | c1_1(X21)) )),
  inference(duplicate_literal_removal,[],[f159])).
fof(f159,plain,(
  ( ! [X21,X20] : (hskp12 | ~c0_1(X20) | c3_1(X20) | c2_1(X20) | ~ndr1_0 | ~c0_1(X21) | c3_1(X21) | c1_1(X21) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f382,plain,(
  ~spl0_19 | spl0_39 | spl0_25 | spl0_24),
  inference(avatar_split_clause,[],[f161,f318,f323,f380,f299])).
fof(f161,plain,(
  ( ! [X18] : (hskp2 | hskp18 | ~c0_1(X18) | c3_1(X18) | c1_1(X18) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f378,plain,(
  ~spl0_19 | spl0_37 | spl0_38 | spl0_1),
  inference(avatar_split_clause,[],[f162,f219,f375,f371,f299])).
fof(f162,plain,(
  ( ! [X17] : (hskp20 | hskp19 | ~c2_1(X17) | c3_1(X17) | c1_1(X17) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f373,plain,(
  ~spl0_19 | spl0_37 | spl0_3 | spl0_14),
  inference(avatar_split_clause,[],[f163,f276,f227,f371,f299])).
fof(f163,plain,(
  ( ! [X16] : (hskp14 | hskp12 | ~c2_1(X16) | c3_1(X16) | c1_1(X16) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f369,plain,(
  spl0_35 | ~spl0_19 | spl0_34 | spl0_11),
  inference(avatar_split_clause,[],[f213,f263,f359,f299,f363])).
fof(f213,plain,(
  ( ! [X14,X15] : (hskp9 | ~c3_1(X14) | ~c0_1(X14) | c1_1(X14) | ~ndr1_0 | ~c2_1(X15) | ~c0_1(X15) | c1_1(X15)) )),
  inference(duplicate_literal_removal,[],[f164])).
fof(f164,plain,(
  ( ! [X14,X15] : (hskp9 | ~c3_1(X14) | ~c0_1(X14) | c1_1(X14) | ~ndr1_0 | ~c2_1(X15) | ~c0_1(X15) | c1_1(X15) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f368,plain,(
  spl0_35 | ~spl0_19 | spl0_36 | spl0_18),
  inference(avatar_split_clause,[],[f214,f294,f366,f299,f363])).
fof(f214,plain,(
  ( ! [X12,X13] : (hskp17 | ~c3_1(X12) | ~c2_1(X12) | ~c0_1(X12) | ~ndr1_0 | ~c2_1(X13) | ~c0_1(X13) | c1_1(X13)) )),
  inference(duplicate_literal_removal,[],[f165])).
fof(f165,plain,(
  ( ! [X12,X13] : (hskp17 | ~c3_1(X12) | ~c2_1(X12) | ~c0_1(X12) | ~ndr1_0 | ~c2_1(X13) | ~c0_1(X13) | c1_1(X13) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f361,plain,(
  ~spl0_19 | spl0_34 | spl0_31 | spl0_3),
  inference(avatar_split_clause,[],[f166,f227,f346,f359,f299])).
fof(f166,plain,(
  ( ! [X11] : (hskp12 | hskp3 | ~c3_1(X11) | ~c0_1(X11) | c1_1(X11) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f349,plain,(
  spl0_29 | ~spl0_19 | spl0_30 | spl0_31),
  inference(avatar_split_clause,[],[f216,f346,f343,f299,f340])).
fof(f216,plain,(
  ( ! [X8,X7] : (hskp3 | ~c3_1(X7) | ~c1_1(X7) | ~c0_1(X7) | ~ndr1_0 | ~c0_1(X8) | c3_1(X8) | c2_1(X8)) )),
  inference(duplicate_literal_removal,[],[f168])).
fof(f168,plain,(
  ( ! [X8,X7] : (hskp3 | ~c3_1(X7) | ~c1_1(X7) | ~c0_1(X7) | ~ndr1_0 | ~c0_1(X8) | c3_1(X8) | c2_1(X8) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f338,plain,(
  spl0_26 | ~spl0_19 | spl0_28 | spl0_9),
  inference(avatar_split_clause,[],[f217,f254,f336,f299,f328])).
fof(f217,plain,(
  ( ! [X6,X5] : (hskp21 | ~c3_1(X5) | ~c0_1(X5) | c2_1(X5) | ~ndr1_0 | ~c1_1(X6) | c3_1(X6) | c2_1(X6)) )),
  inference(duplicate_literal_removal,[],[f169])).
fof(f169,plain,(
  ( ! [X6,X5] : (hskp21 | ~c3_1(X5) | ~c0_1(X5) | c2_1(X5) | ~ndr1_0 | ~c1_1(X6) | c3_1(X6) | c2_1(X6) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f326,plain,(
  ~spl0_19 | spl0_23 | spl0_25 | spl0_1),
  inference(avatar_split_clause,[],[f171,f219,f323,f315,f299])).
fof(f171,plain,(
  ( ! [X3] : (hskp20 | hskp18 | ~c1_1(X3) | ~c0_1(X3) | c2_1(X3) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f321,plain,(
  ~spl0_19 | spl0_23 | spl0_5 | spl0_24),
  inference(avatar_split_clause,[],[f172,f318,f236,f315,f299])).
fof(f172,plain,(
  ( ! [X2] : (hskp2 | hskp15 | ~c1_1(X2) | ~c0_1(X2) | c2_1(X2) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f279,plain,(
  spl0_13 | spl0_7 | spl0_14),
  inference(avatar_split_clause,[],[f177,f276,f245,f272])).
fof(f177,plain,(
  hskp14 | hskp5 | hskp25),
  inference(cnf_transformation,[],[f6])).
fof(f270,plain,(
  spl0_11 | spl0_12 | spl0_1),
  inference(avatar_split_clause,[],[f178,f219,f267,f263])).
fof(f178,plain,(
  hskp20 | hskp23 | hskp9),
  inference(cnf_transformation,[],[f6])).
fof(f252,plain,(
  spl0_7 | spl0_8),
  inference(avatar_split_clause,[],[f180,f249,f245])).
fof(f180,plain,(
  hskp11 | hskp5),
  inference(cnf_transformation,[],[f6])).
fof(f230,plain,(
  spl0_1 | spl0_2 | spl0_3),
  inference(avatar_split_clause,[],[f182,f227,f223,f219])).
fof(f182,plain,(
  hskp12 | hskp6 | hskp20),
  inference(cnf_transformation,[],[f6])).
% SZS output end Proof for SYN448+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7291
% Time elapsed: 0.034 s
% ------------------------------
% ------------------------------
